<!DOCTYPE html>
<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Увод у процесирање низова података у Пајтону</title>
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" type="text/css" href="../_static/gallery.css" />
    <link rel="stylesheet" type="text/css" href="../_static/associations.css" />
    <link rel="stylesheet" type="text/css" href="../_static/editor.css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="https://fonts.googleapis.com/css?family=Source+Sans+Pro" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja-course.css" type="text/css" />
    <link rel="stylesheet" href="../_static/nbstyle.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/language_data.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/karelCorner.js"></script>
    <script type="text/javascript" src="../_static/karelRobot.js"></script>
    <script type="text/javascript" src="../_static/karelWorld.js"></script>
    <script type="text/javascript" src="../_static/karelRobotDrawer.js"></script>
    <script type="text/javascript" src="../_static/karelUI.js"></script>
    <script type="text/javascript" src="../_static/karel.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/notes.js"></script>
    <script type="text/javascript" src="../_static/pygamelib-init.js"></script>
    <script type="text/javascript" src="../_static/gallery.js"></script>
    <script type="text/javascript" src="../_static/associations.js"></script>
    <script type="text/javascript" src="../_static/associations-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/associations-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/editor.js"></script>
    <script type="text/javascript" src="../_static/jszip.js"></script>
    <script type="text/javascript" src="../_static/editor-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/editor-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/translations.js"></script>
    <script async="async" type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="../_static/course-errors.js"></script>
    <script type="text/javascript" src="../_static/petlja.js"></script>
    <script type="text/javascript" src="../_static/require.js"></script>
    <meta charset='utf-8'>
    <meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
    <meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
    <link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />
    <script type="text/javascript">
      eBookConfig = {};
      eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
      eBookConfig.app = eBookConfig.host + '/runestone';
      eBookConfig.ajaxURL = eBookConfig.app + '/ajax/';
      eBookConfig.course = '';
      eBookConfig.logLevel = '0';
      eBookConfig.loginRequired = 'false';
      eBookConfig.build_info = "";
      eBookConfig.isLoggedIn = false;
      eBookConfig.useRunestoneServices = false;
      eBookConfig.python3 = true;
      eBookConfig.basecourse = 'Настава програмирања за Други разред гимназије';
      eBookConfig.runestone_version = '';
      eBookConfig.imagesDir = '../_images/';
      eBookConfig.staticDir = '../_static/';
      if (typeof (Sk) != "undefined")
        Sk.imgPath = eBookConfig.imagesDir;
    </script>

  </head>

  <body>



<div id="navbar" class="content-header">
  <div class="right-side-nav">
    
    <div id="usefulLinksToggle" class="useful-links-toggle">
        <a class="dropdown-toggle dropdown-toggle-burger" role="button" href="#" id="usefulLinksDropdown" data-toggle="dropdown" aria-haspopup="true" aria-expanded="false">
            <i class="fas fa-bars fa-2x"></i>
        </a>
        <div class="dropdown-menu dropdown-nav-links shadow-lg" aria-labelledby="usefulLinksDropdown" x-placement="bottom-start" style="position: absolute; transform: translate3d(-55px, 32px, 0px); top: 0px; left: 0px; will-change: transform;">
          <div class="links-title"> Корисни линкови</div>
          <div class="useful-link">
            <a href="https://www.sphinx-doc.org/en/master/usage/restructuredtext/index.html" target="_blank">Sphinx Docs</a></div>
          <div class="useful-link">
            <a href="https://github.com/Petlja/PetljaDoc/blob/master/README.md" target="_blank">Petljadoc</a></div>
          <div class="useful-link">
            <a href="https://github.com/RunestoneInteractive/overview" target="_blank">Runestone Interactive</a></div>
          
        </div>
    </div>

    
</div>

<div class="left-side-nav">
  <div class="course-logo-nav">
      <a target="_blank" href="https://petlja.org">
          <img src="../_static/img/petlja-logo-wt.svg" alt="petlja.org" id="lightPetljaLogo">
      </a>
  </div>
  <div class="course-logo-nav-kabinet">
      <a target="_blank" href="https://petlja.org/net.kabinet">
          <img src="../_static/img/netkabinet-logo-strelica-wt.svg" alt="petlja.org" id="lightNetKabinetLogo">
      </a>
  </div>
</div>


<div class="content-header-label" id = "content-header-label" href="../">Анализа података у Џупитеру и технике програмирања за други разред гимназије</div>

</div>




<div class="petljaCourseContainer">

<div class="petljaCourseContent">
  
           <div class="tableOfContent">
            <h5> Садржај </h5>
            </div>
           <div>
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J0A>
                    <h5>Пре него што почнемо</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J0A/howto.html><div class="studioLecture" id=activity-howto><i class="fas fa-file-alt activity-icon"></i>Како да користиш овај приручник</div></a>
                <a href=../J0A/J0A.html><div class="studioLecture" id=activity-J0A><i class="fas fa-file-alt activity-icon"></i>Покретање Џупитер радних свески</div></a>
                <a href=../J0A/J01.html><div class="studioLecture" id=activity-J01><i class="fas fa-file-alt activity-icon"></i>Увод у Џупитер</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-podsetnik>
                    <h5>Подсетник Пајтон</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../podsetnik/Podsetnik%20-%20Pajton.html><div class="studioLecture" id=activity-Podsetnik%20-%20Pajton><i class="fas fa-file-alt activity-icon"></i>Подсетник Пајтон</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y01>
                    <h5>Низови података</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y01/L801.html><div class="studioLecture" id=activity-L801><i class="fas fa-file-alt activity-icon"></i>Увод у процесирање низова података</div></a>
                <a href=../Y01/Kreiranje%20lista%20u%20Pajtonu%20-%20List%20comprehension.html><div class="studioLecture" id=activity-Kreiranje%20lista%20u%20Pajtonu%20-%20List%20comprehension><i class="fas fa-file-alt activity-icon"></i>Креирање листа у Пајтону</div></a>
                <a href=../Y01/Низови%20-%20додатни%20задаци.html><div class="studioLecture" id=activity-Низови%20-%20додатни%20задаци><i class="far fa-question-circle activity-icon"></i>Низови - додатни задаци</div></a>
                <a href=../Y01/test801.html><div class="studioLecture" id=activity-test801><i class="far fa-question-circle activity-icon"></i>Квиз - Листе</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y05>
                    <h5>Сортирање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y05/L805.html><div class="studioLecture" id=activity-L805><i class="fas fa-file-alt activity-icon"></i>Сортирање</div></a>
                <a href=../Y05/Сортирање%20-%20додатни%20задаци.html><div class="studioLecture" id=activity-Сортирање%20-%20додатни%20задаци><i class="far fa-question-circle activity-icon"></i>Сортирање - додатни задаци</div></a>
                <a href=../Y05/test805.html><div class="studioLecture" id=activity-test805><i class="far fa-question-circle activity-icon"></i>Квиз - Сортирање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y06>
                    <h5>Филтрирање и претраживање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y06/L806.html><div class="studioLecture" id=activity-L806><i class="fas fa-file-alt activity-icon"></i>Филтрирање и претраживање</div></a>
                <a href=../Y06/test806.html><div class="studioLecture" id=activity-test806><i class="far fa-question-circle activity-icon"></i>Квиз - Филтрирање и претраживање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-predstavljanje>
                    <h5>Представљање низова података</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../predstavljanje/Predstavljanje%20nizova%20podataka.html><div class="studioLecture" id=activity-Predstavljanje%20nizova%20podataka><i class="fas fa-file-alt activity-icon"></i>Представљање низова података</div></a>
                <a href=../predstavljanje/Graficko_predstavljanje.html><div class="studioLecture" id=activity-Graficko_predstavljanje><i class="fas fa-file-alt activity-icon"></i>Графичко представљање низова података</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-tekstualni>
                    <h5>Рад са текстуалним подацима</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../tekstualni/Rad%20sa%20tekstualnim%20podacima.html><div class="studioLecture" id=activity-Rad%20sa%20tekstualnim%20podacima><i class="fas fa-file-alt activity-icon"></i>Рад са текстуалним подацима</div></a>
                <a href=../tekstualni/tekstualne_datoteke.html><div class="studioLecture" id=activity-tekstualne_datoteke><i class="fas fa-file-alt activity-icon"></i>Коришћење текстуалних датотека</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-rad%20sa%20tabelama>
                    <h5>Рад са табелама</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../rad%20sa%20tabelama/Rad%20sa%20tabelama.html><div class="studioLecture" id=activity-Rad%20sa%20tabelama><i class="fas fa-file-alt activity-icon"></i>Рад са табелама</div></a>
                <a href=../rad%20sa%20tabelama/indeksiranje.html><div class="studioLecture" id=activity-indeksiranje><i class="fas fa-file-alt activity-icon"></i>Индексирање табела</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-obrada_tabele>
                    <h5>Обрада и приказ табеларних података</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../obrada_tabele/obrada_prikaz.html><div class="studioLecture" id=activity-obrada_prikaz><i class="fas fa-file-alt activity-icon"></i>Обрада и приказ табеларних података</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J10>
                    <h5>Џупитер и Ексел</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J10/J10.html><div class="studioLecture" id=activity-J10><i class="fas fa-file-alt activity-icon"></i>Џупитер и Ексел</div></a>
                <a href=../J10/testJ10.html><div class="studioLecture" id=activity-testJ10><i class="far fa-question-circle activity-icon"></i>Џупитер и Ексел</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-projektni%20zadatak>
                    <h5>Пример скупа података за пројекто учење</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../projektni%20zadatak/projekat.html><div class="studioLecture" id=activity-projekat><i class="fas fa-file-alt activity-icon"></i>Пример скупа података за пројекто учење</div></a>
                </div>
            </div>
                </div>
</div>



<div class="lectureContent">

  <div id="errors">
  </div>
  
  <div class="lectureContentMaterial">
  
    <div id="main-content">
    
  <div class="section" id="id1">
<h1>Увод у процесирање низова података у Пајтону<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>У овој лекцији ћемо:</p>
<ol class="arabic simple">
<li><p>поновити како се записују низови података,</p></li>
<li><p>поновити како се дефинишу функције,</p></li>
<li><p>написаћемо своју функцију која рачуна најмањи број у датом низу бројева,</p></li>
<li><p>поновити неколико функција које су уграђене у Пајтон, и</p></li>
<li><p>поновити како се листа може учитати од корисника (са стандардног улаза).</p></li>
</ol>
<div class="section" id="id2">
<h2>Записивање низова података<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>
<p>Да се подсетимо, низови података се у Пајтону записују унутар угластих заграда,
у облику <em>листе</em>. На пример, оцене неког ученика из једног предмета могу да се запишу као низ овако:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">Dejanove_ocene</span> <span class="o">=</span> <span class="p">[</span><span class="mi">2</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">3</span><span class="p">,</span> <span class="mi">5</span><span class="p">]</span>
</pre></div>
</div>
<p>Ово је пример нумеричког низа података, где су сви подаци бројеви. Низ података, наравно, не
мора бити нумерички. На пример:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">predmeti</span> <span class="o">=</span> <span class="p">[</span><span class="s2">&quot;matematika&quot;</span><span class="p">,</span> <span class="s2">&quot;srpski&quot;</span><span class="p">,</span> <span class="s2">&quot;likovno&quot;</span><span class="p">,</span> <span class="s2">&quot;istorija&quot;</span><span class="p">,</span> <span class="s2">&quot;fizičko&quot;</span><span class="p">]</span>
</pre></div>
</div>
<p>или овако:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">Ocene</span> <span class="o">=</span> <span class="p">[</span><span class="s2">&quot;Dejan&quot;</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">3</span><span class="p">]</span>
</pre></div>
</div>
<p>Раније смо видели да подаци у низу могу бити и други низови података:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">Ocene_VIII3</span> <span class="o">=</span> <span class="p">[</span>
  <span class="p">[</span><span class="s2">&quot;Dejan&quot;</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">3</span><span class="p">],</span>
  <span class="p">[</span><span class="s2">&quot;Milica&quot;</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">],</span>
  <span class="p">[</span><span class="s2">&quot;Milan&quot;</span><span class="p">,</span> <span class="mi">3</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">5</span><span class="p">]</span>
<span class="p">]</span>
</pre></div>
</div>
<p>О низу података можемо да размишљамо као о улици у којој свака кућица има свој кућни број. Тај број називамо индексом.
Специфичност Пајтона је у томе да кућни бројеви, тј. индекси, увек почињу од 0, рецимо овако:</p>
<a class="reference internal image-reference" href="../_images/Y101.jpg"><img alt="../_images/Y101.jpg" src="../_images/Y101.jpg" style="width: 600px;" /></a>
<p>Број елемената низа се добија позивом уграђене функције <code class="docutils literal notranslate"><span class="pre">len</span></code>
(од енглеског <em>length</em> што значи „дужина”):</p>

<div data-childcomponent="primer1-5" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer1-5 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
predmeti = ["matematika", "srpski", "likovno", "istorija", "fizičko"]
print(len(predmeti))
</textarea>
</div>
<p>Ако желимо да приступимо појединачној кућици довољно је да наведемо њену „адресу” која се,
као и поштанска адреса, састоји од имена улице иза кога се наводи кућни број, рецимо овако:</p>

<div data-childcomponent="primer1-6" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer1-6 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
Dejanove_ocene = [2, 4, 5, 3, 5]
print(Dejanove_ocene[0])
</textarea>
</div>
<p>Вредности бројева у кућицама се могу мењати веома једноставно.
Рецимо, ако Дејан је поправио оцену која је уписана у прву кућицу
(<em>чији индекс је 0</em>) то можемо регистровати овако:</p>

<div data-childcomponent="primer1-7" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer1-7 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
Dejanove_ocene = [2, 4, 5, 3, 5]
Dejanove_ocene[0] = 5
print(Dejanove_ocene[0])
</textarea>
</div>
<p>Стање у улици <code class="docutils literal notranslate"><span class="pre">Dejanove_ocene</span></code> сада изгледа овако:</p>
<a class="reference internal image-reference" href="../_images/Y102.jpg"><img alt="../_images/Y102.jpg" src="../_images/Y102.jpg" style="width: 600px;" /></a>
<p>Ако узмемо да низ <code class="docutils literal notranslate"><span class="pre">predmeti</span></code> садржи имена предмета из којих је Дејан оцењен онда извештај о Дејановим оценама можемо
исписати на следећи начин:</p>

<div data-childcomponent="primer1-8" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer1-8 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="predmeti = [&quot;matematika&quot;, &quot;srpski&quot;, &quot;likovno&quot;, &quot;istorija&quot;, &quot;fizičko&quot;]
Dejanove_ocene = [2, 4, 5, 3, 5]
print(&quot;Dejanove ocene:&quot;)
for i in range(5):
   print(predmeti[i], Dejanove_ocene[i])
"   >


</textarea>
</div>

    <div class="note-wrapper infonote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/infonote-img.svg" class="note-image infonote-image" /> 
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P01</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Dejanove_ocene.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
<div class="section" id="id3">
<h2>Стандардне функције<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h2>
<p>Како би се писање програма поједноставило за велики број важних задатака који се веома
често користе у Пајтону већ постоје функције које их обављају. На пример, функција
<code class="docutils literal notranslate"><span class="pre">len</span></code> рачуна дужину низа, као што смо већ видели.
Одређивање најмањег и највећег елемента низа бројева спадају међу важне задатке,
и зато Пајтон има уграђене функције <code class="docutils literal notranslate"><span class="pre">min</span></code> и <code class="docutils literal notranslate"><span class="pre">max</span></code> које могу да одреде најмањи,
односно, највећи елемент неколико бројева који су унапред наведени, или целог низа бројева:</p>

<div data-childcomponent="primer1-9" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer1-9 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print(min[2,3])
</textarea>
</div>

<div data-childcomponent="primer1-10" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer1-10 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print(max[4,5,-2])
</textarea>
</div>

<div data-childcomponent="primer1-11" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer1-11 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
Dejanove_ocene = [2, 4, 5, 3, 5]
print("Najmanja Dejanova ocena je", min(Dejanove_ocene))
print("Najveca Dejanova ocena je", max(Dejanove_ocene))
</textarea>
</div>
<div class="section" id="id4">
<h3>Пример 1.<a class="headerlink" href="#id4" title="Permalink to this headline">¶</a></h3>
<p>Променљива <code class="docutils literal notranslate"><span class="pre">Zikine_ocene</span></code> садржи Жикине оцене из информатике. Она је сакривена негде на овој страни и ти је не видиш.
У следећем малом програму поправи програмски код тако да добијеш највећу и најмању Жикину оцену из информатике,
као и колико оцена је Жика добио из информатике.</p>

<div data-childcomponent="primer1-12" class="course-box course-box-problem course-content">
<p>У три наредбе испод уместо 0 упиши одговарајуће ИЗРАЗЕ који ће
на основу низа оцена у променљивој <code class="docutils literal notranslate"><span class="pre">Zikine_ocene</span></code> (који је сакривен од тебе!)
ИЗРАЧУНАТИ најмању и највећу Жикину оцену из информатике, као и број оцена које
је Жика добио из информатике</p>

<textarea data-component="activecode" id=primer1-12 data-lang="python" 
     data-include="Zikine_ocene" data-timelimit=25000    
    data-audio=''   
          data-runortest="najmanja najveca broj Zikine_ocene"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
najmanja = 0 # поправи овај ред
najveca = 0 # поправи овај ред
broj = 0 # поправи овај ред
# -*- acsection: after-main -*-
print("Žikina najmanja ocena iz informatike je", najmanja)
print("Žikina najveća ocena iz informatike je", najveca)
print("Žika je dobio", broj, "ocena iz informatike")
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(acMainSection(Zikine_ocene=Zikine_ocene)["najmanja"], min(Zikine_ocene), "Вредност променљиве 'najmanja' треба да буде %s" % min(Zikine_ocene))
        self.assertEqual(acMainSection(Zikine_ocene=Zikine_ocene)["najveca"], max(Zikine_ocene), "Вредност променљиве 'najveca' треба да буде %s" % max(Zikine_ocene))
        self.assertEqual(acMainSection(Zikine_ocene=Zikine_ocene)["broj"], len(Zikine_ocene), "Вредност променљиве 'broj' треба да буде %s" % len(Zikine_ocene))
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id5">
<h3>Пример 2.<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h3>
<p>Функција <code class="docutils literal notranslate"><span class="pre">sum</span></code> рачуна збир елемената низа, рецимо овако:</p>

<div data-childcomponent="primer1-13" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer1-13 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
brojevi = [3,2,1,4,6]
print(sum(brojevi))
</textarea>
</div>
<p>Средња оцена се рачуна као</p>
<div class="math notranslate nohighlight">
\[\frac{\mathrm{збир\ свих\ оцена}}{\mathrm{укупан\ број\ оцена}}\]</div>
<p>У следећем малом програму поправи означени ред тако да добијеш израз који рачуна средњу вредност
Жикиних оцена из информатике. Као и раније, Жикине оцене из информатике су смештене у променљиву
<code class="docutils literal notranslate"><span class="pre">Zikine_ocene</span></code> која је сакривена од тебе, али је програм види.</p>

<div data-childcomponent="primer1-14" class="course-box course-box-problem course-content">
<p>У наредби испод уместо 0 упиши ИЗРАЗ који ће на основу низа оцена
у променљивој <code class="docutils literal notranslate"><span class="pre">Zikine_ocene</span></code> (који је сакривен од тебе!)
ИЗРАЧУНАТИ Жикину средњу оцену из информатике. Немој заокруживати ову вредност!</p>

<textarea data-component="activecode" id=primer1-14 data-lang="python" 
     data-include="Zikine_ocene" data-timelimit=25000    
    data-audio=''   
          data-runortest="srednja_vrednost Zikine_ocene"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
srednja_vrednost = 0 # поправи овај ред
# -*- acsection: after-main -*-
print("Prosek Žikinih ocena iz informatike je", srednja_vrednost)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        x = sum(Zikine_ocene)/len(Zikine_ocene)
        self.assertEqual(acMainSection(Zikine_ocene=Zikine_ocene)["srednja_vrednost"], x, "Вредност променљиве 'srednja_vrednost' треба да буде %s" % x)
myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="id6">
<h2>Дефинисање функција и минимум низа<a class="headerlink" href="#id6" title="Permalink to this headline">¶</a></h2>
<p>Када неки компликован процес треба поновити више пута, а не постоји већ дефинисана функција
која то може да уради за нас, постоји могућност да дефинишемо своју функцију.
Функција је мала целина која мора бити именована како бисмо касније могли да је
брзо и лако позовемо кад год нам затреба. На пример, ево функције која исписује извештај о оценама неког
ученика:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">izvestaj_o_ocenama</span><span class="p">(</span><span class="n">ime_ucenika</span><span class="p">,</span> <span class="n">spisak_predmeta</span><span class="p">,</span> <span class="n">ocene</span><span class="p">):</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;Učenik:&quot;</span><span class="p">,</span> <span class="n">ime_ucenika</span><span class="p">)</span>
    <span class="n">n</span> <span class="o">=</span> <span class="nb">len</span><span class="p">(</span><span class="n">spisak_predmeta</span><span class="p">)</span>
    <span class="k">for</span> <span class="n">i</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="n">n</span><span class="p">):</span>
       <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;-&quot;</span><span class="p">,</span> <span class="n">spisak_predmeta</span><span class="p">[</span><span class="n">i</span><span class="p">],</span> <span class="n">ocene</span><span class="p">[</span><span class="n">i</span><span class="p">])</span>
</pre></div>
</div>
<p>Извештај о Дејановим оценама сада лако можемо добити позивом ове функције:</p>

<div data-childcomponent="primer1-15" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer1-15 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="# Прво дефинишемо функцију
def izvestaj_o_ocenama(ime_ucenika, spisak_predmeta, ocene):
    print(&quot;Učenik:&quot;, ime_ucenika)
    n = len(spisak_predmeta)
    for i in range(n):
       print(&quot;-&quot;, spisak_predmeta[i], ocene[i])

# Сада је можемо позвати више пута ако нам треба извештај о оценама
# за неколико ученика

predmeti = [&quot;matematika&quot;, &quot;srpski&quot;, &quot;likovno&quot;, &quot;istorija&quot;, &quot;fizičko&quot;]
Dejanove_ocene = [2, 4, 5, 3, 5]
Milenine_ocene = [5, 5, 4, 5, 4]

izvestaj_o_ocenama(&quot;Dejan&quot;, predmeti, Dejanove_ocene)
izvestaj_o_ocenama(&quot;Milena&quot;, predmeti, Milenine_ocene)
"   >


</textarea>
</div>

    <div class="note-wrapper infonote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/infonote-img.svg" class="note-image infonote-image" /> 
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P01</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Ocene_izvestaj.py</span></code> и изврши га.</p>

    </p></div></div>
<p>Функције се најчешће користе за аутоматизацију релативно компликованих рачунских процеса.</p>
<div class="section" id="id7">
<h3>Пример 3.<a class="headerlink" href="#id7" title="Permalink to this headline">¶</a></h3>
<p>Написати функцију <code class="docutils literal notranslate"><span class="pre">min_niza(A)</span></code> која одређује најмањи елемент низа <code class="docutils literal notranslate"><span class="pre">A</span></code>
и враћа га као резултат свог рада, али без коришћења уграђене функције <code class="docutils literal notranslate"><span class="pre">min</span></code>.</p>
<p><em>Решење.</em> Погледајмо, прво, како бисмо овај проблем решили када тражимо мањи од два броја, А0 и А1:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">manji_od_dva</span><span class="p">(</span><span class="n">A0</span><span class="p">,</span> <span class="n">A1</span><span class="p">):</span>
    <span class="k">if</span> <span class="n">A0</span> <span class="o">&lt;</span> <span class="n">A1</span><span class="p">:</span> <span class="k">return</span> <span class="n">A0</span>
    <span class="k">else</span><span class="p">:</span> <span class="k">return</span> <span class="n">A1</span>
</pre></div>
</div>
<p>Дакле, ако је А0 мањи, то је резултат рада. У супротном је А1 резултат рада.</p>
<p>Ова идеја ради у случају два броја, али није погодна за уопштења.
Зато ћемо написати нову верзију функције <code class="docutils literal notranslate"><span class="pre">manji_od_dva</span></code> која на први поглед делује чудно,
али ће ускоро постати јасно зашто је погодна.</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">manji_od_dva_ponovo</span><span class="p">(</span><span class="n">A0</span><span class="p">,</span> <span class="n">A1</span><span class="p">):</span>
    <span class="n">manji</span> <span class="o">=</span> <span class="n">A0</span>
    <span class="k">if</span> <span class="n">A1</span> <span class="o">&lt;</span> <span class="n">manji</span><span class="p">:</span> <span class="n">manji</span> <span class="o">=</span> <span class="n">A1</span>
    <span class="k">return</span> <span class="n">manji</span>
</pre></div>
</div>
<p>Идеја се састоји у томе да гледамо елементе један по један, као да нам долазе на траци. Када стигне први елемент, то је најмањи елемент кога смо до сада уочили. Када стигне други елемент проверимо како се он односи према до сада најмањем уоченом елементу и ако је мањи од најмањег до сада, он постаје нови најмањи елемент.</p>
<p>Ево како се идеја коју смо управо описали проширује на три елемента:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">najmanji_od_tri</span><span class="p">(</span><span class="n">A0</span><span class="p">,</span> <span class="n">A1</span><span class="p">,</span> <span class="n">A2</span><span class="p">):</span>
    <span class="n">najmanji</span> <span class="o">=</span> <span class="n">A0</span>
    <span class="k">if</span> <span class="n">A1</span> <span class="o">&lt;</span> <span class="n">najmanji</span><span class="p">:</span> <span class="n">najmanji</span> <span class="o">=</span> <span class="n">A1</span>
    <span class="k">if</span> <span class="n">A2</span> <span class="o">&lt;</span> <span class="n">najmanji</span><span class="p">:</span> <span class="n">najmanji</span> <span class="o">=</span> <span class="n">A2</span>
    <span class="k">return</span> <span class="n">najmanji</span>
</pre></div>
</div>
<p>и на ситуацију са пет елемената:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">najmanji_od_pet</span><span class="p">(</span><span class="n">A0</span><span class="p">,</span> <span class="n">A1</span><span class="p">,</span> <span class="n">A2</span><span class="p">,</span> <span class="n">A3</span><span class="p">,</span> <span class="n">A4</span><span class="p">):</span>
    <span class="n">najmanji</span> <span class="o">=</span> <span class="n">A0</span>
    <span class="k">if</span> <span class="n">A1</span> <span class="o">&lt;</span> <span class="n">najmanji</span><span class="p">:</span> <span class="n">najmanji</span> <span class="o">=</span> <span class="n">A1</span>
    <span class="k">if</span> <span class="n">A2</span> <span class="o">&lt;</span> <span class="n">najmanji</span><span class="p">:</span> <span class="n">najmanji</span> <span class="o">=</span> <span class="n">A2</span>
    <span class="k">if</span> <span class="n">A3</span> <span class="o">&lt;</span> <span class="n">najmanji</span><span class="p">:</span> <span class="n">najmanji</span> <span class="o">=</span> <span class="n">A3</span>
    <span class="k">if</span> <span class="n">A4</span> <span class="o">&lt;</span> <span class="n">najmanji</span><span class="p">:</span> <span class="n">najmanji</span> <span class="o">=</span> <span class="n">A4</span>
    <span class="k">return</span> <span class="n">najmanji</span>
</pre></div>
</div>
<p>Видимо да се корак који садржи <code class="docutils literal notranslate"><span class="pre">if</span></code> понавља за сваки елемент посебно.
Тако се ова идеја лако уопштава на произвољне низове:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">min_niza</span><span class="p">(</span><span class="n">A</span><span class="p">):</span>
    <span class="n">najmanji</span> <span class="o">=</span> <span class="n">A</span><span class="p">[</span><span class="mi">0</span><span class="p">]</span>
    <span class="k">for</span> <span class="n">i</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="nb">len</span><span class="p">(</span><span class="n">A</span><span class="p">)):</span>
        <span class="k">if</span> <span class="n">A</span><span class="p">[</span><span class="n">i</span><span class="p">]</span> <span class="o">&lt;</span> <span class="n">najmanji</span><span class="p">:</span> <span class="n">najmanji</span> <span class="o">=</span> <span class="n">A</span><span class="p">[</span><span class="n">i</span><span class="p">]</span>
    <span class="k">return</span> <span class="n">najmanji</span>
</pre></div>
</div>
<p>Да проверимо како ова функција ради:</p>

<div data-childcomponent="primer1-16" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer1-16 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="def min_niza(A):
    najmanji = A[0]
    for i in range(len(A)):
        if A[i] &lt; najmanji: najmanji = A[i]
    return najmanji

print(min_niza([3,1,5,-1,2]))
"   >


</textarea>
</div>

    <div class="note-wrapper infonote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/infonote-img.svg" class="note-image infonote-image" /> 
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P01</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Min_niza.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
</div>
<div class="section" id="id8">
<h2>Учитавање низова<a class="headerlink" href="#id8" title="Permalink to this headline">¶</a></h2>
<p>Постоје ситуације када је потребно од корисника учитати низ података и обрадити га. У следећем примеру
ћемо се подсетити како се низови бројева учитавају од корисника, односно уносе помоћу тастатуре.</p>
<div class="section" id="id9">
<h3>Пример.<a class="headerlink" href="#id9" title="Permalink to this headline">¶</a></h3>
<p>Написати програм који од корисника учитава оцене из неколико предмета, а онда исписује из којих предмета има најлошију оцену, а из којих има
најбољу оцену.</p>

<div data-childcomponent="Primer_1_4-NOVO" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=Primer_1_4-NOVO data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
             data-passivecode="true"  data-hidehistory=true
    data-includesrc="# Učitavanje podataka
n = int(input(&quot;Koliko ima predmeta? &quot;))
print(&quot;Za&quot;, n, &quot;predmeta unesi imena i ocene&quot;)
predmeti = []
ocene = []
for i in range(n):
  ime = input(&quot;Ime &quot;+ str(i+1) + &quot;. predmeta: &quot;)
  ocena = int(input(&quot;Ocena: &quot;))
  predmeti.append(ime)
  ocene.append(ocena)

# Računanje minimuma i maksimuma
min_ocena = min(ocene)
max_ocena = max(ocene)

# Ispis predmeta iz kojih je učenih najbolje, odnosno, najlošije ocenjen
if min_ocena == max_ocena:
  print(&quot;Ti si primer konstantnog kvaliteta!&quot;)
else:
  print(&quot;Tvoji najbolje ocenjeni predmeti su:&quot;)
  for i in range(n):
    if(ocene[i] == max_ocena):
      print(&quot;-&quot;, predmeti[i])
  print(&quot;Tvoji najlošije ocenjeni predmeti su:&quot;)
  for i in range(n):
    if(ocene[i] == min_ocena):
      print(&quot;-&quot;, predmeti[i])
"   >


</textarea>
</div>
<p>У овом програму имамо два низа: низ у коме памтимо имена предмета и низ у коме памтимо оцене. Оба низа су на почетку празни, а током рада
for-циклуса се на крај сваког од њих дода по један податак: име новог предмета и оцена из тог предмета. Да се подсетимо, наредба</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">L</span><span class="o">.</span><span class="n">append</span><span class="p">(</span><span class="n">x</span><span class="p">)</span>
</pre></div>
</div>
<p>додаје податак <code class="docutils literal notranslate"><span class="pre">x</span></code> на крај листе <code class="docutils literal notranslate"><span class="pre">L</span></code>.</p>
<p>Након тога позивом стандардних функција <code class="docutils literal notranslate"><span class="pre">min</span></code> и <code class="docutils literal notranslate"><span class="pre">mix</span></code> одредимо најлошију и најбољу оцену коју је ученик добио, па у последњем делу програма
(у „трећој фази”) испишемо имена предмета из којих је ученик најбоље оцењен и имена предмета из којих је ученик најлошије оцењен.</p>

    <div class="note-wrapper infonote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/infonote-img.svg" class="note-image infonote-image" /> 
        <div class="course-content">
            <p>
<p>Да би видео како програм ради покрени IDLE окружење из фолдера <code class="docutils literal notranslate"><span class="pre">P01</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Naj_predmeti.py</span></code> и изврши га.</p>
<p>Ево и кратке видео демонстрације:</p>

    <div  id="GnsDOougiPA" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/GnsDOougiPA/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-GnsDOougiPA" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-GnsDOougiPA">
        </div>
    </div>
    
    
    
    </p></div></div>
</div>
</div>
<div class="section" id="id10">
<h2>Задаци<a class="headerlink" href="#id10" title="Permalink to this headline">¶</a></h2>
<div class="section" id="id11">
<h3>Задатак 1.<a class="headerlink" href="#id11" title="Permalink to this headline">¶</a></h3>
<p>Напиши функцију <code class="docutils literal notranslate"><span class="pre">najveci_od_tri(A0,</span> <span class="pre">A1,</span> <span class="pre">A2)</span></code> која рачуна  и враћа највећи од три дата броја, али без употребе стандардне функције <code class="docutils literal notranslate"><span class="pre">max</span></code>.
Напиши тело функције, па онда провери како функција ради.</p>

<div data-childcomponent="zadatak1-1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak1-1 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def najveci_od_tri(A0, A1, A2):
    # Овде напиши функцију која рачуна највећи од бројева A0, A1, A2
    return -1234  # поправи овај ред!

# Провера
test1 = najveci_od_tri(5, 7, 11)
test2 = najveci_od_tri(5, 75, 11)
test3 = najveci_od_tri(55, 7, 11)
# -*- acsection: after-main -*-
print(test1, test2, test3)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        rez1 = max([5, 7,  11])
        rez2 = max([5, 75, 11])
        rez3 = max([55, 7, 11])
        run_test = acMainSection(test1=test1,test2=test2,test3=test3)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id12">
<h3>Задатак 2.<a class="headerlink" href="#id12" title="Permalink to this headline">¶</a></h3>
<p>Напиши функцију <code class="docutils literal notranslate"><span class="pre">najveci_od_pet(A0,</span> <span class="pre">A1,</span> <span class="pre">A2,</span> <span class="pre">A3,</span> <span class="pre">A4)</span></code>
која рачуна и враћа највећи од пет датих бројева, али без употребе стандардне функције <code class="docutils literal notranslate"><span class="pre">max</span></code>.
Напиши тело функције, па онда провери како функција ради.</p>

<div data-childcomponent="zadatak1-2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak1-2 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3 test4 test5"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def najveci_od_pet(A0, A1, A2, A3, A4):
    # Овде напиши функцију која рачуна највећи од бројева A0, A1, A2, A3, A4
    return -1234  # поправи овај ред!

# Провера
test1 = najveci_od_pet(1, 2, 3, 4, 55)
test2 = najveci_od_pet(1, 2, 3, 44, 5)
test3 = najveci_od_pet(1, 2, 33, 4, 5)
test4 = najveci_od_pet(1, 22, 3, 4, 5)
test5 = najveci_od_pet(11, 2, 3, 4, 5)
# -*- acsection: after-main -*-
print(test1, test2, test3, test4, test5)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        rez1 = max(1, 2, 3, 4, 55)
        rez2 = max(1, 2, 3, 44, 5)
        rez3 = max(1, 2, 33, 4, 5)
        rez4 = max(1, 22, 3, 4, 5)
        rez5 = max(11, 2, 3, 4, 5)
        run_test=acMainSection(test1=test1,test2=test2,test3=test3,test4=test4,test5=test5)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
        self.assertEqual(run_test["test4"], rez4, "Вредност променљиве 'test4' треба да буде %s" % rez4)
        self.assertEqual(run_test["test5"], rez5, "Вредност променљиве 'test5' треба да буде %s" % rez5)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id13">
<h3>Задатак 3.<a class="headerlink" href="#id13" title="Permalink to this headline">¶</a></h3>
<p>Напиши функцију <code class="docutils literal notranslate"><span class="pre">max_niza(A)</span></code> која за дати низ <code class="docutils literal notranslate"><span class="pre">А</span></code>
рачуна и враћа највећи елемент, али без употребе стандардне функције <code class="docutils literal notranslate"><span class="pre">max</span></code>.
Напиши тело функције, па онда провери како функција ради.</p>

<div data-childcomponent="zadatak1-3" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak1-3 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3 test4 test5"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def max_niza(A):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = max_niza([1, 2, 3, 4, 55, 6])
test2 = max_niza([1, 2, 3, 4, 55, 777])
test3 = max_niza([111, 2, 3, 4])
test4 = max_niza([1, 222])
test5 = max_niza([22])
# -*- acsection: after-main -*-
print(test1, test2, test3, test4, test5)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        rez1 = max([1, 2, 3, 4, 55, 6])
        rez2 = max([1, 2, 3, 4, 55, 777])
        rez3 = max([111, 2, 3, 4])
        rez4 = max([1, 222])
        rez5 = max([22])
        run_test=acMainSection(test1=test1,test2=test2,test3=test3,test4=test4,test5=test5)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
        self.assertEqual(run_test["test4"], rez4, "Вредност променљиве 'test4' треба да буде %s" % rez4)
        self.assertEqual(run_test["test5"], rez5, "Вредност променљиве 'test5' треба да буде %s" % rez5)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id14">
<h3>Задатак 4.<a class="headerlink" href="#id14" title="Permalink to this headline">¶</a></h3>
<p>На контролној провери из информатике у једној школи се ради шест задатака од којих сваки носи по 20 поена, али наставник даје оцену
на основу пет најбоље урађених задатака. Напиши функцију <code class="docutils literal notranslate"><span class="pre">zbir_najboljih_5(Z1,</span> <span class="pre">Z2,</span> <span class="pre">Z3,</span> <span class="pre">Z4,</span> <span class="pre">Z5,</span> <span class="pre">Z6)</span></code> која
за списак поена по задацима које је освојио неки ученик на тој контролној провери враћа збир поена пет најбоље урађених задатака.
Напиши тело функције, па онда провери како функција ради.
(Упутство: Задатак можеш решити употребом неколико if-ова, али и употребом стандардних функција <code class="docutils literal notranslate"><span class="pre">sum</span></code> и <code class="docutils literal notranslate"><span class="pre">min</span></code>)</p>

<div data-childcomponent="zadatak1-4" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak1-4 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def zbir_najboljih_5(Z1, Z2, Z3, Z4, Z5, Z6):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = zbir_najboljih_5(20, 20, 20, 20, 20, 20)
test2 = zbir_najboljih_5(10,  0, 15, 20,  0, 20)
test3 = zbir_najboljih_5( 5,  7, 10, 15, 17, 20)
# -*- acsection: after-main -*-
print(test1, test2, test3)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        L = [20, 20, 20, 20, 20, 20]; rez1 = sum(L) - min(L)
        L = [10,  0, 15, 20,  0, 20]; rez2 = sum(L) - min(L)
        L = [ 5,  7, 10, 15, 17, 20]; rez3 = sum(L) - min(L)
        run_test = acMainSection(test1=test1,test2=test2,test3=test3)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id15">
<h3>Задатак 5.<a class="headerlink" href="#id15" title="Permalink to this headline">¶</a></h3>
<p>Напиши функцију <code class="docutils literal notranslate"><span class="pre">srednji_od_tri(A0,</span> <span class="pre">A1,</span> <span class="pre">A2)</span></code> која рачуна и враћа средњи од три дата броја.
Напиши тело функције, па онда провери како функција ради.
(Упутство: Задатак можеш решити употребом неколико if-ова, али и употребом функција <code class="docutils literal notranslate"><span class="pre">sum</span></code>, <code class="docutils literal notranslate"><span class="pre">min</span></code> и <code class="docutils literal notranslate"><span class="pre">max</span></code>!)</p>

<div data-childcomponent="zadatak1-5" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak1-5 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3 test4"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def srednji_od_tri(A0, A1, A2):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = srednji_od_tri(2, 7, 5)
test2 = srednji_od_tri(2, 5, 5)
test3 = srednji_od_tri(5, 5, 1)
test4 = srednji_od_tri(2, 2, 2)
# -*- acsection: after-main -*-
print(test1, test2, test3, test4)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        L = [2, 7, 5]; rez1 = sum(L) - min(L) - max(L)
        L = [2, 5, 5]; rez2 = sum(L) - min(L) - max(L)
        L = [5, 5, 1]; rez3 = sum(L) - min(L) - max(L)
        L = [2, 2, 2]; rez4 = sum(L) - min(L) - max(L)
        run_test=acMainSection(test1=test1,test2=test2,test3=test3,test4=test4)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
        self.assertEqual(run_test["test4"], rez4, "Вредност променљиве 'test4' треба да буде %s" % rez4)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id16">
<h3>Задатак 6.<a class="headerlink" href="#id16" title="Permalink to this headline">¶</a></h3>
<p>Један наставник информатике закључује оцене тако што пре рачунања средње оцене одбаци највећу и најмању оцену,
па средњу вредност рачуна на основу преосталих оцена. Напиши функцију <code class="docutils literal notranslate"><span class="pre">srednja_vrednost_bez_ekstrema(ocene)</span></code> која
за списак оцена које је неки ученик добио из информатике рачуна средњу оцену на описани начин.
Напиши тело функције, па онда провери како функција ради.
(Упутство: Задатак можеш решити употребом стандардних функција <code class="docutils literal notranslate"><span class="pre">sum</span></code>, <code class="docutils literal notranslate"><span class="pre">min</span></code>, <code class="docutils literal notranslate"><span class="pre">max</span></code> и <code class="docutils literal notranslate"><span class="pre">len</span></code>)</p>

<div data-childcomponent="zadatak1-6" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak1-6 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def srednja_vrednost_bez_ekstrema(ocene):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = srednja_vrednost_bez_ekstrema([5, 5, 5, 5, 5])
test2 = srednja_vrednost_bez_ekstrema([1, 2, 5])
test3 = srednje_vrednost_bez_ekstrema([1, 2, 3, 4, 5])
# -*- acsection: after-main -*-
print(test1, test2, test3)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        L = [5, 5, 5, 5, 5]; rez1 = (sum(L) - min(L) - max(L))/(len(L) - 2)
        L = [1, 2, 5];       rez2 = (sum(L) - min(L) - max(L))/(len(L) - 2)
        L = [1, 2, 3, 4, 5]; rez3 = (sum(L) - min(L) - max(L))/(len(L) - 2)
        run_test = acMainSection(test1=test1,test2=test2,test3=test3)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="id17">
<h2>Задаци из Алгоритамске збирке задатака<a class="headerlink" href="#id17" title="Permalink to this headline">¶</a></h2>
<p>У следећим задацима можете да увежбате основне операције са низовима.</p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/ispis_u_obratnom_redosledu">https://petlja.org/biblioteka/r/problemi/Zbirka/ispis_u_obratnom_redosledu</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/parni_i_neparni_elementi">https://petlja.org/biblioteka/r/problemi/Zbirka/parni_i_neparni_elementi</a></p>
<div class="line-block">
<div class="line"><br /></div>
</div>
</div>
<div class="section" id="id18">
<h2>(Администрација)<a class="headerlink" href="#id18" title="Permalink to this headline">¶</a></h2>
<p><em>Овде не треба ништа да радиш! Слободно пређи на следећу лекцију!</em></p>

<div data-childcomponent="Zikine_ocene" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=Zikine_ocene data-lang="python" data-autorun="true"
    data-hidecode="true"  data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
Zikine_ocene = [2, 4, 3, 5, 4, 5, 4, 3]
</textarea>
</div>
</div>
</div>


    </div>
    
    <div>
<div class="lecture-prev-next-toggle">
  <a  href="index.html" id="prevLectureLink" class="lecture-toggler-arrows"><i class="fas fa-arrow-left fa"></i> <span> Претходна лекција </span> </a>
  
  <a id="nextLectureLink" class="lecture-toggler-arrows"><span> Следећа лекција</span>  <i class="fas fa-arrow-right fa"></i> </a>
</div>

<script type="text/javascript">
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
</div>

    
  </div>
</div>

<div class="petljadoc-resources">
  <span class="pull-left">(Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>) <br> © 2022 Petlja</span>
</div>


<div class="lectureToolbox">

  <label class="font-size-toggle" id="fontSizeToggle">

      <i class="fas fa-font"></i><label>Слова</label>
  </label>
  <div class="font-change d-none" id="fontChangeModal">
      <span class="dec-letter-icon" id="decLetterIcon">A-</span>
      <label id="currentFontSize">16</label>
      <span class="inc-letter-icon" id="incLetterIcon">A+</span>
  </div>
</div>

</div>

  </body>
</html>