<!DOCTYPE html>
<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Бројачи и суме</title>
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" type="text/css" href="../_static/simanim.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pycode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/p5js.css" />
    <link rel="stylesheet" type="text/css" href="../_static/gallery.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dbDirective.css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="https://fonts.googleapis.com/css?family=Source+Sans+Pro" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja-course.css" type="text/css" />
    <link rel="stylesheet" href="../_static/nbstyle.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/language_data.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/karelCorner.js"></script>
    <script type="text/javascript" src="../_static/karelRobot.js"></script>
    <script type="text/javascript" src="../_static/karelWorld.js"></script>
    <script type="text/javascript" src="../_static/karelRobotDrawer.js"></script>
    <script type="text/javascript" src="../_static/karelUI.js"></script>
    <script type="text/javascript" src="../_static/karel.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/notes.js"></script>
    <script type="text/javascript" src="../_static/pygamelib-init.js"></script>
    <script type="text/javascript" src="../_static/blockly/blockly_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/blocks_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/python_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/msg-sr.js"></script>
    <script type="text/javascript" src="../_static/blockpy/utilities.js"></script>
    <script type="text/javascript" src="../_static/blockpy/python_errors.js"></script>
    <script type="text/javascript" src="../_static/blockpy/ast_node_visitor.js"></script>
    <script type="text/javascript" src="../_static/blockpy/abstract_interpreter.js"></script>
    <script type="text/javascript" src="../_static/blockpy/pytifa.js"></script>
    <script type="text/javascript" src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script type="text/javascript" src="../_static/blockpy/python_to_blockly.js"></script>
    <script type="text/javascript" src="../_static/blockpy/imported.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/class.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/comment.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/dict.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/if.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/io.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/lists.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/sets.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/loops.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/parking.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/tuple.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/turtles.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/text.js"></script>
    <script type="text/javascript" src="../_static/blockpy-modal.js"></script>
    <script type="text/javascript" src="../_static/simanim.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/pyodide/v0.16.1/full/pyodide.js"></script>
    <script type="text/javascript" src="../_static/pycode.js"></script>
    <script type="text/javascript" src="//toolness.github.io/p5.js-widget/p5-widget.js"></script>
    <script type="text/javascript" src="../_static/p5js.js"></script>
    <script type="text/javascript" src="../_static/gallery.js"></script>
    <script type="text/javascript" src="../_static/dbDirective.js"></script>
    <script type="text/javascript" src="../_static/sql.js"></script>
    <script type="text/javascript" src="../_static/translations.js"></script>
    <script async="async" type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="../_static/course-errors.js"></script>
    <script type="text/javascript" src="../_static/petlja.js"></script>
    <script type="text/javascript" src="../_static/require.js"></script>
    <meta charset='utf-8'>
    <meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
    <meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
    <script type="text/javascript">
      eBookConfig = {};
      eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
      eBookConfig.app = eBookConfig.host + '/runestone';
      eBookConfig.ajaxURL = eBookConfig.app + '/ajax/';
      eBookConfig.course = 'Настава програмирања за Други разред гимназије';
      eBookConfig.logLevel = '0';
      eBookConfig.loginRequired = 'false';
      eBookConfig.build_info = "";
      eBookConfig.isLoggedIn = false;
      eBookConfig.useRunestoneServices = false;
      eBookConfig.python3 = true;
      eBookConfig.basecourse = 'Настава програмирања за Други разред гимназије';
      eBookConfig.runestone_version = '';
      eBookConfig.imagesDir = '../_images/';
      eBookConfig.staticDir = '../_static/';
      if (typeof (Sk) != "undefined")
        Sk.imgPath = eBookConfig.imagesDir;
    </script>

  </head>

  <body>



<div id="navbar" class="content-header">
  <div class="right-side-nav">
    
</div>

<div class="left-side-nav">
  <div class="course-logo-nav">
      <a target="_blank" href="https://petlja.org">
          <img src="../_static/img/petlja-logo-wt.svg" alt="petlja.org" id="lightPetljaLogo">
      </a>
  </div>
  <div class="course-logo-nav-kabinet">
      <a target="_blank" href="https://petlja.org/net.kabinet">
          <img src="../_static/img/netkabinet-logo-strelica-wt.svg" alt="petlja.org" id="lightNetKabinetLogo">
      </a>
  </div>
</div>


<div class="content-header-label" id = "content-header-label" href="../">Анализа података у Џупитеру и технике програмирања за други разред гимназије</div>

</div>




<div class="petljaCourseContainer">

<div class="petljaCourseContent">
  
           <div class="tableOfContent">
            <h5> Садржај </h5>
            </div>
           <div>
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J00>
                    <h5>Рад са подацима у текстуалном програмском језику (ПТПЈ)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J00/J00.html><div class="studioLecture" id=activity-J00><i class="fas fa-file-alt activity-icon"></i>Исходи</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J0A>
                    <h5>Припрема</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J0A/J0A.html><div class="studioLecture" id=activity-J0A><i class="fas fa-file-alt activity-icon"></i>Припрема</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J01>
                    <h5>ПТПЈ.1. Увод у Џупитер (*Jupyter*)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J01/J01.html><div class="studioLecture" id=activity-J01><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J02>
                    <h5>ПТПЈ.2. Низови података и линијски дијаграми</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J02/J02.html><div class="studioLecture" id=activity-J02><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J02/testJ02.html><div class="studioLecture" id=activity-testJ02><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J03>
                    <h5>ПТПЈ.3. Хистограми и боје</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J03/J03.html><div class="studioLecture" id=activity-J03><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J03/testJ03.html><div class="studioLecture" id=activity-testJ03><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J04>
                    <h5>ПТПЈ.4. Просек и медијана низа бројева</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J04/J04.html><div class="studioLecture" id=activity-J04><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J04/testJ04.html><div class="studioLecture" id=activity-testJ04><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J05>
                    <h5>ПТПЈ.5. Фреквенцијска анализа и секторски дијаграми</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J05/J05.html><div class="studioLecture" id=activity-J05><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J05/testJ05.html><div class="studioLecture" id=activity-testJ05><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J06>
                    <h5>ПТПЈ.6. Табеларно представљени подаци</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J06/J06.html><div class="studioLecture" id=activity-J06><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J06/testJ06.html><div class="studioLecture" id=activity-testJ06><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J07>
                    <h5>ПТПЈ.7. Индексирање и транспоновање табеле</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J07/J07.html><div class="studioLecture" id=activity-J07><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J07/testJ07.html><div class="studioLecture" id=activity-testJ07><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J08>
                    <h5>ПТПЈ.8. Модификације табеле и записивање табеле у датотеку</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J08/J08.html><div class="studioLecture" id=activity-J08><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J08/testJ08.html><div class="studioLecture" id=activity-testJ08><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J09>
                    <h5>ПТПЈ.9. Сортирање, филтрирање и фреквенцијска анализа</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J09/J09.html><div class="studioLecture" id=activity-J09><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J09/testJ09.html><div class="studioLecture" id=activity-testJ09><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J10>
                    <h5>ПТПЈ.10. Џупитер и Ексел</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J10/J10.html><div class="studioLecture" id=activity-J10><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J10/testJ10.html><div class="studioLecture" id=activity-testJ10><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y00>
                    <h5>Програмирање (ПРОГ)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y00/Y00.html><div class="studioLecture" id=activity-Y00><i class="fas fa-file-alt activity-icon"></i>Исходи</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y01>
                    <h5>ПРОГ.1. Увод у процесирање низова података у Пајтону</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y01/L801.html><div class="studioLecture" id=activity-L801><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y01/test801.html><div class="studioLecture" id=activity-test801><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y02>
                    <h5>ПРОГ.2. Бројачи и суме</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y02/L802.html><div class="studioLecture" id=activity-L802><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y02/test802.html><div class="studioLecture" id=activity-test802><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y03>
                    <h5>ПРОГ.3. Низови симбола (*стрингови* или *ниске*)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y03/L803.html><div class="studioLecture" id=activity-L803><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y03/test803.html><div class="studioLecture" id=activity-test803><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y04>
                    <h5>ПРОГ.4. Угнежђене петље</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y04/L804.html><div class="studioLecture" id=activity-L804><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y04/test804.html><div class="studioLecture" id=activity-test804><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y05>
                    <h5>ПРОГ.5. Сортирање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y05/L805.html><div class="studioLecture" id=activity-L805><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y05/test805.html><div class="studioLecture" id=activity-test805><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y06>
                    <h5>ПРОГ.6. Филтрирање и претраживање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y06/L806.html><div class="studioLecture" id=activity-L806><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y06/test806.html><div class="studioLecture" id=activity-test806><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y07>
                    <h5>ПРОГ.7. Рад са текстуалним датотекама</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y07/L807.html><div class="studioLecture" id=activity-L807><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y07/test807.html><div class="studioLecture" id=activity-test807><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                </div>
</div>



<div class="lectureContent">

  <div id="errors">
  </div>
  
  <div class="lectureContentMaterial">
  
    <div id="main-content">
    
  <div class="section" id="id1">
<h1>Бројачи и суме<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>У овој лекцији</p>
<ol class="arabic simple">
<li><p>уводимо појам бројача,</p></li>
<li><p>показујемо како се рачунају суме, и</p></li>
<li><p>то знање примењујемо на проблем рачунања просека низа бројева.</p></li>
</ol>
<div class="section" id="id2">
<h2>Бројач<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>
<p>Често је потребно да пребројимо колико у низу има елемената који имају неку особину.
Стандардни начин да се ова врста проблема реши је да се једна променљива задужи за
посао бројања и онда се она зове <em>бројач</em>.</p>
<div class="section" id="id3">
<h3>Пример.<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">prebroj_petice</span></code> која у датом списку оцена броји петице. На пример,</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">prebroj_petice</span><span class="p">([</span><span class="mi">3</span><span class="p">,</span> <span class="mi">1</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">5</span><span class="p">])</span>
</pre></div>
</div>
<p>треба да врати 3 зато што наведени списак оцена садржи три петице.</p>
<p><em>Решење.</em></p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">prebroj_petice</span><span class="p">(</span><span class="n">spisak</span><span class="p">):</span>
    <span class="n">broj</span> <span class="o">=</span> <span class="mi">0</span>
    <span class="k">for</span> <span class="n">ocena</span> <span class="ow">in</span> <span class="n">spisak</span><span class="p">:</span>
        <span class="k">if</span> <span class="n">ocena</span> <span class="o">==</span> <span class="mi">5</span><span class="p">:</span> <span class="n">broj</span> <span class="o">+=</span> <span class="mi">1</span>
    <span class="k">return</span> <span class="n">broj</span>
</pre></div>
</div>
<p>Функција <code class="docutils literal notranslate"><span class="pre">prebroj_petice</span></code> прими списак оцена у коме треба да преброји петице и онда се понаша овако.</p>
<ol class="arabic simple">
<li><p>Променљива <code class="docutils literal notranslate"><span class="pre">broj</span></code> ће бројати петице. Пошто још нисмо кренули са провером оцена, постављамо њену вредност на 0 (за сада нисмо наишли ни на једну петицу).</p></li>
<li><p>Наредбом <code class="docutils literal notranslate"><span class="pre">for</span> <span class="pre">ocena</span> <span class="pre">in</span> <span class="pre">spisak:</span></code> пролазимо кроз <code class="docutils literal notranslate"><span class="pre">spisak</span></code> тако што променљива <code class="docutils literal notranslate"><span class="pre">ocena</span></code> редом узима вредности које су наведене на списку. Према томе, за сваку оцену са списка ћемо проверити да ли је једнака са 5, и ако јесте увећаћемо вредност променљиве <code class="docutils literal notranslate"><span class="pre">broj</span></code>.</p></li>
<li><p>На крају враћамо вредност променљиве <code class="docutils literal notranslate"><span class="pre">broj</span></code> као резултат рада функције.</p></li>
</ol>
<p>Да погледамо како функција ради:</p>

<div data-childcomponent="primer2-1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="def prebroj_petice(spisak):
    broj = 0
    for ocena in spisak:
        if ocena == 5: broj += 1
    return broj

print(prebroj_petice([3, 1, 5, 4, 5, 2, 5]))
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P02</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Prebroj_petice.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
<div class="section" id="id4">
<h3>Пример.<a class="headerlink" href="#id4" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">prebroj_ocene</span></code> која броји колико пута се у датом списку оцена појављује дата оцена.
На пример,</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">prebroj_ocene</span><span class="p">([</span><span class="mi">3</span><span class="p">,</span> <span class="mi">1</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">1</span><span class="p">,</span> <span class="mi">2</span><span class="p">],</span> <span class="mi">1</span><span class="p">)</span>
</pre></div>
</div>
<p>треба да врати 2 зато што наведени списак оцена садржи две јединице.</p>
<p><em>Решење.</em></p>

<div data-childcomponent="primer2-1b" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-1b data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="def prebroj_ocene(spisak, n):
    broj = 0
    for ocena in spisak:
        if ocena == n: broj += 1
    return broj

print(prebroj_ocene([3, 1, 5, 4, 5, 2, 5, 1, 2], 1))
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P02</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Prebroj_ocene.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
<div class="section" id="id5">
<h3>Пример*.<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">prebroj_cifre</span></code> која за дати позитиван цео број утврђује колико он има цифара. На пример,</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">prebroj_cifre</span><span class="p">(</span><span class="mi">12487</span><span class="p">)</span>
</pre></div>
</div>
<p>треба да врати 5 зато што број 12487 има 5 цифара.</p>
<p><em>Решење.</em></p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">prebroj_cifre</span><span class="p">(</span><span class="n">n</span><span class="p">):</span>
    <span class="n">br_cif</span> <span class="o">=</span> <span class="mi">0</span>
    <span class="k">while</span> <span class="n">n</span> <span class="o">&gt;</span> <span class="mi">0</span><span class="p">:</span>
        <span class="n">br_cif</span> <span class="o">+=</span> <span class="mi">1</span>
        <span class="n">n</span> <span class="o">//=</span> <span class="mi">10</span>
    <span class="k">return</span> <span class="n">br_cif</span>
</pre></div>
</div>
<p>Идеја програма је следећа.</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">све</span> <span class="n">док</span> <span class="n">број</span> <span class="s2">&quot;има још цифара&quot;</span><span class="p">:</span>
    <span class="n">увећај</span> <span class="n">бројач</span> <span class="n">цифара</span> <span class="n">за</span> <span class="mi">1</span>
    <span class="s2">&quot;откини&quot;</span> <span class="n">последњу</span> <span class="n">цифру</span> <span class="n">броја</span>
</pre></div>
</div>
<p>Условом <code class="docutils literal notranslate"><span class="pre">n</span> <span class="pre">&gt;</span> <span class="pre">0</span></code> проверавамо да ли је остала још нека цифра у броју,
док командом <code class="docutils literal notranslate"><span class="pre">n</span> <span class="pre">//=</span> <span class="pre">10</span></code> „откинемо” последњу цифру броја. Ево примера:</p>

<div data-childcomponent="primer2-2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
n = 12487
n //= 10
print(n)
</textarea>
</div>
<p>Ето: „откинули” смо последњу цифру броја <em>n</em>. Да пробамо да „откинемо” последње две цифре броја <em>n</em>:</p>

<div data-childcomponent="primer2-3" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-3 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
n = 12487
n //= 10
n //= 10
print(n)
</textarea>
</div>
<p>Дакле, у сваком пролазу кроз петљу броју <em>n</em> (крволочно) „откинемо” по једну цифру, па када падне на нулу знамо да смо му откинули све цифре, и тако смо сигурни да смо их све и пребројали.</p>
<p>Да проверимо функцију:</p>

<div data-childcomponent="primer2-4" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-4 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="def prebroj_cifre(n):
    br_cif = 0
    while n &gt; 0:
        br_cif += 1
        n //= 10
    return br_cif

print(prebroj_cifre(12487))
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P02</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Prebroj_cifre.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
<div class="section" id="id6">
<h3>Пример.<a class="headerlink" href="#id6" title="Permalink to this headline">¶</a></h3>
<p>Подсетимо се да уграђена функција <code class="docutils literal notranslate"><span class="pre">len</span></code> одређује број елемената у низу. Она се може испрограмирати овако (мада је, наравно, увек ефикасније позвати уграђену функцију):</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">duzina_niza</span><span class="p">(</span><span class="n">niz</span><span class="p">):</span>
    <span class="n">broj</span> <span class="o">=</span> <span class="mi">0</span>
    <span class="k">for</span> <span class="n">x</span> <span class="ow">in</span> <span class="n">niz</span><span class="p">:</span>
        <span class="n">broj</span> <span class="o">+=</span> <span class="mi">1</span>
    <span class="k">return</span> <span class="n">broj</span>
</pre></div>
</div>
</div>
</div>
<div class="section" id="id7">
<h2>Збир елемената низа<a class="headerlink" href="#id7" title="Permalink to this headline">¶</a></h2>
<p>Често је потребно сабрати елементе неког низа бројева, што се често може урадити још током учитавања низа.
Ево примера.</p>
<div class="section" id="id8">
<h3>Пример 1.<a class="headerlink" href="#id8" title="Permalink to this headline">¶</a></h3>
<p>Напиши Пајтон програм који од корисника учитава цео број <span class="math notranslate nohighlight">\(n\)</span>, потом учитава <span class="math notranslate nohighlight">\(n\)</span> реалних бројева
и исписује њихов збир.</p>

<div data-childcomponent="primer2-5" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-5 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
             data-passivecode="true"  data-hidehistory=true
    data-includesrc="n = int(input())
zbir = 0
for i in range(n):
    x = float(input())
    zbir += x
print(zbir)
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши програм у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P02</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Saberi_n.py</span></code> и изврши га.</p>
<p>Ево и кратке видео демонстрације:</p>

    <div  id="B5Vj5-nycq8" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/B5Vj5-nycq8/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-B5Vj5-nycq8" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-B5Vj5-nycq8">
        </div>
    </div>
    
    
    
    </p></div></div>
</div>
<div class="section" id="id9">
<h3>Пример 2.<a class="headerlink" href="#id9" title="Permalink to this headline">¶</a></h3>
<p>Напиши Пајтон програм који од корисника учитава реалне бројеве све док корисник не унесе реч „stop” након чега
исписује збир унетих бројева.</p>

<div data-childcomponent="primer2-5-5" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-5-5 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
             data-passivecode="true"  data-hidehistory=true
    data-includesrc="zbir = 0
unos = input()
while unos != &quot;stop&quot;:
    x = float(unos)
    zbir += x
    unos = input()
print(zbir)
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши програм у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P02</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Sabiraj_do_stop.py</span></code> и изврши га.</p>
<p>Ево и кратке видео демонстрације:</p>

    <div  id="UlH-NvFuUPA" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/UlH-NvFuUPA/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-UlH-NvFuUPA" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-UlH-NvFuUPA">
        </div>
    </div>
    
    
    
    </p></div></div>
</div>
<div class="section" id="id10">
<h3>Пример 3.<a class="headerlink" href="#id10" title="Permalink to this headline">¶</a></h3>
<p>Милица добија џепарац који је сваког месеца исти. За следећи месец Милица је испланирала низ трошкова.
Напиши Пајтон програм који од корисника учитава висину Миличиног џепарца, и онда списак трошкова које је милица планирала,
и исписује поруку о томе да ли су њени планови реални. Програм учитава трошкове све док корисник не унесе реч „stop”, а
исписује „Dovoljan je” или „Nije dovoljan” у зависности од тога да ли је Миличин џепарац довољан за њене планове или не.</p>

<div data-childcomponent="primer2-5-6" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-5-6 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
             data-passivecode="true"  data-hidehistory=true
    data-includesrc="dzeparac = float(input())
troskovi = 0
unos = input()
while unos != &quot;stop&quot;:
    x = float(unos)
    troskovi += x
    unos = input()
if dzeparac &gt;= troskovi:
    print(&quot;Dovoljan je&quot;)
else:
    print(&quot;Nije dovoljan&quot;)

"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши програм у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P02</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Milicin_dzeparac.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
</div>
<div class="section" id="id11">
<h2>Просек елемената низа<a class="headerlink" href="#id11" title="Permalink to this headline">¶</a></h2>
<p>Да се подсетимо, <em>просек</em> низа бројева се рачуна овако:</p>
<div class="math notranslate nohighlight">
\[\hbox{просек низа } [x_1, x_2, \dots, x_n] = \frac{x_1 + x_2 + \dots + x_n}{n}.\]</div>
<p>Дакле, да бисмо добили просек низа бројева саберемо све елементе низа и тако добијени број поделимо дужином низа.</p>
<p>У Пајтону се обе ове активности могу обавити позивом одговарајуће уграђене функције: функција <code class="docutils literal notranslate"><span class="pre">len</span></code>
рачуна дужину низа, док функција <code class="docutils literal notranslate"><span class="pre">sum</span></code> рачуна збир елемената низа бројева.
Како ћемо у много наврата рачунати просек, показаћемо како се дефинише функција која га рачуна:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">def</span> <span class="nf">prosek</span><span class="p">(</span><span class="n">L</span><span class="p">):</span>
    <span class="k">return</span> <span class="nb">sum</span><span class="p">(</span><span class="n">L</span><span class="p">)</span> <span class="o">/</span> <span class="nb">len</span><span class="p">(</span><span class="n">L</span><span class="p">)</span>
</pre></div>
</div>
<div class="section" id="id12">
<h3>Пример 1.<a class="headerlink" href="#id12" title="Permalink to this headline">¶</a></h3>
<p>Мирко из музичког има 5, 4, 5 и још једну оцену која се не види јер се преко ње разлило мастило.
Просек његових оцена из музичког је 4. Која је то оцена преко које се разлило мастило?</p>
<p><em>Решење.</em> Испробаћемо разне могућности:</p>

<div data-childcomponent="primer2-7" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-7 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
def prosek(L):
    return sum(L) / len(L)

print("Prvi pokusaj:",    prosek([5,4,5,  5]))
print("Drugi pokusaj:",   prosek([5,4,5,  4]))
print("Treci pokusaj:",   prosek([5,4,5,  3]))
print("Cetvrti pokusaj:", prosek([5,4,5,  2]))
print("Peti pokusaj:",    prosek([5,4,5,  1]))
</textarea>
</div>
<p>Која је то оцена?</p>
</div>
<div class="section" id="id13">
<h3>Пример 2.<a class="headerlink" href="#id13" title="Permalink to this headline">¶</a></h3>
<p>Написати програм који са стандардног улаза (од корисника) учитава природан број <code class="docutils literal notranslate"><span class="pre">n</span></code>, потом низ од <code class="docutils literal notranslate"><span class="pre">n</span></code>
реалних бројева и онда утврђује колико од учитаних реалних бројева је веће од њиховог просека.</p>

<div data-childcomponent="primer2-6-iznad-proseka-1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-6-iznad-proseka-1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
             data-passivecode="true"  data-hidehistory=true
    data-includesrc="n = int(input())
brojevi = []
for i in range(n):
  x = float(input())
  brojevi.append(x)

prosek = sum(brojevi)/n
k = 0 # brojač

for i in range(n):
  if brojevi[i] &gt; prosek:
    k += 1

print(k)
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Покрени IDLE окружење, из фолдера <code class="docutils literal notranslate"><span class="pre">P02</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Iznad_proseka.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
<div class="section" id="id14">
<h3>Пример 3.<a class="headerlink" href="#id14" title="Permalink to this headline">¶</a></h3>
<p>Написати програм који са стандардног улаза (од корисника) учитава природан број <code class="docutils literal notranslate"><span class="pre">n</span></code>, потом низ од <code class="docutils literal notranslate"><span class="pre">n</span></code>
тачака у равни и онда рачуна и штампа координате тежишта учитаног низа тачака. Координате једне тачке уносимо
у истом реду раздвојене једном празнином, на пример овако:</p>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="mi">4</span>
<span class="o">-</span><span class="mf">1.5</span> <span class="mf">2.2</span>
<span class="mf">3.0</span> <span class="mf">1.5</span>
<span class="mf">4.25</span> <span class="o">-</span><span class="mf">2.29</span>
<span class="mf">3.21</span> <span class="mf">8.0</span>
</pre></div>
</div>
<p><em>Решење.</em></p>
<p>Тачка у равни се може описати својим <em>координатама</em> што је пар реалних бројева, на пример овако: <span class="math notranslate nohighlight">\(A(2.5, -1.4)\)</span>.
У програму нам зато требају два низа: низ <span class="math notranslate nohighlight">\(x\)</span>-координата и низ <span class="math notranslate nohighlight">\(y\)</span>-координата учитаних тачака.
Координате тежишта коначног скупа тачака се рачуна тако што се израчуна просек
<span class="math notranslate nohighlight">\(x\)</span>-координата и просек <span class="math notranslate nohighlight">\(y\)</span>-координата – то су <span class="math notranslate nohighlight">\(x\)</span>-координата и <span class="math notranslate nohighlight">\(y\)</span>-координата тежишта.</p>

<div data-childcomponent="primer2-6-iznad-proseka" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer2-6-iznad-proseka data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
             data-passivecode="true"  data-hidehistory=true
    data-includesrc="n = int(input())
x_koord = []
y_koord = []
for i in range(n):
  tacka = input()
  koord = tacka.split()
  x = float(koord[0])
  y = float(koord[1])
  x_koord.append(x)
  y_koord.append(y)

T_x = sum(x_koord) / n
T_y = sum(y_koord) / n

print(T_x, T_y)

"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Покрени IDLE окружење, из фолдера <code class="docutils literal notranslate"><span class="pre">P02</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Teziste.py</span></code> и изврши га.</p>
<p>Ево и кратке видео демонстрације:</p>

    <div  id="voAJXV7kNIc" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/voAJXV7kNIc/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-voAJXV7kNIc" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-voAJXV7kNIc">
        </div>
    </div>
    
    
    
    </p></div></div>
</div>
</div>
<div class="section" id="id15">
<h2>Задаци<a class="headerlink" href="#id15" title="Permalink to this headline">¶</a></h2>
<div class="section" id="id16">
<h3>Задатак 1.<a class="headerlink" href="#id16" title="Permalink to this headline">¶</a></h3>
<p>Напиши Пајтон функцију <code class="docutils literal notranslate"><span class="pre">broj_pozitivnih</span></code> која утврђује колико у датом низу има позитивних бројева. На пример,
<code class="docutils literal notranslate"><span class="pre">broj_pozitivnih([2,</span> <span class="pre">-1,</span> <span class="pre">3,</span> <span class="pre">0,</span> <span class="pre">-5,</span> <span class="pre">1,</span> <span class="pre">9])</span></code> треба да врати 4.</p>

<div data-childcomponent="zadatak2-1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak2-1 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def broj_pozitivnih(L):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = broj_pozitivnih([1, 2, 3, 4, 5])
test2 = broj_pozitivnih([-1, -2, -3, -4])
test3 = broj_pozitivnih([2, -1, 3, 0, -5, 1, 9])
# -*- acsection: after-main -*-
print(test1, test2, test3)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        run_test = acMainSection(test1=test1,test2=test2,test3=test3)
        self.assertEqual(run_test["test1"], 5, "Вредност променљиве 'test1' треба да буде %s" % 5)
        self.assertEqual(run_test["test2"], 0, "Вредност променљиве 'test2' треба да буде %s" % 0)
        self.assertEqual(run_test["test3"], 4, "Вредност променљиве 'test3' треба да буде %s" % 4)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id17">
<h3>Задатак 2.<a class="headerlink" href="#id17" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">broj_prestupnih</span></code> која прими две године и утврди колико преступних година има у
наведеном интервалу. На пример, <code class="docutils literal notranslate"><span class="pre">broj_prestupnih(1987,</span> <span class="pre">2019)</span></code> треба да врати 8.
У томе ће ти помоћи функција <code class="docutils literal notranslate"><span class="pre">prestupna_godina</span></code> која проверава да ли је година преступна.</p>

<div data-childcomponent="zadatak2-2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak2-2 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
def prestupna_godina(g):
    if g % 400 == 0 or (g % 100 != 0 and g % 4 == 0):
        return True
    else:
        return False
# -*- acsection: main -*-

def broj_prestupnih(g1, g2):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = broj_prestupnih(1987, 2019)
test2 = broj_prestupnih(2019, 2021)
test3 = broj_prestupnih(1997, 2003)
# -*- acsection: after-main -*-
print(test1, test2, test3)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        def __brp(g1, g2):
           n = 0
           g = g1
           while g <= g2:
               if prestupna_godina(g): n += 1
               g += 1
           return n
        rez1 = __brp(1987, 2019)
        rez2 = __brp(2019, 2021)
        rez3 = __brp(1997, 2003)
        run_test = acMainSection(test1=test1,test2=test2,test3=test3)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id18">
<h3>Задатак 3.<a class="headerlink" href="#id18" title="Permalink to this headline">¶</a></h3>
<p>Напиши функцију <code class="docutils literal notranslate"><span class="pre">prebroj_max(A)</span></code> која утврђује колико пута се у датом низу бројева појављује највећа вредност низа.
(Напомена: највећи елемент низа се може израчунати помоћу функције <code class="docutils literal notranslate"><span class="pre">max</span></code>.)</p>

<div data-childcomponent="zadatak2-3" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak2-3 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def prebroj_max(A):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = prebroj_max([1, 2, 3, 2, 3, 3, 3])
test2 = prebroj_max([0, 0, 0])
test3 = prebroj_max([5])
# -*- acsection: after-main -*-
print(test1, test2, test3)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        run_test = acMainSection(test1=test1,test2=test2,test3=test3)
        self.assertEqual(run_test["test1"], 4, "Вредност променљиве 'test1' треба да буде %s" % 4)
        self.assertEqual(run_test["test2"], 3, "Вредност променљиве 'test2' треба да буде %s" % 3)
        self.assertEqual(run_test["test3"], 1, "Вредност променљиве 'test3' треба да буде %s" % 1)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id19">
<h3>Задатак 4.<a class="headerlink" href="#id19" title="Permalink to this headline">¶</a></h3>
<p>Милoraд у дневнику има неколико оцена из информатике. Напиши Пајтон функцију <code class="docutils literal notranslate"><span class="pre">koliko_do_odlicnog_uspeha</span></code>
која прима низ Милорадових оцена и онда одређује који је најмањи број петица које Милорад треба да добије из информатике
да би наставник морао да му закључи 5. (Према Правилнику о оцењивању, наставник мора да закључи оцену 5
ученику чији просек оцена је барем 4,50.)</p>
<p>На пример, <code class="docutils literal notranslate"><span class="pre">koliko_do_odlicnog_uspeha([3,3])</span></code> треба да врати 6 јер је (3 + 3 + 5 + 5 + 5 + 5 + 5 + 5) / 8 = 4,50.</p>

<div data-childcomponent="zadatak2-6" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak2-6 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3 test4"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
def prosek(A):
    return sum(A) / len(A)
# -*- acsection: main -*-

def koliko_do_odlicnog_uspeha(A):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = koliko_do_odlicnog_uspeha([4])
test2 = koliko_do_odlicnog_uspeha([4, 4])
test3 = koliko_do_odlicnog_uspeha([3, 3])
test4 = koliko_do_odlicnog_uspeha([5, 5, 5])
# -*- acsection: after-main -*-
print(test1, test2, test3, test4)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        def __pr(A):
            return sum(A) / len(A)
        def __kdou(L):
            n=0
            while __pr(L) < 4.5:
                n += 1
                L.append(5)
            return n
        rez1 = __kdou([4])
        rez2 = __kdou([4, 4])
        rez3 = __kdou([3, 3])
        rez4 = __kdou([5, 5, 5])
        run_test = acMainSection(test1=test1,test2=test2,test3=test3,test4=test4)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
        self.assertEqual(run_test["test4"], rez4, "Вредност променљиве 'test4' треба да буде %s" % rez4)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id20">
<h3>Задатак 5.<a class="headerlink" href="#id20" title="Permalink to this headline">¶</a></h3>
<p>Напиши функцију <code class="docutils literal notranslate"><span class="pre">zbir_cifara</span></code> која рачуна збир цифара датог позитивног целог броја. На пример,
<code class="docutils literal notranslate"><span class="pre">zbir_cifara(12487)</span></code> треба да врати 22.</p>

<div data-childcomponent="zadatak2-4" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak2-4 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def zbir_cifara(n):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = zbir_cifara(12487)
test2 = zbir_cifara(1248700)
test3 = zbir_cifara(9)
# -*- acsection: after-main -*-
print(test1, test2, test3)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        def __zb(n):
            s=0
            while n > 0:
                s += n % 10
                n //= 10
            return s
        rez1 = __zb(12487)
        rez2 = __zb(1248700)
        rez3 = __zb(9)
        run_test = acMainSection(test1=test1,test2=test2,test3=test3)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id21">
<h3>Задатак 6*.<a class="headerlink" href="#id21" title="Permalink to this headline">¶</a></h3>
<p><em>Судбински број</em> неке особе се добија овако:</p>
<ol class="arabic simple">
<li><p>Запишемо датум рођења те особе као један низ цифара, на пример: 15. мај 2001. –&gt; 15052001</p></li>
<li><p>Саберемо цифре тог броја.</p></li>
<li><p>Ако смо добили једноцифрени број, то је судбински број особе.</p></li>
<li><p>Ако нисмо добили једноцифрени број, саберемо цифре добијеног броја, и тако све док не добијемо једноцифрени број.</p></li>
</ol>
<p><strong>Важна напомена!</strong> Ово је само вежбица из програмирања. Немојте придавати никакав значај добијеном броју! Нумерологија је изашла из моде пре око хиљаду година.</p>
<p>Напиши функцију <code class="docutils literal notranslate"><span class="pre">sudbinski_broj</span></code> која рачуна судбински број неке особе на основу броја који одговара њеном датуму рођења.</p>

<div data-childcomponent="zadatak2-5b" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak2-5b data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def sudbinski_broj(n):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = sudbinski_broj(10102010)
test2 = sudbinski_broj(29091999)
# -*- acsection: after-main -*-
print(test1, test2)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        def __zb(n):
            s=0
            while n > 0:
                s += n % 10
                n //= 10
            return s
        def __sb(n):
            while n > 9: n = __zb(n)
            return n
        rez1 = __sb(10102010)
        rez2 = __sb(29091999)
        run_test = acMainSection(test1=test1,test2=test2)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="pomoc1"   data-showtitle="Помоћ" data-hidetitle="Сакриј помоћ">
    
    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Не заборави да функције можеш дефинисати и у оквиру других функција.</p>

    </p></div></div>

    </div>
    </div>
</div>
<div class="section" id="id22">
<h2>Задаци из Алгоритамске збирке задатака<a class="headerlink" href="#id22" title="Permalink to this headline">¶</a></h2>
<p>Решавањем следећих задатака можете да увежбате додатне вештине у раду са листама.</p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/prosecno_odstupanje_od_minimuma">https://petlja.org/biblioteka/r/problemi/Zbirka/prosecno_odstupanje_od_minimuma</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/minimalno_odstupanje_od_proseka">https://petlja.org/biblioteka/r/problemi/Zbirka/minimalno_odstupanje_od_proseka</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/translacija_tacaka">https://petlja.org/biblioteka/r/problemi/Zbirka/translacija_tacaka</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/ciklicno_pomeranje_za_jedno_mesto">https://petlja.org/biblioteka/r/problemi/Zbirka/ciklicno_pomeranje_za_jedno_mesto</a></p>
</div>
</div>


    </div>
    
    <div>
<div class="lecture-prev-next-toggle">
  <a  href="index.html" id="prevLectureLink" class="lecture-toggler-arrows"><i class="fas fa-arrow-left fa"></i> <span> Претходно лекција </span> </a>
  
  <a id="nextLectureLink" class="lecture-toggler-arrows"><span> Следећа лекција</span>  <i class="fas fa-arrow-right fa"></i> </a>
</div>

<script type="text/javascript">
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
</div>

    
  </div>
</div>
<div class="petljadoc-resources">
  <span class="pull-left">(Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>) <br> © 2022 Petlja</span>
</div>

<div class="lectureToolbox">

  <label class="font-size-toggle" id="fontSizeToggle">

      <i class="fas fa-font"></i><label>Слова</label>
  </label>
  <div class="font-change d-none" id="fontChangeModal">
      <span class="dec-letter-icon" id="decLetterIcon">A-</span>
      <label id="currentFontSize">16</label>
      <span class="inc-letter-icon" id="incLetterIcon">A+</span>
  </div>
</div>

</div>

  </body>
</html>