<!DOCTYPE html>
<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Низови симбола (стрингови или ниске)</title>
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" type="text/css" href="../_static/simanim.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pycode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/p5js.css" />
    <link rel="stylesheet" type="text/css" href="../_static/gallery.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dbDirective.css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="https://fonts.googleapis.com/css?family=Source+Sans+Pro" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja-course.css" type="text/css" />
    <link rel="stylesheet" href="../_static/nbstyle.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/language_data.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/karelCorner.js"></script>
    <script type="text/javascript" src="../_static/karelRobot.js"></script>
    <script type="text/javascript" src="../_static/karelWorld.js"></script>
    <script type="text/javascript" src="../_static/karelRobotDrawer.js"></script>
    <script type="text/javascript" src="../_static/karelUI.js"></script>
    <script type="text/javascript" src="../_static/karel.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/notes.js"></script>
    <script type="text/javascript" src="../_static/pygamelib-init.js"></script>
    <script type="text/javascript" src="../_static/blockly/blockly_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/blocks_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/python_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/msg-sr.js"></script>
    <script type="text/javascript" src="../_static/blockpy/utilities.js"></script>
    <script type="text/javascript" src="../_static/blockpy/python_errors.js"></script>
    <script type="text/javascript" src="../_static/blockpy/ast_node_visitor.js"></script>
    <script type="text/javascript" src="../_static/blockpy/abstract_interpreter.js"></script>
    <script type="text/javascript" src="../_static/blockpy/pytifa.js"></script>
    <script type="text/javascript" src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script type="text/javascript" src="../_static/blockpy/python_to_blockly.js"></script>
    <script type="text/javascript" src="../_static/blockpy/imported.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/class.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/comment.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/dict.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/if.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/io.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/lists.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/sets.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/loops.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/parking.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/tuple.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/turtles.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/text.js"></script>
    <script type="text/javascript" src="../_static/blockpy-modal.js"></script>
    <script type="text/javascript" src="../_static/simanim.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/pyodide/v0.16.1/full/pyodide.js"></script>
    <script type="text/javascript" src="../_static/pycode.js"></script>
    <script type="text/javascript" src="//toolness.github.io/p5.js-widget/p5-widget.js"></script>
    <script type="text/javascript" src="../_static/p5js.js"></script>
    <script type="text/javascript" src="../_static/gallery.js"></script>
    <script type="text/javascript" src="../_static/dbDirective.js"></script>
    <script type="text/javascript" src="../_static/sql.js"></script>
    <script type="text/javascript" src="../_static/translations.js"></script>
    <script async="async" type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="../_static/course-errors.js"></script>
    <script type="text/javascript" src="../_static/petlja.js"></script>
    <script type="text/javascript" src="../_static/require.js"></script>
    <meta charset='utf-8'>
    <meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
    <meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
    <link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />
    <script type="text/javascript">
      eBookConfig = {};
      eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
      eBookConfig.app = eBookConfig.host + '/runestone';
      eBookConfig.ajaxURL = eBookConfig.app + '/ajax/';
      eBookConfig.course = 'Настава програмирања за Други разред гимназије';
      eBookConfig.logLevel = '0';
      eBookConfig.loginRequired = 'false';
      eBookConfig.build_info = "";
      eBookConfig.isLoggedIn = false;
      eBookConfig.useRunestoneServices = false;
      eBookConfig.python3 = true;
      eBookConfig.basecourse = 'Настава програмирања за Други разред гимназије';
      eBookConfig.runestone_version = '';
      eBookConfig.imagesDir = '../_images/';
      eBookConfig.staticDir = '../_static/';
      if (typeof (Sk) != "undefined")
        Sk.imgPath = eBookConfig.imagesDir;
    </script>

  </head>

  <body>



<div id="navbar" class="content-header">
  <div class="right-side-nav">
    
</div>

<div class="left-side-nav">
  <div class="course-logo-nav">
      <a target="_blank" href="https://petlja.org">
          <img src="../_static/img/petlja-logo-wt.svg" alt="petlja.org" id="lightPetljaLogo">
      </a>
  </div>
  <div class="course-logo-nav-kabinet">
      <a target="_blank" href="https://petlja.org/net.kabinet">
          <img src="../_static/img/netkabinet-logo-strelica-wt.svg" alt="petlja.org" id="lightNetKabinetLogo">
      </a>
  </div>
</div>


<div class="content-header-label" id = "content-header-label" href="../">Анализа података у Џупитеру и технике програмирања</div>

</div>




<div class="petljaCourseContainer">

<div class="petljaCourseContent">
  
           <div class="tableOfContent">
            <h5> Садржај </h5>
            </div>
           <div>
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J00>
                    <h5>Рад са подацима у текстуалном програмском језику (ПТПЈ)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J00/J00.html><div class="studioLecture" id=activity-J00><i class="fas fa-file-alt activity-icon"></i>Исходи</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J0A>
                    <h5>Припрема</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J0A/J0A.html><div class="studioLecture" id=activity-J0A><i class="fas fa-file-alt activity-icon"></i>Припрема</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J01>
                    <h5>ПТПЈ.1. Увод у Џупитер (*Jupyter*)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J01/J01.html><div class="studioLecture" id=activity-J01><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J02>
                    <h5>ПТПЈ.2. Низови података и линијски дијаграми</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J02/J02.html><div class="studioLecture" id=activity-J02><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J02/testJ02.html><div class="studioLecture" id=activity-testJ02><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J03>
                    <h5>ПТПЈ.3. Хистограми и боје</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J03/J03.html><div class="studioLecture" id=activity-J03><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J03/testJ03.html><div class="studioLecture" id=activity-testJ03><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J04>
                    <h5>ПТПЈ.4. Просек и медијана низа бројева</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J04/J04.html><div class="studioLecture" id=activity-J04><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J04/testJ04.html><div class="studioLecture" id=activity-testJ04><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J05>
                    <h5>ПТПЈ.5. Фреквенцијска анализа и секторски дијаграми</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J05/J05.html><div class="studioLecture" id=activity-J05><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J05/testJ05.html><div class="studioLecture" id=activity-testJ05><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J06>
                    <h5>ПТПЈ.6. Табеларно представљени подаци</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J06/J06.html><div class="studioLecture" id=activity-J06><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J06/testJ06.html><div class="studioLecture" id=activity-testJ06><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J07>
                    <h5>ПТПЈ.7. Индексирање и транспоновање табеле</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J07/J07.html><div class="studioLecture" id=activity-J07><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J07/testJ07.html><div class="studioLecture" id=activity-testJ07><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J08>
                    <h5>ПТПЈ.8. Модификације табеле и записивање табеле у датотеку</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J08/J08.html><div class="studioLecture" id=activity-J08><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J08/testJ08.html><div class="studioLecture" id=activity-testJ08><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J09>
                    <h5>ПТПЈ.9. Сортирање, филтрирање и фреквенцијска анализа</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J09/J09.html><div class="studioLecture" id=activity-J09><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J09/testJ09.html><div class="studioLecture" id=activity-testJ09><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J10>
                    <h5>ПТПЈ.10. Џупитер и Ексел</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J10/J10.html><div class="studioLecture" id=activity-J10><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J10/testJ10.html><div class="studioLecture" id=activity-testJ10><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y00>
                    <h5>Програмирање (ПРОГ)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y00/Y00.html><div class="studioLecture" id=activity-Y00><i class="fas fa-file-alt activity-icon"></i>Исходи</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y01>
                    <h5>ПРОГ.1. Увод у процесирање низова података у Пајтону</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y01/L801.html><div class="studioLecture" id=activity-L801><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y01/test801.html><div class="studioLecture" id=activity-test801><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y02>
                    <h5>ПРОГ.2. Бројачи и суме</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y02/L802.html><div class="studioLecture" id=activity-L802><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y02/test802.html><div class="studioLecture" id=activity-test802><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y03>
                    <h5>ПРОГ.3. Низови симбола (*стрингови* или *ниске*)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y03/L803.html><div class="studioLecture" id=activity-L803><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y03/test803.html><div class="studioLecture" id=activity-test803><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y04>
                    <h5>ПРОГ.4. Угнежђене петље</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y04/L804.html><div class="studioLecture" id=activity-L804><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y04/test804.html><div class="studioLecture" id=activity-test804><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y05>
                    <h5>ПРОГ.5. Сортирање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y05/L805.html><div class="studioLecture" id=activity-L805><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y05/test805.html><div class="studioLecture" id=activity-test805><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y06>
                    <h5>ПРОГ.6. Филтрирање и претраживање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y06/L806.html><div class="studioLecture" id=activity-L806><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y06/test806.html><div class="studioLecture" id=activity-test806><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y07>
                    <h5>ПРОГ.7. Рад са текстуалним датотекама</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y07/L807.html><div class="studioLecture" id=activity-L807><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y07/test807.html><div class="studioLecture" id=activity-test807><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                </div>
</div>



<div class="lectureContent">

  <div id="errors">
  </div>
  
  <div class="lectureContentMaterial">
  
    <div id="main-content">
    
  <div class="section" id="id1">
<h1>Низови симбола (<em>стрингови</em> или <em>ниске</em>)<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>У овој лекцији ћемо:</p>
<ol class="arabic simple">
<li><p>упознати стрингове као низове карактера</p></li>
<li><p>видети како се приступа елементима стринга и како се издвајају делови стринга, и</p></li>
<li><p>упознати основне операције са стринговима.</p></li>
</ol>
<div class="section" id="id2">
<h2>Стрингови (ниске)<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>
<p>Иако су настали као машине за обраду нумеричких података рачунари су доживели врхунас свог успеха
када се схватило да они могу да обрађују и <em>симболичке</em> податке. Обрада симболичких података данас
представља доминантан начин употребе рачунара.</p>
<p><em>Стринг</em> је реч књижевног енглеског језика (енгл. <em>string</em>) чије основно значење је
„објекти исте врсте нанизани на нит”. Зато се стрингови често на српском зову и <em>ниске</em>.
У Пајтону стринг је низ симбола. На пример, ово су стрингови:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="s2">&quot;Zdravo! Kako si?&quot;</span>
<span class="s1">&#39;Основна школа &quot;Јован Јовановић-Змај&quot;&#39;</span>
<span class="s2">&quot;:-)&quot;</span>
<span class="s1">&#39;&#39;</span>
<span class="s2">&quot;&quot;</span>
</pre></div>
</div>
<p>Видимо да се стрингови пишу тако што се између апострофа или наводника (свеједно је) наведе произвољан низ симбола.
Последња два примера нам показују како изгледа <em>празан стринг</em>, односно, стринг који нема ниједно слово.
Други пример нам показује како да се снађемо ако стринг треба да садржи наводнике: просто га оградимо апострофима.</p>
<p>Број симбола у стрингу се зове <em>дужина стринга</em>. Дужине стрингова из претходног примера су, редом, 16, 36, 3, 0 и 0.</p>
<p>Стрингове додељујемо променљивим као што смо и навикли, исписујемо их командом <cite>print</cite>, а учитавамо командом <cite>input</cite>.
Команда <cite>input</cite> сачека да корисник унесе низ симбола и да на крају притисне [ENTER]. Она потом врати све симболе
које је корисник унео као један низ симбола (такорећи <em>стринг</em>). На пример:</p>

<div data-childcomponent="primer3-1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="print(&quot;Unesi svoje ime, prvo u nominativu pa onda u vokativu&quot;)
ime = input(&quot;nominativ: &quot;)
ime_v = input(&quot;vokativ: &quot;)
print(&quot;Aha! Znao sam!&quot;)
print(&quot;Ti se zoveš&quot;, ime)
print(&quot;Zdravo,&quot;, ime_v)
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P03</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Vokativ.py</span></code> и изврши га.</p>
<p>Ево и кратке видео демонстрације:</p>

    <div  id="EhYScc6wZOo" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/EhYScc6wZOo/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-EhYScc6wZOo" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-EhYScc6wZOo">
        </div>
    </div>
    
    
    
    </p></div></div>
<p>Уграђена функција <code class="docutils literal notranslate"><span class="pre">len</span></code> враћа дужину стринга и користи се овако:</p>

<div data-childcomponent="primer3-2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
ime = input("Unesi ime:")
print("Ti se zoveš", ime, "i tvoje ime ima", len(ime), "slova")
</textarea>
</div>
<div class="section" id="id3">
<h3>Пример.<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон програм који утврђује колико пута се неки симбол јавља у унетом стрингу.</p>

<div data-childcomponent="primer3-3" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-3 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="print(&quot;Utvrdjujemo koliko puta se neki simbol javlja u unetom stringu&quot;)
c = input(&quot;Unesi simbol: &quot;)
s = input(&quot;Unesi string: &quot;)
broj = 0
for x in s:
    if x == c:
        broj += 1
print(&quot;Simbol&quot;, c, &quot;se u unetom stringu javlja&quot;, broj, &quot;puta&quot;)
"   >


</textarea>
</div>
<p>Конструкција <code class="docutils literal notranslate"><span class="pre">for</span> <span class="pre">x</span> <span class="pre">in</span> <span class="pre">s</span></code> прође кроз цео стринг <code class="docutils literal notranslate"><span class="pre">s</span></code> слово по слово, па за свако појављивање симбола <code class="docutils literal notranslate"><span class="pre">c</span></code> увећамо бројач за 1.</p>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P03</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Broj_simbola.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
<div class="section" id="id4">
<h3>Пример.<a class="headerlink" href="#id4" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон програм који утврђује да ли је у датом стрингу знакова, <code class="docutils literal notranslate"><span class="pre">&lt;</span></code>, једнак броју знакова, <code class="docutils literal notranslate"><span class="pre">&gt;</span></code>.</p>

<div data-childcomponent="primer3-4" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-4 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="s = input(&quot;Unesi string&quot;)
br_jemanje = 0
br_jevece = 0
for x in s:
    if x == &quot;&lt;&quot;: br_jemanje += 1
    if x == &quot;&gt;&quot;: br_jevece += 1
if br_jemanje == br_jevece:
    print(&quot;isti je broj znakova &lt; и &gt;&quot;)
else:
    print(&quot;nije isti broj znakova &lt; и &gt;&quot;)
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P03</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Zagrade.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
</div>
<div class="section" id="id5">
<h2>Приступање елементима стринга и деловима стринга<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h2>
<p>Пошто је стринг низ симбола могуће је приступити појединачним симболима у стрингу, али и деловима стринга.</p>
<a class="reference internal image-reference" href="../_images/Y103.jpg"><img alt="../_images/Y103.jpg" src="../_images/Y103.jpg" style="width: 800px;" /></a>
<p>Стринг је, дакле, низ кућица. У сваку кућицу је уписан тачно један симбол и свака кућица има кућни број.
Као и раније, кућни бројеви крећу од нуле. На пример, ако ставимо</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;Popokatepetl&quot;</span>
</pre></div>
</div>
<p>онда имамо:</p>

<div data-childcomponent="primer3-5" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-5 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
s = "Popokatepetl"
print(s[4])
print(s[0])
</textarea>
</div>
<p>Инересантно је (а касније ћемо видети да је и корисно!) то што низови у Пајтону имају и негативне индексе.
Последња кућица у низу има индекс -1 („први од краја”), претпоследња -2 („други од краја”) и тако даље. Зато је:</p>

<div data-childcomponent="primer3-6" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-6 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
s = "Popokatepetl"
print(s[-1])
print(s[-5])
</textarea>
</div>
<p>Из стринга можемо лако да издвојимо његове делове употребом конструкције <code class="docutils literal notranslate"><span class="pre">s[m:n]</span></code>
која издваја део стринга који почиње од кућице са редним бројем <em>m</em> и завршава се са кућицом чији редни број је <em>n - 1</em>
(обрати пажњу: кренемо од <em>m</em>, али <strong>не идемо</strong> до <em>n</em> већ до кућице пре). На ово ћемо морати мало да се привикнемо. На пример,</p>

<div data-childcomponent="primer3-7" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-7 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
s = "Popokatepetl"
print(s[2:5])
</textarea>
</div>
<p>Дакле, из стринга <em>s</em> смо издвојили симболе <em>s[2]</em>, <em>s[3]</em> и <em>s[4]</em>, односно, кренули смо од кућице број 2 и завршили
<em>непосредно пре</em> кућице број 5.</p>
<p>Ако желимо да издвојимо део стринга који креће од 6. кућице и иде до краја стринга, написаћемо <em>s[6:]</em></p>

<div data-childcomponent="primer3-8" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-8 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
s = "Popokatepetl"
print(s[6:])
</textarea>
</div>
<p>С друге стране, ако желимо да издвојимо део стринга који креће од почетка стринга и иде до 6. кућице (укљућујући и њу) написаћемо
<em>s[:7]</em></p>

<div data-childcomponent="primer3-9" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-9 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
s = "Popokatepetl"
print(s[:7])
</textarea>
</div>
<p><em>Не смемо заборавити да се кућица чији кућни број је наведен иза двотачке не укључује у сегмент!</em></p>
<p>Издвајање делова стринга у комбинацији са негативним индексима нам омогућује да концизно запишемо веома корисне конструкције.
Рецимо, последња четири елемента стринга добијамо овако:</p>

<div data-childcomponent="primer3-10" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-10 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
s = "Popokatepetl"
print(s[-4:])
</textarea>
</div>
<p>што је концизан запис за део стринга „од -4. кућице па до краја”. С друге стране, „све осим последња два” можемо да запишемо овако:</p>

<div data-childcomponent="primer3-11" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-11 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
s = "Popokatepetl"
print(s[:-2])
</textarea>
</div>
<div class="section" id="id6">
<h3>Пример.<a class="headerlink" href="#id6" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">prebroj_podstr(p,</span> <span class="pre">s)</span></code> која утврђује колико пута се стринг <code class="docutils literal notranslate"><span class="pre">p</span></code> јавља као подстринг стринга <code class="docutils literal notranslate"><span class="pre">s</span></code>.</p>

<div data-childcomponent="primer3-12" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-12 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="def prebroj_podstr(p, s):
    broj = 0
    n = len(p)
    for i in range(len(s) - n + 1):
        if p == s[i:i+n]:
            broj += 1
    return broj

# провера
print(prebroj_podstr(&quot;bana&quot;, &quot;oprobana torta od banana&quot;))
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P03</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Prebroj_podstr.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
<div class="section" id="id7">
<h3>Пример.<a class="headerlink" href="#id7" title="Permalink to this headline">¶</a></h3>
<p>У сваком рачунарском систему се подаци пакују у <em>фајлове</em> (или <em>датотеке</em>) и од самих почетака рачунарске ере постоји обичај
да се име фајла организује тако да последњих неколико симбола у имену фајла означава тип података који је у том фајлу похрањен.
Тај део имена фајла се обично одвоји тачком од остатка имена фајла и зове се <em>екстензија</em>. Рецимо:</p>
<table class="docutils align-left" id="id21">
<caption><span class="caption-text">Екстензије и типови података</span><a class="headerlink" href="#id21" title="Permalink to this table">¶</a></caption>
<colgroup>
<col style="width: 50%" />
<col style="width: 50%" />
</colgroup>
<thead>
<tr class="row-odd"><th class="head"><p>Екстензија</p></th>
<th class="head"><p>Тип података</p></th>
</tr>
</thead>
<tbody>
<tr class="row-even"><td><p>.txt</p></td>
<td><p>текстуална датотека</p></td>
</tr>
<tr class="row-odd"><td><p>.log</p></td>
<td><p>текстуална датотека</p></td>
</tr>
<tr class="row-even"><td><p>.png</p></td>
<td><p>слика</p></td>
</tr>
<tr class="row-odd"><td><p>.jpg</p></td>
<td><p>слика</p></td>
</tr>
<tr class="row-even"><td><p>.bmp</p></td>
<td><p>слика</p></td>
</tr>
<tr class="row-odd"><td><p>.mpg</p></td>
<td><p>видео</p></td>
</tr>
<tr class="row-even"><td><p>.mp3</p></td>
<td><p>звук</p></td>
</tr>
</tbody>
</table>
<p>Написати Пајтон програм који од корисника учитава име фајла и на основу горње табеле одређује тип
података који је у њему похрањен. У случају да се екстензија имена фајла не налази у овој табели
програм треба да пријави да се ради о непознатом типу података.</p>

<div data-childcomponent="primer3-13" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-13 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="fajl = input(&quot;Unesi ime fajla: &quot;)
ekstenzija = fajl[-4:] # uzimamo poslednja 4 simbola imena
if ekstenzija in [&quot;.txt&quot;, &quot;.log&quot;]:
    print(&quot;tekstualni fajl&quot;)
elif ekstenzija in [&quot;.png&quot;, &quot;.jpg&quot;, &quot;.bmp&quot;]:
    print(&quot;slika&quot;)
elif ekstenzija == &quot;.mpg&quot;:
    print(&quot;video&quot;)
elif ekstenzija == &quot;.mp3&quot;:
    print(&quot;zvuk&quot;)
else:
    print(&quot;Ovaj format mi nije poznat&quot;)
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P03</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Ekstenzije.py</span></code> и изврши га.</p>
<p>Ево и кратке видео демонстрације:</p>

    <div  id="KgoFGIL_gVI" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/KgoFGIL_gVI/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-KgoFGIL_gVI" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-KgoFGIL_gVI">
        </div>
    </div>
    
    
    
    </p></div></div>
</div>
<div class="section" id="id8">
<h3>Пример.<a class="headerlink" href="#id8" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">broj_reci(s)</span></code> која у датом стрингу броји речи. Реч је непрекидан низ слова.</p>

<div data-childcomponent="primer3-14" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-14 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="def broj_reci(s):
    n = len(s)
    broj = 0
    for i in range(n-1):
        if s[i].isalpha() and not s[i+1].isalpha():
            broj += 1
    if s[n-1].isalpha():
        broj += 1
    return broj

# провера
print(broj_reci(&quot;Ovo je samo primer!&quot;))
print(broj_reci(&quot;Proba&quot;))
print(broj_reci(&quot;!!!!&quot;))
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P03</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Broj_reci.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
</div>
<div class="section" id="id9">
<h2>Операције са стринговима<a class="headerlink" href="#id9" title="Permalink to this headline">¶</a></h2>
<p>У Пајтону, као и у већини модерних програмских језика, стрингови су <em>непроменљиви</em> (то својство се на енглеском зове
<em>immutability</em>). На пример, нека је</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">s</span> <span class="o">=</span> <span class="s2">&quot;Banana&quot;</span>
</pre></div>
</div>
<p>Ако покушамо да променимо прво слово стринга на мало слово ’b’ добићемо грешку!</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">s</span><span class="p">[</span><span class="mi">0</span><span class="p">]</span> <span class="o">=</span> <span class="s2">&quot;b&quot;</span>
<span class="c1"># ОВО НЕ МОЖЕ! СИСТЕМ ЋЕ ПРИЈАВИТИ ГРЕШКУ!</span>
</pre></div>
</div>
<p>Зато се трансформације стрингова увек своде на то да се од старог стринга изгради нови стринг који има својства која желимо.
Једна од основних операција коју при томе користимо је <em>надовезивање стрингова</em> или <em>конкатенација</em>.
Операција надовезивања стрингова се означава знаком <cite>+</cite>. На пример:</p>

<div data-childcomponent="primer3-15" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-15 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
s = "Novi"
t = "Beograd"
print(s + t)
</textarea>
</div>
<p>Важно је напоменути да се стрингови надовезују на најједноставнији могући начин: просто се дода један стринг на крај
оног другог, без уметања празнина или неке друге врсте мудровања.</p>
<p>Користећи операцију надовезивања можемо стрингу додати неки други стринг на почетак или на крај. На пример:</p>

<div data-childcomponent="primer3-16" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-16 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
s = "bana"
s = "opro" + s
print(s)

t = "bana"
t = t + "na"
print(t)
</textarea>
</div>
<p>Још једна интересантна операција је множење стринга бројем:</p>

<div data-childcomponent="primer3-17" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-17 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print("Mir! " * 3)
</textarea>
</div>
<div class="section" id="id10">
<h3>Пример.<a class="headerlink" href="#id10" title="Permalink to this headline">¶</a></h3>
<p>Ђоле треба 100 пута (за казну :-)) да испише реченицу „Ученици не смеју да се гађају кредом!”. Помозите Ђолету!</p>

<div data-childcomponent="primer3-18" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-18 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print("Ученици не смеју да се гађају кредом! " * 100)
</textarea>
</div>
</div>
<div class="section" id="id11">
<h3>Пример.<a class="headerlink" href="#id11" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију која обрће стринг.</p>

<div data-childcomponent="primer3-19" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-19 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="def obrni_string(s):
    novi = &quot;&quot;
    for x in s:
        novi = x + novi
    return novi

# провера
print(obrni_string(&quot;Proba&quot;))
print(obrni_string(&quot;ana voli milovana&quot;))
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P03</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Obrni_string.py</span></code> и изврши га.</p>

    </p></div></div>
<p>Ова функција од стринга <code class="docutils literal notranslate"><span class="pre">s</span></code> у променљивој <code class="docutils literal notranslate"><span class="pre">novi</span></code> изгради нови стринг који садржи иста слова као и стринг <code class="docutils literal notranslate"><span class="pre">s</span></code>,
али у обрнутом редоследу. На почетку поставимо стринг <code class="docutils literal notranslate"><span class="pre">novi</span></code> на празан стринг. Потом у циклусу слова стринга
<code class="docutils literal notranslate"><span class="pre">s</span></code> додајемо на стринг <code class="docutils literal notranslate"><span class="pre">novi</span></code>, али једно по једно на почетак стринга <code class="docutils literal notranslate"><span class="pre">novi</span></code>.
Тако прво слово стринга <code class="docutils literal notranslate"><span class="pre">s</span></code> постаје последње слово стринга <code class="docutils literal notranslate"><span class="pre">novi</span></code>, док последње слово стринга <code class="docutils literal notranslate"><span class="pre">s</span></code>
на крају буде прво слово стринга <code class="docutils literal notranslate"><span class="pre">novi</span></code>. Тако у стрингу <code class="docutils literal notranslate"><span class="pre">novi</span></code> изградимо слово по слово обрнуту верзију
стринга <code class="docutils literal notranslate"><span class="pre">s</span></code>.</p>
</div>
<div class="section" id="id12">
<h3>Пример.<a class="headerlink" href="#id12" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон програм који проверава да ли је дата реченица палиндром.
Приликом провере треба игнорисати празинине и величину слова. На пример, следеће реченице су палиндроми:</p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span>Ана воли Милована!
Сир има мирис?
Сава зидар ради за вас
Маја са Недом оде на сајам.
</pre></div>
</div>
<p><em>Решење.</em> У решењу овог примера ћемо користити уграђену функцију <code class="docutils literal notranslate"><span class="pre">isaplha()</span></code>
која проверава да ли је нешто слово, као и уграђену функцију <code class="docutils literal notranslate"><span class="pre">lower()</span></code> која од стринга прави нови у коме су сва слова мала.</p>

<div data-childcomponent="primer3-20" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer3-20 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="s = input(&quot;Unesi string: &quot;)
novi = &quot;&quot;
obrnut = &quot;&quot;
for x in s.lower():
    if x.isalpha():
        novi = novi + x
        obrnut = x + obrnut
if novi == obrnut:
    print(&quot;jeste palindrom&quot;)
else:
    print(&quot;nije palindrom&quot;)
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P03</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Palindrom.py</span></code> и изврши га.</p>
<p>Ево и кратке видео демонстрације:</p>

    <div  id="17srK5-x7NI" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/17srK5-x7NI/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-17srK5-x7NI" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-17srK5-x7NI">
        </div>
    </div>
    
    
    
    </p></div></div>
</div>
</div>
<div class="section" id="id13">
<h2>Задаци<a class="headerlink" href="#id13" title="Permalink to this headline">¶</a></h2>
<div class="section" id="id14">
<h3>Задатак 1.<a class="headerlink" href="#id14" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">broj_samoglasnika(s)</span></code> која
утврђује број самогласника у стрингу <code class="docutils literal notranslate"><span class="pre">s</span></code>. Претпостављамо да <code class="docutils literal notranslate"><span class="pre">s</span></code> садржи реченицу
српског језика, па треба пребројати само самогласнике који се јављају у српском (а, е, и, о, у),
<em>независно од тога да ли је реченица написана ћирилицом или латиницом!</em>
(Водити рачуна о великим и малим словима!)</p>

<div data-childcomponent="zadatak3-1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak3-1 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3 test4"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def broj_samoglasnika(s):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = broj_samoglasnika("Ovde ima mnogo samoglasnika")
test2 = broj_samoglasnika("Led je čvrst, krt i tvrd.")
test3 = broj_samoglasnika("Čvrst, krt, tvrd zvrk!")
test4 = broj_samoglasnika("")
# -*- acsection: after-main -*-
print(test1, test2, test3, test4)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        def __samogl(c):
            return c.lower() in ["a", "e", "i", "o", "u", "а", "е", "и", "о", "у"]
        def __brsam(s):
            n=0
            for c in s:
                if __samogl(c): n += 1
            return n
        rez1 = __brsam("Ovde ima mnogo samoglasnika")
        rez2 = __brsam("Led je čvrst, krt i tvrd.")
        rez3 = __brsam("Čvrst, krt, tvrd zvrk!")
        rez4 = __brsam("")
        run_test = acMainSection(test1=test1,test2=test2,test3=test3,test4=test4)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
        self.assertEqual(run_test["test4"], rez4, "Вредност променљиве 'test4' треба да буде %s" % rez4)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id15">
<h3>Задатак 2.<a class="headerlink" href="#id15" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">broj_suglasnika(s)</span></code> која
утврђује број сугласника у стрингу <code class="docutils literal notranslate"><span class="pre">s</span></code>. Претпостављамо да <code class="docutils literal notranslate"><span class="pre">s</span></code> садржи реченицу
српског језика, па је сугласник свако слово које је различито од а, е, и, о, у
<em>независно од тога да ли је реченица написана ћирилицом или латиницом!</em>
(Водити рачуна о великим и малим словима!)</p>

<div data-childcomponent="zadatak3-1b" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak3-1b data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3 test4"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def broj_suglasnika(s):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = broj_suglasnika("Ovde ima mnogo samoglasnika")
test2 = broj_suglasnika("Led je čvrst, krt i tvrd.")
test3 = broj_suglasnika("Čvrst, krt, tvrd zvrk!")
test4 = broj_suglasnika("")
# -*- acsection: after-main -*-
print(test1, test2, test3, test4)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        def __samogl(c):
            return c.lower() in ["a", "e", "i", "o", "u", "а", "е", "и", "о", "у"]
        def __sugl(c):
            return c.isalpha() and not __samogl(c)
        def __brsug(s):
            n=0
            for c in s:
                if __sugl(c): n += 1
            return n
        rez1 = __brsug("Ovde ima mnogo samoglasnika")
        rez2 = __brsug("Led je čvrst, krt i tvrd.")
        rez3 = __brsug("Čvrst, krt, tvrd zvrk!")
        rez4 = __brsug("")
        run_test = acMainSection(test1=test1,test2=test2,test3=test3,test4=test4)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
        self.assertEqual(run_test["test4"], rez4, "Вредност променљиве 'test4' треба да буде %s" % rez4)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id16">
<h3>Задатак 3.<a class="headerlink" href="#id16" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">poslednja_rec(s)</span></code> која из датог стринга издваја последњу реч која је
наведена у стрингу. Ако у стрингу не постоји ниједно слово функција треба да врати празан стринг.
(Напомена: реч је непрекидан низ слова.)</p>

<div data-childcomponent="zadatak3-2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak3-2 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3 test4"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def poslednja_rec(s):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = poslednja_rec("Ovaj string ima pet reci")
test2 = poslednja_rec("Ovaj string se zavrsava uzvicnicima!!!")
test3 = poslednja_rec("Test")
test4 = poslednja_rec("1 2 3 4 5")
# -*- acsection: after-main -*-
print(test1, test2, test3, test4)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        rez1 = "reci"
        rez2 = "uzvicnicima"
        rez3 = "Test"
        rez4 = ""
        run_test = acMainSection(test1=test1,test2=test2,test3=test3,test4=test4)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде 'reci'")
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде 'uzvicnicima'")
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде 'Test'")
        self.assertEqual(run_test["test4"], rez4, "Вредност променљиве 'test4' треба да буде празан стринг")
myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="pomoc5"   data-showtitle="Помоћ" data-hidetitle="Сакриј помоћ">
    
    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Не заборави да кроз итерабле (листе, ниске, торке…) можеш пролазити и уназад помоћу <code class="docutils literal notranslate"><span class="pre">for</span></code> петље тако што ћеш у оквиру функције range() искористити аргументе који дефинишу које вредности ће твој бројач узимати. Први аргумент ове функције је почетна вредност бројача, други аргумент је завршна вредност, а трећи аргумент је корак који ћемо направити. Ако је почетна вредност већа од завршне, а корак нам је негативан, наш бројач ће се смањивати како пролазимо кроз листу. Погледај следећи пример:</p>

    </p></div></div>

<div data-childcomponent="range" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=range data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
s = 'String'                      # пошто се броји од нуле, почињемо бројем мање од дужине стринга да не бисмо испали из граница,
for i in range(len(s)-1, -1, -1): # завршавамо са -1 како би нула била последњи број до ког стижемо
   print(s[i])                    # у сваком проласку кроз петљу смањујемо i за 1
</textarea>
</div>

    </div>
    </div>
<div class="section" id="id17">
<h3>Задатак 4*.<a class="headerlink" href="#id17" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">obrni_reci(s)</span></code> која од датог стринга прави нови у коме су речи
послагане од последње према првој. На пример, за стринг</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="s2">&quot;Napisati program koji ispisuje reci unetog teksta u obrnutom poretku&quot;</span>
</pre></div>
</div>
<p>функција треба да врати:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="s2">&quot;poretku obrnutom u teksta unetog reci ispisuje koji program Napisati&quot;</span>
</pre></div>
</div>

<div data-childcomponent="zadatak3-3" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak3-3 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-

def obrni_reci(s):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = obrni_reci("Dve reci")
test2 = obrni_reci("Ovaj string se ne zavrsava uzvicnikom")
test3 = obrni_reci("Test")
# -*- acsection: after-main -*-

print(test1)
print(test2)
print(test3)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        def __obrnir(s):
           s += " "
           t = ""
           w = ""
           for i in range(len(s)-1):
               if s[i].isalpha():
                   w += s[i]
                   if not(s[i+1].isalpha()):
                       if t == "":
                           t = w
                       else:
                           t = w + " " + t
                       w = ""
           return t
        rez1 = __obrnir("Dve reci")
        rez2 = __obrnir("Ovaj string se ne zavrsava uzvicnikom")
        rez3 = __obrnir("Test")
        run_test = acMainSection(test1=test1,test2=test2,test3=test3)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде '%s'" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде '%s'" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде '%s'" % rez3)
myTests().main()
</textarea>
</div>
</div>
<div class="section" id="id18">
<h3>Задатак 5*.<a class="headerlink" href="#id18" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију која проверава да ли су у датом стрингу изломљене заграде <code class="docutils literal notranslate"><span class="pre">&lt;</span></code> и <code class="docutils literal notranslate"><span class="pre">&gt;</span></code> <em>балансиране</em>.
То значи да:</p>
<ol class="arabic simple">
<li><p>у стрингу имамо исти број отворених и затоврених изломљених заграда, и</p></li>
<li><p>не постоји почетни сегмент стринга у коме је број затворених заграда већи од броја отворених.</p></li>
</ol>

<div data-childcomponent="zadatak3-4" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak3-4 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3 test4 test5"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def balansirane_zagrade(s):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = balansirane_zagrade("<>") # јесу
test2 = balansirane_zagrade("<<>>") # јесу
test3 = balansirane_zagrade("<<>><<><<><>>>") # јесу
test4 = balansirane_zagrade("><") # нису
test5 = balansirane_zagrade("<<>>><") # нису
# -*- acsection: after-main -*-
print(test1, test2, test3, test4, test5)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        rez1 = True
        rez2 = True
        rez3 = True
        rez4 = False
        rez5 = False
        run_test = acMainSection(test1=test1,test2=test2,test3=test3,test4=test4,test5=test5)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
        self.assertEqual(run_test["test4"], rez4, "Вредност променљиве 'test4' треба да буде %s" % rez4)
        self.assertEqual(run_test["test5"], rez5, "Вредност променљиве 'test5' треба да буде %s" % rez5)
myTests().main()
</textarea>
</div>

    <div data-component="reveal" id="pomoc2"   data-showtitle="Помоћ" data-hidetitle="Сакриј помоћ">
    
    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Не заборави да у оквиру функције можеш више пута да искористиш <code class="docutils literal notranslate"><span class="pre">return</span></code></p>

    </p></div></div>

    </div>
    </div>
<div class="section" id="id19">
<h3>Задатак 6.<a class="headerlink" href="#id19" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">razmeni(s,</span> <span class="pre">din_za_1_eur)</span></code> која конвертује дати износ из динара у евре и обрнуто.
Стринг <code class="docutils literal notranslate"><span class="pre">s</span></code> садржи износ записан у облику <code class="docutils literal notranslate"><span class="pre">број</span> <span class="pre">валута</span></code> на пример овако:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="s2">&quot;125 DIN&quot;</span>
<span class="s2">&quot;217.50 EUR&quot;</span>
<span class="s2">&quot;1527.99 DIN&quot;</span>
<span class="s2">&quot;50 EUR&quot;</span>
</pre></div>
</div>
<p>док аргумент <code class="docutils literal notranslate"><span class="pre">din_za_1_eur</span></code> представља вредност једног евра у динарима.
Функција треба да врати стринг који има облик <code class="docutils literal notranslate"><span class="pre">број</span> <span class="pre">валута</span></code>. При томе користи бити уграђене функције <code class="docutils literal notranslate"><span class="pre">float</span></code> и <code class="docutils literal notranslate"><span class="pre">str</span></code>.
(За неки стринг <code class="docutils literal notranslate"><span class="pre">p</span></code> функција <code class="docutils literal notranslate"><span class="pre">float(p)</span></code> враћа број који је записан стрингом <code class="docutils literal notranslate"><span class="pre">p</span></code> игноришући при том празнине које се
можда јављају у стрингу; с друге стране, функција <code class="docutils literal notranslate"><span class="pre">str</span></code> од неког броја прави стринг који садржи његов запис.)</p>
<p><em>Добијене износе НЕ ТРЕБА заокруживати на две децимале!</em></p>
<p>Функција <code class="docutils literal notranslate"><span class="pre">razmeni</span></code> треба да ради овако:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="gp">&gt;&gt;&gt; </span><span class="n">razmeni</span><span class="p">(</span><span class="s2">&quot;124.99 EUR&quot;</span><span class="p">,</span> <span class="mf">117.90</span><span class="p">)</span>
<span class="go">&#39;14736.321‬ DIN&#39;</span>

<span class="gp">&gt;&gt;&gt; </span><span class="n">razmeni</span><span class="p">(</span><span class="s2">&quot;12000 DIN&quot;</span><span class="p">,</span> <span class="mf">117.90</span><span class="p">)</span>
<span class="go">‭&#39;101.7811704 EUR&#39;</span>
</pre></div>
</div>
<p>Ако корисник унесе неку трећу валуту, програм треба да врати стринг „?VALUTA”:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="gp">&gt;&gt;&gt; </span><span class="n">razmeni</span><span class="p">(</span><span class="s2">&quot;124.99 USD&quot;</span><span class="p">,</span> <span class="mf">117.90</span><span class="p">)</span>
<span class="go">&#39;?VALUTA&#39;</span>
</pre></div>
</div>

<div data-childcomponent="zadatak3-5" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=zadatak3-5 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3 test4"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-

def razmeni(s, din_za_1_eur):
    # Овде напиши функцију
    return -1234  # поправи овај ред!

# Провера
test1 = razmeni("1 EUR", 117.90)
test2 = razmeni("124.99 EUR", 117.90)
test3 = razmeni("12000 DIN", 117.90)
test4 = razmeni("1 USD", 107.90)
# -*- acsection: after-main -*-

print(test1)
print(test2)
print(test3)
print(test4)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        def __razmeni(s, dz1e):
            x = s.split()
            iznos = float(x[0])
            valuta = x[1]
            if valuta.upper() == "DIN": return str(iznos/dz1e) + " EUR"
            if valuta.upper() == "EUR": return str(iznos * dz1e) + " DIN"
            return "?VALUTA"
        rez1 = __razmeni("1 EUR", 117.90)
        rez2 = __razmeni("124.99 EUR", 117.90)
        rez3 = __razmeni("12000 DIN", 117.90)
        rez4 = __razmeni("1 USD", 107.90)
        run_test = acMainSection(test1=test1,test2=test2,test3=test3,test4=test4)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде '%s'" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде '%s'" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде '%s'" % rez3)
        self.assertEqual(run_test["test4"], rez4, "Вредност променљиве 'test4' треба да буде '%s'" % rez4)
myTests().main()
</textarea>
</div>
</div>
</div>
<div class="section" id="id20">
<h2>Задаци из Алгоритамске збирке задатака<a class="headerlink" href="#id20" title="Permalink to this headline">¶</a></h2>
<p>Решавањем следећих задатака можете да увежбате додатне вештине у раду са стринговима.</p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/prezime_pa_ime">https://petlja.org/biblioteka/r/problemi/Zbirka/prezime_pa_ime</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/da_li_je_podniska">https://petlja.org/biblioteka/r/problemi/Zbirka/da_li_je_podniska</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/izbacivanje_podniski">https://petlja.org/biblioteka/r/problemi/Zbirka/izbacivanje_podniski</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/rec_frankenstajn">https://petlja.org/biblioteka/r/problemi/Zbirka/rec_frankenstajn</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/cezarov_kod">https://petlja.org/biblioteka/r/problemi/Zbirka/cezarov_kod</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/vrednost_izraza">https://petlja.org/biblioteka/r/problemi/Zbirka/vrednost_izraza</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/spreadsheet_kolone">https://petlja.org/biblioteka/r/problemi/Zbirka/spreadsheet_kolone</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/brojevne_osnove">https://petlja.org/biblioteka/r/problemi/Zbirka/brojevne_osnove</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/arapski_u_rimski">https://petlja.org/biblioteka/r/problemi/Zbirka/arapski_u_rimski</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/rimski_u_arapski">https://petlja.org/biblioteka/r/problemi/Zbirka/rimski_u_arapski</a></p>
</div>
</div>


    </div>
    
    <div>
<div class="lecture-prev-next-toggle">
  <a  href="index.html" id="prevLectureLink" class="lecture-toggler-arrows"><i class="fas fa-arrow-left fa"></i> <span> Претходно лекција </span> </a>
  
  <a id="nextLectureLink" class="lecture-toggler-arrows"><span> Следећа лекција</span>  <i class="fas fa-arrow-right fa"></i> </a>
</div>

<script type="text/javascript">
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
</div>

    
  </div>
</div>
<div class="petljadoc-resources">
  <span class="pull-left">(Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>) <br> © 2022 Petlja</span>
</div>

<div class="lectureToolbox">

  <label class="font-size-toggle" id="fontSizeToggle">

      <i class="fas fa-font"></i><label>Слова</label>
  </label>
  <div class="font-change d-none" id="fontChangeModal">
      <span class="dec-letter-icon" id="decLetterIcon">A-</span>
      <label id="currentFontSize">16</label>
      <span class="inc-letter-icon" id="incLetterIcon">A+</span>
  </div>
</div>

</div>

  </body>
</html>