<!DOCTYPE html>
<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Подсетник - Пајтон</title>
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/qchoice.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" type="text/css" href="../_static/gallery.css" />
    <link rel="stylesheet" type="text/css" href="../_static/associations.css" />
    <link rel="stylesheet" type="text/css" href="../_static/editor.css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="https://fonts.googleapis.com/css?family=Source+Sans+Pro" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja-course.css" type="text/css" />
    <link rel="stylesheet" href="../_static/nbstyle.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/language_data.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/karelCorner.js"></script>
    <script type="text/javascript" src="../_static/karelRobot.js"></script>
    <script type="text/javascript" src="../_static/karelWorld.js"></script>
    <script type="text/javascript" src="../_static/karelChat.js"></script>
    <script type="text/javascript" src="../_static/karelRobotDrawer.js"></script>
    <script type="text/javascript" src="../_static/karelUI.js"></script>
    <script type="text/javascript" src="../_static/karel.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/notes.js"></script>
    <script type="text/javascript" src="../_static/pygamelib-init.js"></script>
    <script type="text/javascript" src="../_static/gallery.js"></script>
    <script type="text/javascript" src="../_static/associations.js"></script>
    <script type="text/javascript" src="../_static/associations-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/associations-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/editor.js"></script>
    <script type="text/javascript" src="../_static/jszip.js"></script>
    <script type="text/javascript" src="../_static/editor-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/editor-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/translations.js"></script>
    <script async="async" type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="../_static/course-errors.js"></script>
    <script type="text/javascript" src="../_static/petlja.js"></script>
    <script type="text/javascript" src="../_static/require.js"></script>
    <meta charset='utf-8'>
    <meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
    <meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
    <link rel="shortcut icon" href="../_static/favicon.ico" type="image/ico" />
    <script type="text/javascript">
      eBookConfig = {};
      eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
      eBookConfig.app = eBookConfig.host + '/runestone';
      eBookConfig.ajaxURL = eBookConfig.app + '/ajax/';
      eBookConfig.course = 'Настава програмирања за Други разред гимназије';
      eBookConfig.logLevel = '0';
      eBookConfig.loginRequired = 'false';
      eBookConfig.build_info = "";
      eBookConfig.isLoggedIn = false;
      eBookConfig.useRunestoneServices = false;
      eBookConfig.python3 = true;
      eBookConfig.basecourse = 'Настава програмирања за Други разред гимназије';
      eBookConfig.runestone_version = '';
      eBookConfig.imagesDir = '../_images/';
      eBookConfig.staticDir = '../_static/';
      if (typeof (Sk) != "undefined")
        Sk.imgPath = eBookConfig.imagesDir;
    </script>

  </head>

  <body>



<div id="navbar" class="content-header">
  <div class="right-side-nav">
    
    <div id="usefulLinksToggle" class="useful-links-toggle">
        <a class="dropdown-toggle dropdown-toggle-burger" role="button" href="#" id="usefulLinksDropdown" data-toggle="dropdown" aria-haspopup="true" aria-expanded="false">
            <i class="fas fa-bars fa-2x"></i>
        </a>
        <div class="dropdown-menu dropdown-nav-links shadow-lg" aria-labelledby="usefulLinksDropdown" x-placement="bottom-start" style="position: absolute; transform: translate3d(-55px, 32px, 0px); top: 0px; left: 0px; will-change: transform;">
          <div class="links-title"> Корисни линкови</div>
          <div class="useful-link">
            <a href="https://www.sphinx-doc.org/en/master/usage/restructuredtext/index.html" target="_blank">Sphinx Docs</a></div>
          <div class="useful-link">
            <a href="https://github.com/Petlja/PetljaDoc/blob/master/README.md" target="_blank">Petljadoc</a></div>
          <div class="useful-link">
            <a href="https://github.com/RunestoneInteractive/overview" target="_blank">Runestone Interactive</a></div>
          
        </div>
    </div>

    
</div>

<div class="left-side-nav">
  <div class="course-logo-nav">
      <a target="_blank" href="https://petlja.org">
          <img src="../_static/img/petlja-logo-wt.svg" alt="petlja.org" id="lightPetljaLogo">
      </a>
  </div>
  <div class="course-logo-nav-kabinet">
      <a target="_blank" href="https://petlja.org/net.kabinet">
          <img src="../_static/img/netkabinet-logo-strelica-wt.svg" alt="petlja.org" id="lightNetKabinetLogo">
      </a>
  </div>
</div>


<div class="content-header-label" id = "content-header-label" href="../">Анализа података у Џупитеру и технике програмирања за други разред гимназије</div>

</div>




<div class="petljaCourseContainer">

<div class="petljaCourseContent">
  
           <div class="tableOfContent">
            <h5> Садржај </h5>
            </div>
           <div>
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J0A>
                    <h5>Пре него што почнемо</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J0A/howto.html><div class="studioLecture" id=activity-howto><i class="fas fa-file-alt activity-icon"></i>Како да користиш овај приручник</div></a>
                <a href=../J0A/J0A.html><div class="studioLecture" id=activity-J0A><i class="fas fa-file-alt activity-icon"></i>Покретање Џупитер радних свески</div></a>
                <a href=../J0A/J01.html><div class="studioLecture" id=activity-J01><i class="fas fa-file-alt activity-icon"></i>Увод у Џупитер</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-podsetnik>
                    <h5>Подсетник Пајтон</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../podsetnik/Podsetnik%20-%20Pajton.html><div class="studioLecture" id=activity-Podsetnik%20-%20Pajton><i class="fas fa-file-alt activity-icon"></i>Подсетник Пајтон</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y01>
                    <h5>Низови података</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y01/L801.html><div class="studioLecture" id=activity-L801><i class="fas fa-file-alt activity-icon"></i>Увод у процесирање низова података</div></a>
                <a href=../Y01/03_zadaci.html><div class="studioLecture" id=activity-03_zadaci><i class="fas fa-file-alt activity-icon"></i>Задаци - процесирање низова</div></a>
                <a href=../Y01/Kreiranje%20lista%20u%20Pajtonu%20-%20List%20comprehension.html><div class="studioLecture" id=activity-Kreiranje%20lista%20u%20Pajtonu%20-%20List%20comprehension><i class="fas fa-file-alt activity-icon"></i>Креирање листа у Пајтону</div></a>
                <a href=../Y01/Низови%20-%20додатни%20задаци.html><div class="studioLecture" id=activity-Низови%20-%20додатни%20задаци><i class="far fa-question-circle activity-icon"></i>Низови - додатни задаци</div></a>
                <a href=../Y01/test801.html><div class="studioLecture" id=activity-test801><i class="far fa-question-circle activity-icon"></i>Квиз - Листе</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y05>
                    <h5>Сортирање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y05/L805.html><div class="studioLecture" id=activity-L805><i class="fas fa-file-alt activity-icon"></i>Сортирање</div></a>
                <a href=../Y05/04_zadaci.html><div class="studioLecture" id=activity-04_zadaci><i class="fas fa-file-alt activity-icon"></i>Задаци - сортирање</div></a>
                <a href=../Y05/Сортирање%20-%20додатни%20задаци.html><div class="studioLecture" id=activity-Сортирање%20-%20додатни%20задаци><i class="far fa-question-circle activity-icon"></i>Сортирање - додатни задаци</div></a>
                <a href=../Y05/test805.html><div class="studioLecture" id=activity-test805><i class="far fa-question-circle activity-icon"></i>Квиз - Сортирање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y06>
                    <h5>Филтрирање и претраживање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y06/L806.html><div class="studioLecture" id=activity-L806><i class="fas fa-file-alt activity-icon"></i>Филтрирање и претраживање</div></a>
                <a href=../Y06/05_zadaci.html><div class="studioLecture" id=activity-05_zadaci><i class="fas fa-file-alt activity-icon"></i>Задаци - филтрирање и претраживање</div></a>
                <a href=../Y06/test806.html><div class="studioLecture" id=activity-test806><i class="far fa-question-circle activity-icon"></i>Квиз - Филтрирање и претраживање</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-predstavljanje>
                    <h5>Представљање низова података</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../predstavljanje/Predstavljanje%20nizova%20podataka.html><div class="studioLecture" id=activity-Predstavljanje%20nizova%20podataka><i class="fas fa-file-alt activity-icon"></i>Представљање низова података</div></a>
                <a href=../predstavljanje/Graficko_predstavljanje.html><div class="studioLecture" id=activity-Graficko_predstavljanje><i class="fas fa-file-alt activity-icon"></i>Графичко представљање низова података</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-tekstualni>
                    <h5>Рад са текстуалним подацима</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../tekstualni/Rad%20sa%20tekstualnim%20podacima.html><div class="studioLecture" id=activity-Rad%20sa%20tekstualnim%20podacima><i class="fas fa-file-alt activity-icon"></i>Рад са текстуалним подацима</div></a>
                <a href=../tekstualni/tekstualne_datoteke.html><div class="studioLecture" id=activity-tekstualne_datoteke><i class="fas fa-file-alt activity-icon"></i>Коришћење текстуалних датотека</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-rad%20sa%20tabelama>
                    <h5>Рад са табелама</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../rad%20sa%20tabelama/Rad%20sa%20tabelama.html><div class="studioLecture" id=activity-Rad%20sa%20tabelama><i class="fas fa-file-alt activity-icon"></i>Рад са табелама</div></a>
                <a href=../rad%20sa%20tabelama/indeksiranje.html><div class="studioLecture" id=activity-indeksiranje><i class="fas fa-file-alt activity-icon"></i>Индексирање табела</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-obrada_tabele>
                    <h5>Обрада и приказ табеларних података</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../obrada_tabele/obrada_prikaz.html><div class="studioLecture" id=activity-obrada_prikaz><i class="fas fa-file-alt activity-icon"></i>Обрада и приказ табеларних података</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J10>
                    <h5>Џупитер и Ексел</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J10/J10.html><div class="studioLecture" id=activity-J10><i class="fas fa-file-alt activity-icon"></i>Џупитер и Ексел</div></a>
                <a href=../J10/testJ10.html><div class="studioLecture" id=activity-testJ10><i class="far fa-question-circle activity-icon"></i>Џупитер и Ексел</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-projektni%20zadatak>
                    <h5>Пример скупа података за пројекто учење</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../projektni%20zadatak/projekat.html><div class="studioLecture" id=activity-projekat><i class="fas fa-file-alt activity-icon"></i>Пример скупа података за пројекто учење</div></a>
                </div>
            </div>
                </div>
</div>



<div class="lectureContent">

  <div id="errors">
  </div>
  
  <div class="lectureContentMaterial">
  
    <div id="main-content">
    
  <div class="section" id="id1">
<h1>Подсетник - Пајтон<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>У овој свесци вам дајемо кратак подсетник на основне елемента и
правила рада у Пајтону. Ако сте заборавили како се програмира у овом
програмском језику или сте, из неког разлога, раније пропустили те
лекције, овај подсетник би требало да вам помогне. Као и код било ког
другог програмског језика, најтеже је памтити синтаксу, односно
називе инструкција и правила како се оне комбинују. Овде имате
примере неких основних комбинација.</p>
<div class="section" id="id2">
<h2>Типови података<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>
<p>Пајтон у свом стандардном окружењу, без библиотека има три основна типа,
односно групе типова података: нумеричке, знаковне и логичке, али у
оквиру њих постоји и финија подела.</p>
<div class="section" id="id3">
<h3>Нумерички подаци<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h3>
<p>Пајтон познаје три типа нумеричких, односно бројчаних података:
целобројне вредности (<em>integer</em>), децималне бројеве (<em>float</em>) и
комплексне бројеве (<em>complex</em>). Ког је типа променљива, увек можемо да
проверимо помоћу функције <code class="docutils literal notranslate"><span class="pre">type()</span></code>.</p>

<div data-childcomponent="integer" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=integer data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
i=10
print(type(i))
</textarea>
</div>
<p>Број 10 је свакако цео број, али такође припада и скупу рационалних,
реалних, комплексних итд. Разлог због ког је Пајтон тип ове променљиве
препознао баш као <code class="docutils literal notranslate"><span class="pre">int</span></code> је то што бира најекономичнију варијанту. Број
10 ће имати најмањи запис и трошити најмање меморије ако га чувамо као
“интиџер”. Уколико број има децимални запис (макар из децималне тачке
стајала и нула) тип тог податка ће бити препознат као <code class="docutils literal notranslate"><span class="pre">float</span></code>.</p>

<div data-childcomponent="float" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=float data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
f=10.0
print(type(f))
</textarea>
</div>
<p>Пајтон ради и са комплексним бројевима. Имагинарни део комплесног броја
се записује као број иза кога иде слово ј, без размака.</p>

<div data-childcomponent="complex" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=complex data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
c=5j
print(type(c))
</textarea>
</div>
<p>Без обзира што Пајтон разликује целе бројеве, реалне и комплексне, он
рачуна са њима заједно.</p>

<div data-childcomponent="mnozenje" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=mnozenje data-lang="python" 
     data-include="integer float complex" data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
print(i*f*c)
</textarea>
</div>
</div>
<div class="section" id="id4">
<h3>Знаковни подаци<a class="headerlink" href="#id4" title="Permalink to this headline">¶</a></h3>
<p>Знаковне податке обично називамо стринговима или нискама. Одатле ознака
<code class="docutils literal notranslate"><span class="pre">str</span></code>. То су низови алфа-нумеричких карактера записани као текст. Када
променљивој додељујемо вредност која је типа стринг, онда тај податак
стављамо под једноструке или деоструке знаке навода. Битно је да на
почетку и на крају буде иста врста наводника.</p>

<div data-childcomponent="string1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=string1 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
s='Програмирање у Пајтону'
print(type(s))
</textarea>
</div>
<p>Рачунске операције са стринговима нису исте као са бројевима. Сабирање
стрингова постоји и значи додавање једног низа карактера на други.
Множење два стринга није могуће, али стринг и интиџер могу да се
помноже. Остале комбинације нису могуће. Погледајте примере.</p>

<div data-childcomponent="string2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=string2 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = "Пај"
b = "тон"
print(a+b)
</textarea>
</div>

<div data-childcomponent="sting3" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=sting3 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = "тон"
print(a*8)
</textarea>
</div>
</div>
<div class="section" id="id5">
<h3>Логичке вредности<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h3>
<p>Тип променљиве која садржи вредност логичког исказа називамо “булеан”
(<em>boolean</em>). Булеани имају своје вредности које могу да буду или тачно
(<em>True</em>) или нетачно (<em>False</em>). Свако поређење вредности променљивих
(било ког типа) помоћу оператора једнако (==), није једнако (!=), мање
(&lt;), веће (&gt;) резултује логичком вредношћу.</p>

<div data-childcomponent="bool1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=bool1 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = 5==5.0
print(a)
</textarea>
</div>

<div data-childcomponent="bool2" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=bool2 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
print(type(5==5.0))
</textarea>
</div>
<p>Операције за рад са логичким вредностима су логичко и (<code class="docutils literal notranslate"><span class="pre">and</span></code> или
<code class="docutils literal notranslate"><span class="pre">&amp;</span></code>), логичко или (<code class="docutils literal notranslate"><span class="pre">or</span></code> или <code class="docutils literal notranslate"><span class="pre">|</span></code>) и не (<code class="docutils literal notranslate"><span class="pre">not</span></code>). Резултати ових
логичких операција су такође логичке вредности.</p>

<div data-childcomponent="bool3" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=bool3 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = (5==5) & (5!=6)
print(a)
</textarea>
</div>
<p>Осим резултата поређења вредности променљивих, Пајтон “тривијалним”
вредностима нумеричких и знаковних променљивих (нпр, 0, ’’), као и
празним листама, скуповима или речницима придружује логичку вредност
<em>False</em>, док свим осталим даје <em>True</em>.</p>

<div data-childcomponent="bool4" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=bool4 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = bool('')
print(a)
</textarea>
</div>

<div data-childcomponent="bool5" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=bool5 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
a = bool([False])
print(a)
</textarea>
</div>
</div>
<div class="section" id="id6">
<h3>Претварање вредности у други тип податка<a class="headerlink" href="#id6" title="Permalink to this headline">¶</a></h3>
<p>У неким случајевима, тип податка је могуће претворити у други. На пример
стринг “2.75” лако може да се претвори у реалан број 2.75. Обрнуто
такође важи.</p>

<div data-childcomponent="stringtofloat" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=stringtofloat data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = float("2.75")
print(a)
</textarea>
</div>

<div data-childcomponent="floattostring" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=floattostring data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = str(1/7)
print(a)
</textarea>
</div>
<p>Претварање интиџера у децимални број је једнозначно, само се промени
формат, а вредност остаје иста. У супротном смеру морамо да занемаримо
цифре иза децималне тачке због чега се вредност променљиве мења.</p>

<div data-childcomponent="inttofloat" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=inttofloat data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = float(3)
print(a)
</textarea>
</div>

<div data-childcomponent="floattoint" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=floattoint data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = int(10/3)
print(a)
</textarea>
</div>
<p>Булеане је такође могуће претворити у нумеричке вредности или стрингове.
У супротном смеру, то значи губитак информације и свођење на бинарну
вредност која говори да ли је почетна вредност нетривијална.</p>

<div data-childcomponent="booltostr" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=booltostr data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = str(6>4)
print(a)
</textarea>
</div>

<div data-childcomponent="floattobool" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=floattobool data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a = bool(2.5)
print(a)
</textarea>
</div>
</div>
</div>
<div class="section" id="id7">
<h2>Колекције<a class="headerlink" href="#id7" title="Permalink to this headline">¶</a></h2>
<p>Пајтон различите типове скупова података назива колекцијама. Постоје
четири освновне колекције које имамо на располагању чим покренемо
Пајтон. То су листе, речници, торке и скупови. Ако бисте учитали модул
<code class="docutils literal notranslate"><span class="pre">collections</span></code> добили бисте још неке, али су нам ове четири које се
подразумевају више него довољне за почетак.</p>
<div class="section" id="list">
<h3>Листе (<code class="docutils literal notranslate"><span class="pre">list</span></code>)<a class="headerlink" href="#list" title="Permalink to this headline">¶</a></h3>
<p>Листу правимо тако што унутар угластих заграда наводимо различите
податке раздвајајући их зарезом. Елементи листе могу бити подаци свих
типова па чак и саме листе. Подаци могу да се понављају. Важна
карактеристика листа је да су елементи нумерисани. Сваки елемент има
свој индекс који почиње од нуле. Елементи из листе се позивају преко
имена листе и индекса у угластим заградама.</p>

<div data-childcomponent="lsita" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=lsita data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
lista=[2,3,'četiri']
print(lista[2])
</textarea>
</div>
</div>
<div class="section" id="dict">
<h3>Речници (<code class="docutils literal notranslate"><span class="pre">dict</span></code>)<a class="headerlink" href="#dict" title="Permalink to this headline">¶</a></h3>
<p>Друга важна колекција су речници. Ту су елементи уређени парови који
увек имају кључ (<em>key</em>) и вредност (<em>value</em>). Кључ и вредност су
раздвојени знаком <code class="docutils literal notranslate"><span class="pre">:</span></code>, док су парови раздвојени зарезом унитар
витичастих заграда. Елементи не могу да се понављају и сваки кључ може
да има само једну вредност. Елементи речника нису нумерисани. Уместо
индекса овде користимо кључ за референцирање.</p>

<div data-childcomponent="recnik" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=recnik data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
dic={'Земља':'Earth','Месец':'Moon','Сунце':'Sun'}
print(dic['Месец'])
</textarea>
</div>
</div>
<div class="section" id="tuple">
<h3>Торке (<code class="docutils literal notranslate"><span class="pre">tuple</span></code>)<a class="headerlink" href="#tuple" title="Permalink to this headline">¶</a></h3>
<p>Торке су врло сличне листама. Елементи се углавном набрајају унутар
обичних заграда, али може и без њих. Једина важна разлика у односу на
листе је што вредности елемената кад су једном дефинисане не могу више
да се мењају. Ако хоћете другачију торку, направите нову.</p>

<div data-childcomponent="tupple" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=tupple data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
torka=(2,3,3,5)
print(torka[3])
</textarea>
</div>
<p>Врло важно својство колекција је да дозвољавају да колекције једне врсте
претварамо у другу. На пример, ако од торке хоћемо да направимо листу,
довољно је да је трансформишемо помоћу функције <code class="docutils literal notranslate"><span class="pre">list()</span></code>.</p>

<div data-childcomponent="torkaulistu" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=torkaulistu data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
torka=(2,3,3,5)
torka = list(torka)
print(torka)
</textarea>
</div>
<p>Торке имају могућност да већем броју променљивих у једној линији доделе
вредности. Потребно је само да наведемо имена променљивих одвојених
зарезом, онда знак једнако и потом низ вредности које придружујемо
променљивим. Веома је корисна могућност да помоћу торки двема
променљивим заменимо вредности. У другим програмским језицима је за
замену вредности потребно увести и трећу, помоћну промељиву. У Пајтону
је то крајње једноставно: само им заменимо места. Погледајте пример.</p>

<div data-childcomponent="torka_menjanje" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=torka_menjanje data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a,b=3,5
a,b=b,a
print(a,b)
</textarea>
</div>
</div>
<div class="section" id="set">
<h3>Скупови (<code class="docutils literal notranslate"><span class="pre">set</span></code>)<a class="headerlink" href="#set" title="Permalink to this headline">¶</a></h3>
<p>Скупови су, као и математички скупови, колекције различитих елемената
који не могу да се понављају. Ове колекције нису нумерисане и не могу се
индексирати. Над скуповима су дефинисане операције које су нам познате
из математике: пресек, унија и разлика.</p>

<div data-childcomponent="set1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=set1 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
s1={'beli','plavi','crveni','zeleni'}
s2={'plavi','crni','sivi'}
unija=s1.union(s2)  # unija prvog i drugog skupa
presek=s1.intersection(s2) # presek prvog i drugog skupa
razlika=s1.difference(s2) # razlika prvog i drugog skupa
print(presek)
</textarea>
</div>
<p>Важна особина скупова је то што скупу не можете поново да придружите
елемент који се већ налази у скупу. Нема дупликата. Кад год нам је
потребно да нађемо јединствене елементе неке друге колекције, можемо
помоћу скупа да уклонимо дупликате. На пример, ако узмемо листу у којој
има дупликата, претворимо је у скуп, а по потом поново у листу, у тој
другој листи ће остати само јединствене вредности. Погледајте пример.</p>

<div data-childcomponent="lista_set" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=lista_set data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
lista1=[1,2,2,2,3,4,5,6]
skup=set(lista1)  # listu pretvaramo u skup kako bismo ukunuli duplikate
lista2=list(skup) # sad skup vreaćamo u listu
print(lista2)
</textarea>
</div>
</div>
</div>
<div class="section" id="id8">
<h2>Петље<a class="headerlink" href="#id8" title="Permalink to this headline">¶</a></h2>
<p>Када део програма треба да понавимо више пута, пожељно је да тај кôд
ставимо у петљу. Петља је програмска структура која нам омогућава да
контролошемо колико ће се пута блок инструкција извршавати. Свако
извршавање кôда у петљи назива се <strong>итерација</strong>. Променљива која броји
колико пута је блок инструкција извршен назива се итератор или
<strong>бројач</strong>.</p>
<p>Најчешће коришћена петља у Пајтону је <em>for</em> петља. Ова петља почиње
инструкцијом <code class="docutils literal notranslate"><span class="pre">for</span></code>, иза ког иде услов да се бројач (итератор) налази у
колекцији која следи. Оператор <code class="docutils literal notranslate"><span class="pre">in</span></code> између бројача и колекције
означава услов да вредност бројача припада том скупу вредности. Прва
линија петље завршава се знаком две тачке <code class="docutils literal notranslate"><span class="pre">:</span></code>. Следи блок инструкција
који је увучен.</p>
<p>Петља у наредној ћелији је пример једноставног програма који исписује
бројеве од 0 до 9.</p>

<div data-childcomponent="petlja1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=petlja1 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
for i in range(10):
    print(i)
</textarea>
</div>
<p>Скуп вредности које узима променљива током итерирања овде је био
дефинисан са <code class="docutils literal notranslate"><span class="pre">range(10)</span></code>, што је распон бројева од 0 до 9. Како у
Пајтону индексирање увек почиње од нуле, вредности ове “итерабле” иду од
0 до 9. Реч итерабла се односи на све колекције кроз које Пајтон може да
пролази секвенцијално, односно тачно утврђеним редоследом од прве до
последње вредности. Листе су такође итерабле па и њих можемо да
користимо за конструисање петљи. Када итерирамо по колекцији типа
<em>range</em>, итератори су бројеви. Ако користимо листу као итераблу, онда су
вредности итератора вредности из листе које редом узимамо.</p>

<div data-childcomponent="lista_stringova" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=lista_stringova data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
for s in ['први','други','трећи']:
    print(s)
</textarea>
</div>
<p>Итерабле такође могу да буду и скупови, стрингови, торке итд. Стринг
није ништа друго до листа карактера. Зато можемо да итерирамо кроз
стринг, односно да правимо петљу у којој ће итератор узимати вредности
из стринга. Ево једноставног примера:</p>

<div data-childcomponent="petlja_string" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=petlja_string data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
for c in "Python":
    print(c)
</textarea>
</div>
<p>Осим <em>for</em> петље, Пајтон има и <em>while</em> петљу. Код ње није битно да
итератор буде у одређеном скупу. <em>while</em> петља се извршава док год је
припадајући услов испуњен, односно док је логички исказ тачан. Ево
примера како при пример са <em>for</em> петљом можемо да напишемо користећи
инструкцију <code class="docutils literal notranslate"><span class="pre">while</span></code>.</p>

<div data-childcomponent="while_petlja" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=while_petlja data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
i=0          # početna vrednost iteratora/brojača
while i<10:  # dok god je brojač manji od 10
    print(i) # štampaj njegovu vrednost
    i+=1     # i povećavaj je za 1
</textarea>
</div>

    <div class="note-wrapper suggestionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/suggestionnote-img.svg" class="note-image suggestionnote-image" /> 
        <div class="course-content">
            
<p>Приметите да овде “ручно” морамо да мењамо вредност индекса <code class="docutils literal notranslate"><span class="pre">i</span></code>. Израз
<code class="docutils literal notranslate"><span class="pre">i+=1</span></code> је скраћени запис од <code class="docutils literal notranslate"><span class="pre">i=i+1</span></code>. То значи да у свакој итерацији
вредност индекса <code class="docutils literal notranslate"><span class="pre">i</span></code> повећавамо за 1.</p>

    </div></div>
<div class="section" id="id9">
<h3>Угњеждене петље<a class="headerlink" href="#id9" title="Permalink to this headline">¶</a></h3>
<p>Унутар једне петље лако можете да ставите другу. Пазите само да
итератори буду различити. Типичан пример проблема који можете да решите
петљом у петљи је да направите табелу где рачунате резултате неке
операције између свака два елемента низа. Конкретно, можемо да направимо
таблицу множења. Потребно је само сваки елемент једног низа помножити са
сваким елементом другог. Следи пример једноставног програма који то
рачуна и исписује на екрану.</p>

<div data-childcomponent="ugnježđene" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=ugnježđene data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
n=5
for i in range(1,n+1):      # za svako i od 1 do 5
    for j in range (1,n+1): # za svako j od 1 do 5
        print(i*j,end="\t") # štampaj i*j i pomeri kurzor za jedan tab
    print("\n")             # kad j dođe do 5 premesti kurzor na početak sledećeg reda
</textarea>
</div>
<p>Ако занемарите мало чудан начин исписа, програм је крајње једноставан.
Узимамо све бројеве од 1 до 5, множимо сваки са сваким и приказујемо то
на екрану. Оно што је необично код исписа је што смо ставили да се испис
завршава “табом”, а не новим редом. То нам омогућава да испишемо 5
резултата у једном реду. Ред завршавамо тек када унутрашња петља стигне
до краја. Контролни знак за таб је “\t”, а за нови ред “\n”.</p>
</div>
</div>
<div class="section" id="id10">
<h2>Условно гранање<a class="headerlink" href="#id10" title="Permalink to this headline">¶</a></h2>
<p>Да би програм могао да одлучи који део кôда треба да изврши потребно је
да провери испуњеност траженог критеријума, односно да евалуира
(вреднује) одговарајући логички исказ. У зависности од тога да ли је
услов испуњен или не, програм ће извршавати различите делове кôда. У
Пајтону се најчешће користи <code class="docutils literal notranslate"><span class="pre">if</span> <span class="pre">...</span> <span class="pre">then</span> <span class="pre">...</span> <span class="pre">else</span></code> конструкција. Прво
проверавамо да ли је испуњен услов наведен иза <code class="docutils literal notranslate"><span class="pre">if</span></code>, ако јесте онда
програм извршава оно што иде после <code class="docutils literal notranslate"><span class="pre">then</span></code>, а ако није испуњен извршава
оно што следи после <code class="docutils literal notranslate"><span class="pre">else</span></code>. Погледајте пример.</p>

<div data-childcomponent="uslovno" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=uslovno data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a=int(input())       # traži unos sa tastature i prevaraga u integer
if a%2==0:           # ako je ostatak pri deljenju sa 2 jednak nuli
    print('паран')   # štampaj da je broj paran
else:                # ako nije
    print('непаран') # štampaj da je neparan
</textarea>
</div>
<p>У овом примеру прво тражимо да неко помоћу тастатуре укуца број, онда
тај унос претворимо у интиџер и придружимо га промељивој <code class="docutils literal notranslate"><span class="pre">a</span></code>. У
следећој линији проверавамо да ли је тај број дељив са 2, односно да ли
је остатак при дељењу са 2 једнак нули. Потом иде оно што треба извршти
ако је услов испуњен, тј. да програм испише да је унети број паран. У
противном исписује да је непаран.</p>

    <div class="note-wrapper suggestionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/suggestionnote-img.svg" class="note-image suggestionnote-image" /> 
        <div class="course-content">
            
<p>Обратите пажњу на синтаксу. У Пајтону се инструкција <code class="docutils literal notranslate"><span class="pre">then</span></code> уопште не
користи. Уместо тога користе се знак <code class="docutils literal notranslate"><span class="pre">:</span></code> и увучени блокови
инструкција. Видите како се после реда који се завршава са <code class="docutils literal notranslate"><span class="pre">:</span></code>
појављује увучени део кôда. Први блок инструкција (који може да има
произвољно много линија кôда) се извршава ако је услов испуњен, а други
блок (после <code class="docutils literal notranslate"><span class="pre">else:</span></code>) ако услов није испуњен.</p>

    </div></div>
<p>Уколико гранање на две могућности није довољно, можемо да уведемо и нове
блокове помоћу <code class="docutils literal notranslate"><span class="pre">elif</span></code> инструкције. Нема ограничења колико ових услова
може да буде. Када први услов није испуњен, онда можемо да испитамо да
ли је неки други услов испуњен помоћу <code class="docutils literal notranslate"><span class="pre">elif</span></code>. Погледајте пример.</p>

<div data-childcomponent="elif" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=elif data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a=5
b=6
if a==b:                           # ako su a i b jednaki
    print("a и b су једнаки.")     # ako su a i b jednaki
elif a<b:                          # ako je a manje od b
    print("a је мање од b.")       # napiši da je a manje od b
else:                              # u protivnom
    print("a је веће од b.")       # napiši da je a veće od b
</textarea>
</div>
<p>Други начин да проверимо више услова је да у <em>if-then-else</em> блоковима
уведемо нове услове и ново гранање. Ако то радите обратите пажњу да ти
нови блокови морају да буду још више увучени. Ако су они први били по 4
размака, онда ови угњеждени морају да буду по 8. Едитор у коме радите
обично води рачуна о увлачењу блокова, али није лоше да проверите да ли
је све увучено како треба. За разлику од неких других програмских
језика, Пајтон неће да ради и пријављује грешку ако блокови нису
правилно увучени. Ево примера како смо претходни пример могли да урадимо
са угњежденим условима.</p>

<div data-childcomponent="viseuslova" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=viseuslova data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a=6
b=6
if a<b:
    print("a је мање од b.")
else:
    if a==b:
        print("a и b су једнаки.")
    else:
        print("a је веће од b.")
</textarea>
</div>
<div class="section" id="id11">
<h3>Гранање због грешке<a class="headerlink" href="#id11" title="Permalink to this headline">¶</a></h3>
<p>Постоји могућност да направимо гранање које би нас штитило од ненамерне
грешке у извршавању програма. У идеалном случају, то нам није потребно,
али пракса говори да се грешке дешавају и да програми због тога престају
да раде. Ако можемо да претпоставимо каква је могућа грешка, онда можемо
да је контролишемо. Примера ради, пробајте да унесете децимални број у
горе приказани програм који исписује да ли је број паран или непаран.
Шта се дешава? Програм приказује информацију о “ValueError” грешци и
престаје да ради. Било би добро када бисмо могли реагујемо на грешку
тако што се у том случају изврши део програма предвиђен за такве
случајеве.</p>

    <div class="note-wrapper infonote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/infonote-img.svg" class="note-image infonote-image" /> 
        <div class="course-content">
            
<p>Инструкције <code class="docutils literal notranslate"><span class="pre">try</span></code> и <code class="docutils literal notranslate"><span class="pre">except</span></code> омогућавају да се програм регуларно
заврши и онда када дође до грешке. Програм покушава да уради оно што је
у блоку <code class="docutils literal notranslate"><span class="pre">try</span></code>. Уколико дође до грешке, онда извршава оно у блоку
<code class="docutils literal notranslate"><span class="pre">except</span></code>.</p>

    </div></div>
<p>Демонстрираћемо како <em>try-except</em> ради на примеру који смо
већ помињали.</p>

<div data-childcomponent="try" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=try data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
a=input()                          # unesi podatak sa tastature
try:                               # proba
    a=int(a)                       # ako je a celobrojna vrednost u decimalnom obliku, promeni mu tip u integer
    if a%2==0:                     # i onda proveri da li je broj paran ili neparan
        print('паран')
    else:
        print('непаран')
except:                                     # ako promena tipa u integer nije uspela
    print('Унели сте број који није цео.')  # onda samo ispiši poruku
</textarea>
</div>
<p>Гранање <em>try-except</em> има сложенију структуру него што је приказано у
примеру. Два додатна блока <code class="docutils literal notranslate"><span class="pre">else</span></code> и <code class="docutils literal notranslate"><span class="pre">finally</span></code> дају додатне
могућности за управљање грешкама, али их овде нећемо објашњавати. Ако
вас интересује, пронађите само упутство и прочитајте сами како раде.</p>
</div>
</div>
<div class="section" id="id12">
<h2>Функције<a class="headerlink" href="#id12" title="Permalink to this headline">¶</a></h2>
<p>Пајтон има неке уграђене функције као што су <code class="docutils literal notranslate"><span class="pre">sum()</span></code> или <code class="docutils literal notranslate"><span class="pre">min()</span></code>,
али за конкретну намену сами можемо да напишемо нову функцију. Функције
се једноставно дефинишу: у првој линији пишемо
<code class="docutils literal notranslate"><span class="pre">def</span> <span class="pre">&lt;име</span> <span class="pre">функције&gt;(&lt;име</span> <span class="pre">променљиве</span> <span class="pre">у</span> <span class="pre">коју</span> <span class="pre">се</span> <span class="pre">смешта</span> <span class="pre">аргумент&gt;):</span></code>,
онда следи блок инструкција којима обрађујемо и рачунамо нешто са
улазном променљивом и на крају, у последњој линији, пишемо
<code class="docutils literal notranslate"><span class="pre">return</span> <span class="pre">&lt;име</span> <span class="pre">излазне</span> <span class="pre">променљиве&gt;</span></code>. Свака функција може да има више
улазних променљивих, као што може да врати и више излазних.</p>
<p>Први пример функције који ћемо овде приказати је рачунање производа свих
вредности у листи. Оваква функција не постоји у стандардној Пајтоновој
библиотеци. Зато је добро да знамо да је напишемо.</p>

<div data-childcomponent="funct1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=funct1 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
def proizvod(L):     # argument, odnosno lista koju prosleđujemo funkciji biće smeštena u promenljivu
    p=1              # početna vrednost proizvoda je 1
    for x in L:      # za svaku vrednost iz liste
        p=p*x        # proizvod množimo tom vrednošću
    return p         # funkcija vraća proizvod svih vrednosti u listi

proizvod([1,2,3,4,5])
</textarea>
</div>

    <div class="note-wrapper questionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/question-mark.png" class="note-image questionnote-image" /> 
        <div class="course-content">
<p>Задатак: Можете ли, по угледу на претходни пример, сами да напишете функцију <code class="docutils literal notranslate"><span class="pre">zbir()</span></code> која би сабрала све елементе листе?</p>

    </div></div>

<div data-childcomponent="лифт_апсолутна_вредност_input" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=лифт_апсолутна_вредност_input data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="test1 test2 test3"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
# -*- acsection: main -*-
def zbir(L):
   # dopuni ovu fuknciju
# Провера
test1 = zbir([1, 2, 3, 4, 5])
test2 = zbir([100, 200, 500])
test3 = zbir([-1, 0, 0, 4, 5])
# -*- acsection: after-main -*-
print(test1, test2, test3)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        L = [1, 2, 3, 4, 5]; rez1 = sum(L)
        L = [100, 200, 500]; rez2 = sum(L)
        L = [-1, 0, 0, 4, 5]; rez3 = sum(L)
        run_test = acMainSection(test1=test1,test2=test2,test3=test3)
        self.assertEqual(run_test["test1"], rez1, "Вредност променљиве 'test1' треба да буде %s" % rez1)
        self.assertEqual(run_test["test2"], rez2, "Вредност променљиве 'test2' треба да буде %s" % rez2)
        self.assertEqual(run_test["test3"], rez3, "Вредност променљиве 'test3' треба да буде %s" % rez3)
myTests().main()
</textarea>
</div>

    <div class="note-wrapper questionnote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/question-mark.png" class="note-image questionnote-image" /> 
        <div class="course-content">
<p>Како бисте написали функцију чији је улаз листа, а излаз најмања
вредност из те листе? Та функција треба да пронађе најмању вредност у
листи и да је врати. Оваква функција већ постоји у Пајтону и зове се
<code class="docutils literal notranslate"><span class="pre">min()</span></code>, али ћемо искористи прилику да напишемо своју верзију.</p>

    </div></div>
<p>Име ове нове функције ће бити <code class="docutils literal notranslate"><span class="pre">minimum()</span></code>. Ево једне могућности како то може
да се напише.</p>

<div data-childcomponent="minimum" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=minimum data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
def minimum(L):
    m=L[0]            # za početak, najmanje vrednost će biti ona prva u listi
    for x in L[1:]:   # iteracije kreću od druge vrednosti u listi i idu do kraja
        if x<m:       # ako je tekuća vrednost iteratora x manja од m, tj. one koja je do sad bila najmanja
            m=x       # onda neka novi minimum bude
    return m          # na kraju funkcija vraća vrednost
lista=[11,22,33,14,25,36,17,28,39,10,21,32]
print(minimum(lista))
</textarea>
</div>
</div>
<div class="section" id="id13">
<h2>Библиотеке<a class="headerlink" href="#id13" title="Permalink to this headline">¶</a></h2>
<p>Пајтон без додатних библиотека има само основне типове података,
структуре и функције. У тој стандардној библиотеци нема структура у
којима бисмо чували табеле, нема функција које би цртале графиконе или
чак рачунале једноставне математичке функције. За све то је потребно да
учитамо библиотеке, модуле из библиотека или њихове појединачне
функције. Срећом, Пајтон има огроман број врло квалитетних библиотека до
којих се лако долази.</p>
<p>Библиотеке се састоје из модула који у себи садрже разне типове
података, објекте, функције и методе за те објекте итд. Стандардна
библиотека има неколико модула које не учитава аутоматски. То значи да
су модули инсталирани, али да треба да их учитамо. Међу њима су модули
који су нам често потребни као што су <code class="docutils literal notranslate"><span class="pre">math</span></code> (математичке функције),
<code class="docutils literal notranslate"><span class="pre">random</span></code> (генерисање случајних брпјева) или <code class="docutils literal notranslate"><span class="pre">statistics</span></code>
(статистичке функције).</p>
<p>Када учитамо целу библиотеку или модул помоћу инструкције <code class="docutils literal notranslate"><span class="pre">import</span></code>,
можемо да користимо све њене функције и објекте само морамо да наведемо
одакле су преузети. На пример, када учитамо библитеку <code class="docutils literal notranslate"><span class="pre">math</span></code>, добићемо
и функције и константе које уз њу иду. Кад их позивамо прво куцамо име
библиотеке са тачком на крају, тј. <code class="docutils literal notranslate"><span class="pre">math.</span></code> па онда име функције или
константе. Ево примера где програм прво исписује вредност константе
<span class="math notranslate nohighlight">\(\pi\)</span>, а затим и косинус од тог броја:</p>

<div data-childcomponent="ipython3" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=ipython3 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
               data-hidehistory=true
       >
import math
print(math.pi)
print(math.cos(math.pi))
</textarea>
</div>
<p>Има библиотека које имају непрактично дуга имена па би било тешко стално
укуцавати њихова имена кад хоћемо да позивамо њихове функције. У тим
случајевима већ при учитавању библиотеке кажемо Пајтону да ћемо је
убудуће позивати по “надимку”, тј. скраћеној верзији имена коју можемо
сами да смислимо. У следећој ћелији је пример како учитавамо библиотеку
<code class="docutils literal notranslate"><span class="pre">matploylib.pyplot</span></code> одређујући да ћемо је убудуће звати само <code class="docutils literal notranslate"><span class="pre">plt</span></code>.</p>
<div class="highlight-ipython3 notranslate"><div class="highlight"><pre><span></span><span class="kn">import</span> <span class="nn">matplotlib.pyplot</span> <span class="k">as</span> <span class="nn">plt</span>
<span class="n">plt</span><span class="o">.</span><span class="n">plot</span><span class="p">([</span><span class="mi">4</span><span class="p">,</span><span class="mi">2</span><span class="p">,</span><span class="mi">3</span><span class="p">,</span><span class="mi">4</span><span class="p">,</span><span class="mi">3</span><span class="p">])</span>
</pre></div>
</div>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="p">[</span><span class="o">&lt;</span><span class="n">matplotlib</span><span class="o">.</span><span class="n">lines</span><span class="o">.</span><span class="n">Line2D</span> <span class="n">at</span> <span class="mh">0x26d58a0b160</span><span class="o">&gt;</span><span class="p">]</span>
</pre></div>
</div>
<img alt="../_images/output_104_1.png" src="../_images/output_104_1.png" />
<p>Ако нас интересује само једна функција из библиотеке, можемо њу посебно
да учитамо помоћу линије
<code class="docutils literal notranslate"><span class="pre">from</span> <span class="pre">&lt;назив</span> <span class="pre">библиотеке&gt;</span> <span class="pre">import</span> <span class="pre">&lt;назив</span> <span class="pre">функције&gt;</span></code>. Тада при употреби
не морамо да наводимо из које је библиотеке. На пример, ако из целог
модула <code class="docutils literal notranslate"><span class="pre">statistics</span></code> хоћемо само функцију <code class="docutils literal notranslate"><span class="pre">mean()</span></code> која рачуна средњу
вредност низа, то можемо да урадимо на следећи начин:</p>
<div class="highlight-ipython3 notranslate"><div class="highlight"><pre><span></span><span class="kn">from</span> <span class="nn">statistics</span> <span class="kn">import</span> <span class="n">mean</span>
<span class="n">mean</span><span class="p">([</span><span class="mi">3</span><span class="p">,</span><span class="mi">4</span><span class="p">,</span><span class="mi">4</span><span class="p">,</span><span class="mi">6</span><span class="p">])</span>
</pre></div>
</div>
<div class="highlight-default notranslate"><div class="highlight"><pre><span></span><span class="mf">4.25</span>
</pre></div>
</div>

    <div class="note-wrapper infonote-type">
        <div class="note-icon-holder"> </div>
        <img src="../_static/img/infonote-img.svg" class="note-image infonote-image" /> 
        <div class="course-content">
            
<p>Многе важне екстерене библиотеке су већ инсталиране у Џупитер окружењу.
Оне које се најчешће користе су <code class="docutils literal notranslate"><span class="pre">pandas</span></code> за рад са табелама и
фајловима, <code class="docutils literal notranslate"><span class="pre">matplotlib.pyplot</span></code> за цртање графикона и <code class="docutils literal notranslate"><span class="pre">numpy</span></code> за рад
са вишедимензионалним подацима. Уколико тражене библиотеке нема у
окружењу, потребно је да је инсталирате. У већини случајева нема потребе
да библиотеку тражите на интернету и да је преузимате. Довољно је да
укуцате <code class="docutils literal notranslate"><span class="pre">pip</span> <span class="pre">install</span> <span class="pre">&lt;назив</span> <span class="pre">библиотеке&gt;</span></code> и Пајтон ће библиотеку сам
пронаћи, преузети и инсталирати.</p>

    </div></div>
</div>
</div>


    </div>
    
    <div>
<div class="lecture-prev-next-toggle">
  <a  href="index.html" id="prevLectureLink" class="lecture-toggler-arrows"><i class="fas fa-arrow-left fa"></i> <span> Претходна лекција </span> </a>
  
  <a id="nextLectureLink" class="lecture-toggler-arrows"><span> Следећа лекција</span>  <i class="fas fa-arrow-right fa"></i> </a>
</div>

<script type="text/javascript">
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
</div>

    
  </div>
</div>

<div class="petljadoc-resources">
  <span class="pull-left">(Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>) <br> © 2022 Petlja</span>
</div>


<div class="lectureToolbox">

  <label class="font-size-toggle" id="fontSizeToggle">

      <i class="fas fa-font"></i><label>Слова</label>
  </label>
  <div class="font-change d-none" id="fontChangeModal">
      <span class="dec-letter-icon" id="decLetterIcon">A-</span>
      <label id="currentFontSize">16</label>
      <span class="inc-letter-icon" id="incLetterIcon">A+</span>
  </div>
</div>

</div>

  </body>
</html>