<!DOCTYPE html>
<html xmlns="http://www.w3.org/1999/xhtml" lang="sr-Cyrl">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Угнежђене петље</title>
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="https://stackpath.bootstrapcdn.com/bootstrap/4.2.1/css/bootstrap.min.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/karel.css" />
    <link rel="stylesheet" type="text/css" href="../_static/notes.css" />
    <link rel="stylesheet" type="text/css" href="../_static/simanim.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pycode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/p5js.css" />
    <link rel="stylesheet" type="text/css" href="../_static/gallery.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dbDirective.css" />
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="https://use.fontawesome.com/releases/v5.1.1/css/all.css" type="text/css" />
    <link rel="stylesheet" href="https://fonts.googleapis.com/css?family=Source+Sans+Pro" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-4.0.0-dist/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/flatly.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/petlja-course.css" type="text/css" />
    <link rel="stylesheet" href="../_static/nbstyle.css" type="text/css" />
    <script id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/language_data.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/mchoice-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/dragndrop-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/fitb-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/parsons-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/karelCorner.js"></script>
    <script type="text/javascript" src="../_static/karelRobot.js"></script>
    <script type="text/javascript" src="../_static/karelWorld.js"></script>
    <script type="text/javascript" src="../_static/karelRobotDrawer.js"></script>
    <script type="text/javascript" src="../_static/karelUI.js"></script>
    <script type="text/javascript" src="../_static/karel.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/karel-i18n.sr-Cyrl.js"></script>
    <script type="text/javascript" src="../_static/notes.js"></script>
    <script type="text/javascript" src="../_static/pygamelib-init.js"></script>
    <script type="text/javascript" src="../_static/blockly/blockly_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/blocks_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/python_compressed.js"></script>
    <script type="text/javascript" src="../_static/blockly/msg-sr.js"></script>
    <script type="text/javascript" src="../_static/blockpy/utilities.js"></script>
    <script type="text/javascript" src="../_static/blockpy/python_errors.js"></script>
    <script type="text/javascript" src="../_static/blockpy/ast_node_visitor.js"></script>
    <script type="text/javascript" src="../_static/blockpy/abstract_interpreter.js"></script>
    <script type="text/javascript" src="../_static/blockpy/pytifa.js"></script>
    <script type="text/javascript" src="../_static/blockpy/abstract_interpreter_definitions.js"></script>
    <script type="text/javascript" src="../_static/blockpy/python_to_blockly.js"></script>
    <script type="text/javascript" src="../_static/blockpy/imported.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/class.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/comment.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/comprehensions.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/dict.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/if.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/io.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/lists.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/sets.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/loops.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/parking.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/tuple.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/turtles.js"></script>
    <script type="text/javascript" src="../_static/blockpy/blocks/text.js"></script>
    <script type="text/javascript" src="../_static/blockpy-modal.js"></script>
    <script type="text/javascript" src="../_static/simanim.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/pyodide/v0.16.1/full/pyodide.js"></script>
    <script type="text/javascript" src="../_static/pycode.js"></script>
    <script type="text/javascript" src="//toolness.github.io/p5.js-widget/p5-widget.js"></script>
    <script type="text/javascript" src="../_static/p5js.js"></script>
    <script type="text/javascript" src="../_static/gallery.js"></script>
    <script type="text/javascript" src="../_static/dbDirective.js"></script>
    <script type="text/javascript" src="../_static/sql.js"></script>
    <script type="text/javascript" src="../_static/translations.js"></script>
    <script async="async" type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.5/latest.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-4.0.0-dist/js/bootstrap.bundle.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/rangy-core.js"></script>
    <script type="text/javascript" src="../_static/rangy-textrange.js"></script>
    <script type="text/javascript" src="../_static/rangy-cssclassapplier.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="../_static/course-errors.js"></script>
    <script type="text/javascript" src="../_static/petlja.js"></script>
    <script type="text/javascript" src="../_static/require.js"></script>
    <meta charset='utf-8'>
    <meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
    <meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
    <script type="text/javascript">
      eBookConfig = {};
      eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
      eBookConfig.app = eBookConfig.host + '/runestone';
      eBookConfig.ajaxURL = eBookConfig.app + '/ajax/';
      eBookConfig.course = 'Настава програмирања за Други разред гимназије';
      eBookConfig.logLevel = '0';
      eBookConfig.loginRequired = 'false';
      eBookConfig.build_info = "";
      eBookConfig.isLoggedIn = false;
      eBookConfig.useRunestoneServices = false;
      eBookConfig.python3 = true;
      eBookConfig.basecourse = 'Настава програмирања за Други разред гимназије';
      eBookConfig.runestone_version = '';
      eBookConfig.imagesDir = '../_images/';
      eBookConfig.staticDir = '../_static/';
      if (typeof (Sk) != "undefined")
        Sk.imgPath = eBookConfig.imagesDir;
    </script>

  </head>

  <body>



<div id="navbar" class="content-header">
  <div class="right-side-nav">
    
</div>

<div class="left-side-nav">
  <div class="course-logo-nav">
      <a target="_blank" href="https://petlja.org">
          <img src="../_static/img/petlja-logo-wt.svg" alt="petlja.org" id="lightPetljaLogo">
      </a>
  </div>
  <div class="course-logo-nav-kabinet">
      <a target="_blank" href="https://petlja.org/net.kabinet">
          <img src="../_static/img/netkabinet-logo-strelica-wt.svg" alt="petlja.org" id="lightNetKabinetLogo">
      </a>
  </div>
</div>


<div class="content-header-label" id = "content-header-label" href="../">Анализа података у Џупитеру и технике програмирања за други разред гимназије</div>

</div>




<div class="petljaCourseContainer">

<div class="petljaCourseContent">
  
           <div class="tableOfContent">
            <h5> Садржај </h5>
            </div>
           <div>
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J00>
                    <h5>Рад са подацима у текстуалном програмском језику (ПТПЈ)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J00/J00.html><div class="studioLecture" id=activity-J00><i class="fas fa-file-alt activity-icon"></i>Исходи</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J0A>
                    <h5>Припрема</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J0A/J0A.html><div class="studioLecture" id=activity-J0A><i class="fas fa-file-alt activity-icon"></i>Припрема</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J01>
                    <h5>ПТПЈ.1. Увод у Џупитер (*Jupyter*)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J01/J01.html><div class="studioLecture" id=activity-J01><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J02>
                    <h5>ПТПЈ.2. Низови података и линијски дијаграми</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J02/J02.html><div class="studioLecture" id=activity-J02><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J02/testJ02.html><div class="studioLecture" id=activity-testJ02><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J03>
                    <h5>ПТПЈ.3. Хистограми и боје</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J03/J03.html><div class="studioLecture" id=activity-J03><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J03/testJ03.html><div class="studioLecture" id=activity-testJ03><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J04>
                    <h5>ПТПЈ.4. Просек и медијана низа бројева</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J04/J04.html><div class="studioLecture" id=activity-J04><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J04/testJ04.html><div class="studioLecture" id=activity-testJ04><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J05>
                    <h5>ПТПЈ.5. Фреквенцијска анализа и секторски дијаграми</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J05/J05.html><div class="studioLecture" id=activity-J05><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J05/testJ05.html><div class="studioLecture" id=activity-testJ05><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J06>
                    <h5>ПТПЈ.6. Табеларно представљени подаци</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J06/J06.html><div class="studioLecture" id=activity-J06><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J06/testJ06.html><div class="studioLecture" id=activity-testJ06><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J07>
                    <h5>ПТПЈ.7. Индексирање и транспоновање табеле</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J07/J07.html><div class="studioLecture" id=activity-J07><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J07/testJ07.html><div class="studioLecture" id=activity-testJ07><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J08>
                    <h5>ПТПЈ.8. Модификације табеле и записивање табеле у датотеку</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J08/J08.html><div class="studioLecture" id=activity-J08><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J08/testJ08.html><div class="studioLecture" id=activity-testJ08><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J09>
                    <h5>ПТПЈ.9. Сортирање, филтрирање и фреквенцијска анализа</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J09/J09.html><div class="studioLecture" id=activity-J09><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J09/testJ09.html><div class="studioLecture" id=activity-testJ09><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-J10>
                    <h5>ПТПЈ.10. Џупитер и Ексел</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../J10/J10.html><div class="studioLecture" id=activity-J10><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../J10/testJ10.html><div class="studioLecture" id=activity-testJ10><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y00>
                    <h5>Програмирање (ПРОГ)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y00/Y00.html><div class="studioLecture" id=activity-Y00><i class="fas fa-file-alt activity-icon"></i>Исходи</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y01>
                    <h5>ПРОГ.1. Увод у процесирање низова података у Пајтону</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y01/L801.html><div class="studioLecture" id=activity-L801><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y01/test801.html><div class="studioLecture" id=activity-test801><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y02>
                    <h5>ПРОГ.2. Бројачи и суме</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y02/L802.html><div class="studioLecture" id=activity-L802><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y02/test802.html><div class="studioLecture" id=activity-test802><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y03>
                    <h5>ПРОГ.3. Низови симбола (*стрингови* или *ниске*)</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y03/L803.html><div class="studioLecture" id=activity-L803><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y03/test803.html><div class="studioLecture" id=activity-test803><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y04>
                    <h5>ПРОГ.4. Угнежђене петље</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y04/L804.html><div class="studioLecture" id=activity-L804><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y04/test804.html><div class="studioLecture" id=activity-test804><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y05>
                    <h5>ПРОГ.5. Сортирање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y05/L805.html><div class="studioLecture" id=activity-L805><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y05/test805.html><div class="studioLecture" id=activity-test805><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y06>
                    <h5>ПРОГ.6. Филтрирање и претраживање</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y06/L806.html><div class="studioLecture" id=activity-L806><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y06/test806.html><div class="studioLecture" id=activity-test806><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                
                <div class="courseContentsSection">
                <div class="lecture-div courseContentsSectionHeadline" id=lecture-Y07>
                    <h5>ПРОГ.7. Рад са текстуалним датотекама</h5>
                    <i class="fas fa-angle-down caret-position "> </i>
                    <i class="fas fa-angle-up caret-position d-none"></i>
                </div>
                <div class="courseContentsSubSection d-none"><a href=../Y07/L807.html><div class="studioLecture" id=activity-L807><i class="fas fa-file-alt activity-icon"></i>Лекција</div></a>
                <a href=../Y07/test807.html><div class="studioLecture" id=activity-test807><i class="far fa-question-circle activity-icon"></i>Квиз</div></a>
                </div>
            </div>
                </div>
</div>



<div class="lectureContent">

  <div id="errors">
  </div>
  
  <div class="lectureContentMaterial">
  
    <div id="main-content">
    
  <div class="section" id="id1">
<h1>Угнежђене петље<a class="headerlink" href="#id1" title="Permalink to this headline">¶</a></h1>
<p>У овој лекцији показујемо:</p>
<ol class="arabic simple">
<li><p>како раде угнежђене петље,</p></li>
<li><p>како се користе угнежђене петље за исписивање табела са подацима, и</p></li>
<li><p>како се обрађују подаци представљени табелама.</p></li>
</ol>
<div class="section" id="id2">
<h2>Петља у петљи<a class="headerlink" href="#id2" title="Permalink to this headline">¶</a></h2>
<p>Када се у телу петље нађе друга петља, као у овом примеру:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">for</span> <span class="n">i</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="mi">3</span><span class="p">):</span>
    <span class="nb">print</span><span class="p">(</span><span class="s2">&quot;Tablica mnozenja sa&quot;</span><span class="p">,</span> <span class="n">i</span><span class="o">+</span><span class="mi">1</span><span class="p">)</span>
    <span class="c1">######### ovo je petlja u petlji</span>
    <span class="k">for</span> <span class="n">j</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="mi">5</span><span class="p">):</span>
        <span class="nb">print</span><span class="p">(</span><span class="n">i</span><span class="o">+</span><span class="mi">1</span><span class="p">,</span> <span class="s2">&quot;*&quot;</span><span class="p">,</span> <span class="n">j</span><span class="o">+</span><span class="mi">1</span><span class="p">,</span> <span class="s2">&quot;=&quot;</span><span class="p">,</span> <span class="p">(</span><span class="n">i</span><span class="o">+</span><span class="mi">1</span><span class="p">)</span> <span class="o">*</span> <span class="p">(</span><span class="n">j</span><span class="o">+</span><span class="mi">1</span><span class="p">))</span>
</pre></div>
</div>
<p>онда кажемо да је друга петља <em>угнежђена</em>, јер је обухваћена оном првом као у гнезду.
За прву кажемо још и да је <em>спољашња</em> петља, док је ова друга <em>унутрашња</em>. Погледајмо како овај прогам ради:</p>

<div data-childcomponent="primer4-1" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
for i in range(3):
    print("Tablica mnozenja sa", i+1)
    ######### ovo je petlja u petlji
    for j in range(5):
        print(i+1, "*", j+1, "=", (i+1) * (j+1))
</textarea>
</div>
<p>Угнежђене петље се природно јављају када радимо са табелама података, или са подацима који треба да буду организовани
у две димензије (хоризонтално и вертикално).</p>
<div class="section" id="id3">
<h3>Пример.<a class="headerlink" href="#id3" title="Permalink to this headline">¶</a></h3>
<p>Написати програм који приказује таблицу множења 10 x 10.</p>
<p><em>Решење</em>. Први покушај да се овај проблем реши би могао да изгледа овако:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="k">for</span> <span class="n">i</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="mi">10</span><span class="p">):</span>
    <span class="k">for</span> <span class="n">j</span> <span class="ow">in</span> <span class="nb">range</span><span class="p">(</span><span class="mi">10</span><span class="p">):</span>
        <span class="nb">print</span><span class="p">((</span><span class="n">i</span> <span class="o">+</span> <span class="mi">1</span><span class="p">)</span><span class="o">*</span><span class="p">(</span><span class="n">j</span> <span class="o">+</span> <span class="mi">1</span><span class="p">))</span>
</pre></div>
</div>
<p>али то није оно што смо желели. Желимо таблицу множења да испишемо као <em>таблицу</em>.
Проблем је у томе што наредба <code class="docutils literal notranslate"><span class="pre">print</span></code> пређе у нови ред чим нешто испише, а ми желимо да испишемо низ бројева
у једном реду пре него што пређемо у нови ред.</p>
<p>Наредби <code class="docutils literal notranslate"><span class="pre">print</span></code> можемо рећи да не прелази у нови ред тако што као опцију додамо <code class="docutils literal notranslate"><span class="pre">end</span> <span class="pre">=</span> <span class="pre">&quot;&quot;</span></code>.
Тада ће све што наредба испише бити смештено у исти ред. Када завршимо исписивање бројева који треба да буду
у истом реду прелазимо у нови ред наредбом <code class="docutils literal notranslate"><span class="pre">print()</span></code>. Ова наредба, заправо, ништа не испише, већ само пређе у нови ред.</p>

<div data-childcomponent="primer4-3" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-3 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
for i in range(10):
    for j in range(10):
        print((i + 1)*(j + 1), end="")
    print()
</textarea>
</div>
<p>Сада је мало боље, али и даље далеко од онога што желимо. Треба нам мали размак између бројева. То је лако додати:</p>

<div data-childcomponent="primer4-4" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-4 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
for i in range(10):
    for j in range(10):
        print((i + 1)*(j + 1), " ", end="")
    print()
</textarea>
</div>
<p>Добили смо таблицу множења <em>у принципу</em>. Једини је проблем у томе што су редови и даље неуредни.</p>
<p>Да бисмо у таблици поравнали редове није, дакле, довољно само додати мали размак, већ треба обезбедити да
<em>за сваки број буде употребљен исти број места</em> независно од тога да ли број има једну, две или три цифре.
За то ћемо користити уграђену функцију <code class="docutils literal notranslate"><span class="pre">rjust</span></code> (од енглеског <em>right justified</em> што значи „поравнато о десној ивици”)
која приказује податке на посебан начин. На пример,</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="s2">&quot;Python&quot;</span><span class="o">.</span><span class="n">rjust</span><span class="p">(</span><span class="mi">15</span><span class="p">)</span>
</pre></div>
</div>
<p>ће од стринга <code class="docutils literal notranslate"><span class="pre">&quot;Python&quot;</span></code> направити нови стринг <code class="docutils literal notranslate"><span class="pre">&quot;</span>&#160;&#160;&#160;&#160;&#160;&#160;&#160;&#160; <span class="pre">Python&quot;</span></code> који има дужину 15. Функција
<code class="docutils literal notranslate"><span class="pre">rjust</span></code> ће на пошетак стринга додати потребан број празнина тако да дужина новог стринга буде дати број. С друге стране,
ако желимо неки број да испишемо на дати број места поравнато по десној ивици морамо прво од њега да направимо стринг кога
ћемо потом поравнати по десној ивици функцијом <code class="docutils literal notranslate"><span class="pre">rjust</span></code>. Тако,</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="nb">str</span><span class="p">(</span><span class="mi">56</span><span class="p">)</span><span class="o">.</span><span class="n">rjust</span><span class="p">(</span><span class="mi">4</span><span class="p">)</span>
</pre></div>
</div>
<p>ће <em>форматирати</em> број 56 као стринг дужине 4 у коме је садржај поравнат по десној ивици: <code class="docutils literal notranslate"><span class="pre">&quot;</span>&#160; <span class="pre">56&quot;</span></code>.</p>
<p>Коначно, ево програма који исписује таблицу множења <em>као табелу</em>:</p>

<div data-childcomponent="primer4-5" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-5 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="for i in range(10):
    for j in range(10):
        s = str((i + 1)*(j + 1))
        print(s.rjust(4), end=&quot;&quot;)
    print()
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P04</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Tablica_mnoz.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
<div class="section" id="id4">
<h3>Пример.<a class="headerlink" href="#id4" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">trougao_brojeva(n)</span></code> која исписује следећи троугао бројева:</p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span>1
1  2
1  2  3
1  2  3  4
1  2  3  4  5
</pre></div>
</div>
<p>итд. Ово је пример позива функције за <code class="docutils literal notranslate"><span class="pre">n</span> <span class="pre">=</span> <span class="pre">5</span></code>.</p>

<div data-childcomponent="primer4-6" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-6 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="def trougao_brojeva(n):
    for i in range(n + 1):
        for j in range(i):
            s = str(j + 1)
            print(s.rjust(4), end=&quot;&quot;)
        print()

# Провера
trougao_brojeva(6)

"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P04</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Trougao.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
<div class="section" id="id5">
<h3>Пример.<a class="headerlink" href="#id5" title="Permalink to this headline">¶</a></h3>
<p>Написати Пајтон функцију <code class="docutils literal notranslate"><span class="pre">trougao_brojeva_2(n)</span></code> која исписује следећи троугао бројева:</p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span> 1
 2   3
 4   5   6
 7   8   9  10
11  12  13  14  15
</pre></div>
</div>
<p>итд. Ово је пример позива функције за <code class="docutils literal notranslate"><span class="pre">n</span> <span class="pre">=</span> <span class="pre">5</span></code>. Видимо да троугао има n редова, да су дужине редова
1, 2, 3, …, n, и да се бројеви редом уписују у табелу.</p>
<p><em>Решење.</em></p>
<p>Идеја решења је да уведемо бројач који ће садржати следећи број који треба уписати на одговарајуће место,
док ће две петље само регулисати облик троугла, овако:</p>

<div data-childcomponent="primer4-7" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-7 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="def trougao_brojeva_2(n):
    broj = 1
    for i in range(n + 1):
        for j in range(i):
            s = str(broj)
            print(s.rjust(4), end=&quot;&quot;)
            broj += 1
        print()

# Провера
trougao_brojeva_2(6)
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P04</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Trougao_2.py</span></code> и изврши га.</p>

    </p></div></div>
</div>
<div class="section" id="id6">
<h3>Пример.<a class="headerlink" href="#id6" title="Permalink to this headline">¶</a></h3>
<p>Наћи све бројеве са највише четири цифре који су једнаки трећем степену збира својих цифара.</p>
<p><em>Решење.</em></p>
<p>Број са највише четири цифре се може записати у облику <span class="math notranslate nohighlight">\(\langle abcd\rangle\)</span>,
при чему је дозвољено да свака од цифара буде нула (траже се бројеви са највише четири цифре; према томе,
за троцифрени број ће бити <span class="math notranslate nohighlight">\(a = 0\)</span>).</p>
<p>Дакле, тражимо бројеве <span class="math notranslate nohighlight">\(\langle abcd\rangle\)</span> такве да је <span class="math notranslate nohighlight">\(\langle abcd\rangle = (a + b + c + d)^3\)</span>.
Подсетимо се још да је:</p>
<div class="math notranslate nohighlight">
\[\langle abcd\rangle = 1000 a + 100 b + 10 c + d.\]</div>
<p>Сада се програм лако пише:</p>

<div data-childcomponent="primer4-8" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-8 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="for a in range(1, 10):
    for b in range(10):
        for c in range(10):
            for d in range(10):
                if 1000*a + 100*b + 10*c + d == (a + b + c + d)**3:
                    print(1000*a + 100*b + 10*c + d)
"   >


</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P04</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Treci_stepen_cif.py</span></code> и изврши га.</p>
<p>Ево и кратке видео демонстрације:</p>

    <div  id="ykG3h8darso" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/ykG3h8darso/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-ykG3h8darso" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-ykG3h8darso">
        </div>
    </div>
    
    
    
    </p></div></div>
</div>
</div>
<div class="section" id="id7">
<h2>Табеле са подацима<a class="headerlink" href="#id7" title="Permalink to this headline">¶</a></h2>
<p>Веома често су подаци у вези са неким проблемом представљени табелом. Рецимо, ова табела садржи податке о једној групи деце (при чему је, наравно, старост изражена у годинама, маса у килограмима, а висина у центиметрима):</p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span>    Ime Pol Starost Masa Visina
--------------------------------
    Ana  ž    13     46    160
  Bojan  m    14     52    165
  Vlada  m    13     47    157
Gordana  ž    15     54    165
  Dejan  m    15     56    163
  Đorđe  m    13     45    159
  Elena  ž    14     49    161
Žaklina  ž    15     52    164
  Zoran  m    15     57    167
  Ivana  ž    13     45    158
  Jasna  ž    14     51    162
</pre></div>
</div>
<p>Да бисмо могли машински да обрађујемо и анализирамо податке, прво их морамо представити на погодан начин.</p>
<p>Један једноставан начин да се то уради је да сваки ред табеле представимо једном низом, и да потом све те низове
запакујемо у један велики низ, рецимо овако:</p>

<div data-childcomponent="primer4-podaci-o-razredu" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-podaci-o-razredu data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
razred = [["Ana",     "ž", 13, 46, 160],
          ["Bojan",   "m", 14, 52, 165],
          ["Vlada",   "m", 13, 47, 157],
          ["Gordana", "ž", 15, 54, 165],
          ["Dejan",   "m", 15, 56, 163],
          ["Đorđe",   "m", 13, 45, 159],
          ["Elena",   "ž", 14, 49, 161],
          ["Žaklina", "ž", 15, 52, 164],
          ["Zoran",   "m", 15, 57, 167],
          ["Ivana",   "ž", 13, 45, 158],
          ["Jasna",   "ž", 14, 51, 162]]
</textarea>
</div>
<p>(Ако извршиш ову наредбу неће се десити ништа: систем ће само наведене податке спремити у променљиву <code class="docutils literal notranslate"><span class="pre">razred</span></code>.)</p>
<p>Као и у већини модерних програмских језика, индекси елемената низа у Пајтону почињу од нуле. Зато је:</p>

<div data-childcomponent="primer4-9" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-9 data-lang="python" 
     data-include="primer4-podaci-o-razredu" data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
print("Prvi red tabele:", razred[0])
print("Drugi red tabele:", razred[1])
print("Vlada je visok", razred[2][4], "cm")
</textarea>
</div>
<p>Последња наредба заслужује додатни коментар.
Пошто се подаци о Влади налазе у трећем реду табеле (чији индекс је 2, јер индексирање почиње од нуле!),
а висина се налази у том низу на петом месту (чији индекс је 4) израз <code class="docutils literal notranslate"><span class="pre">razred[2][4]</span></code> каже:
<em>из табеле ``razred`` одабери прво ред са индексом 2, па из тог реда податак са индексом 4
(при томе водећи рачуна да индексирање почиње од нуле!)</em></p>
<p>Ако желимо да израчунамо просечну висину ученика у разреду учинићемо то овако:</p>

<div data-childcomponent="primer4-10" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-10 data-lang="python" 
     data-include="primer4-podaci-o-razredu" data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
br_ucenika = len(razred)
zbir = 0
for ucenik in razred:
    zbir += ucenik[4]
print("Prosecna visina je", zbir/br_ucenika)
</textarea>
</div>
<p>Ако желимо да пребројимо колико у разреду има дечака, а колико девојчица, можемо то учинити овако
(имајући у виду да се информација о полу налази на другом месту у реду, чији индекс је 1):</p>

<div data-childcomponent="primer4-11" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-11 data-lang="python" 
     data-include="primer4-podaci-o-razredu" data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
br_decaka = 0
br_devojcica = 0
for ucenik in razred:
    if ucenik[1] == "m":
        br_decaka += 1
    else:
        br_devojcica += 1
print("Decaka ima", br_decaka)
print("Devojcica ima", br_devojcica)
</textarea>
</div>
<p>Ево како можемо одредити највишу девојчицу:</p>

<div data-childcomponent="primer4-12" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-12 data-lang="python" 
     data-include="primer4-podaci-o-razredu" data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
       >
# ne znamo koja devojcica je najvisa, pa stavljamo prazan string u ime
ime = ""

# ne znamo nista o njenoj visini, pa stavljamo vrednost koja je manja od svake razumne vrednosti
visina = 0

for ucenik in razred:
    if ucenik[1] == "ž" and ucenik[4] > visina:
        ime = ucenik[0]
print("Najvisa devojcica je", ime)
</textarea>
</div>

    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Изврши исти програм и у Пајтон окружењу!</p>
<p>Покрени IDLE, из фолдера <code class="docutils literal notranslate"><span class="pre">P04</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Grupa_dece.py</span></code> и изврши га.</p>
<p>Ево и кратке видео демонстрације:</p>

    <div  id="90frnGYJ7bg" class="ytvideo"  style="text-align: center; margin: 15px; cursor:pointer;">
        <div style="background-image: url('https://img.youtube.com/vi/90frnGYJ7bg/mqdefault.jpg'); background-repeat: no-repeat;background-position: center; height: 210px; width: 310px; margin: 0 auto; border: 1px solid #20c997">
            <img src="/_static/images/play_button.svg" style="margin-top: 78px;" /> 
        </div>
      

        </div>
    <div id="modal-90frnGYJ7bg" class="ytvideoclose" style="display: none; position: fixed; top: 0; left: 0; width: 100vw; height: 100vh; background-color: rgba(128, 182, 128, 0.3); z-index: 10000;">
        <div style="background-color: white; position: fixed; top: 4vh; left: 5vw; width:  90vw; height:  90vh; z-index: 100;" id="YTmodal-90frnGYJ7bg">
        </div>
    </div>
    
    
    
    </p></div></div>
</div>
<div class="section" id="id8">
<h2>Задаци<a class="headerlink" href="#id8" title="Permalink to this headline">¶</a></h2>
<div class="section" id="id9">
<h3>Задатак 1.<a class="headerlink" href="#id9" title="Permalink to this headline">¶</a></h3>
<p>(Задатак реши у IDLE окружењу)
Напиши Пајтон програм који исписује таблицу <em>сабирања</em>, али само за цифре 0, 1, …, 9, овако:</p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span> + | 0   1   2   3   4   5   6   7   8   9
---+----------------------------------------
 0 | 0   1   2   3   4   5   6   7   8   9
 1 | 1   2   3   4   5   6   7   8   9  10
 2 | 2   3   4   5   6   7   8   9  10  11
 3 | 3   4   5   6   7   8   9  10  11  12
 4 | 4   5   6   7   8   9  10  11  12  13
 5 | 5   6   7   8   9  10  11  12  13  14
 6 | 6   7   8   9  10  11  12  13  14  15
 7 | 7   8   9  10  11  12  13  14  15  16
 8 | 8   9  10  11  12  13  14  15  16  17
 9 | 9  10  11  12  13  14  15  16  17  18
</pre></div>
</div>

    <div data-component="reveal" id="pomoc3"   data-showtitle="Помоћ" data-hidetitle="Сакриј помоћ">
    
    <div class="course-box course-box-info">
        <div class="course-content">
            <p>
<p>Не заборави да наредбу принт можеш позвати и више од једном и да програм не мораш да решиш помоћу једне (двоструке) петље.</p>

    </p></div></div>

    </div>
    </div>
<div class="section" id="id10">
<h3>Задатак 2.<a class="headerlink" href="#id10" title="Permalink to this headline">¶</a></h3>
<p>(Задатак реши у IDLE окружењу)
Написати Пајтон програм који од корисника учитава природан број <em>n</em> и онда исписује првих <em>n</em> редова следећег троугла бројева:</p>
<p><em>(а)</em></p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span>1
1 1
1 1 1
1 1 1 1
</pre></div>
</div>
<p><em>(б)</em></p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span>1
2 1
3 2 1
4 3 2 1
5 4 3 2 1
</pre></div>
</div>
<p><em>(в)</em></p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span>1
2  3  2
3  4  5  4  3
4  5  6  7  6  5  4
5  6  7  8  9  8  7  6  5
</pre></div>
</div>
</div>
<div class="section" id="id11">
<h3>Задатак 3.<a class="headerlink" href="#id11" title="Permalink to this headline">¶</a></h3>
<p>(Задатак реши у IDLE окружењу)
Написати Пајтон програм који од корисника учитава природан број <em>n</em> и онда бројеве 1, 2, 3, …, <span class="math notranslate nohighlight">\(n^2\)</span>
уписује у квадрат <span class="math notranslate nohighlight">\(n \times n\)</span> „змијасто”. На пример, за <span class="math notranslate nohighlight">\(n = 5\)</span> програм исписује:</p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span> 1   2   3   4   5
10   9   8   7   6
11  12  13  14  15
20  19  18  17  16
21  22  23  24  25
</pre></div>
</div>
</div>
<div class="section" id="id12">
<h3>Задатак 4.<a class="headerlink" href="#id12" title="Permalink to this headline">¶</a></h3>
<p>(Задатак реши у IDLE окружењу)
Наћи све бројеве са највише четири цифре који су једнаки</p>
<p><em>(а)</em> квадрату збира својих цифара;</p>
<p><em>(б)</em> четвртом степену збира својих цифара.</p>
</div>
<div class="section" id="id13">
<h3>Задатак 5.<a class="headerlink" href="#id13" title="Permalink to this headline">¶</a></h3>
<p>(Задатак реши у IDLE окружењу)
Наћи сва могућа дешифровања следећег сабирања:</p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span>АВ + ... + АВ = ПАС
</pre></div>
</div>
<p>где са десне стране имамо најмање два сабирка. Наравно, истим словима одговарају исте цифре,
а различитим словима различите. При томе цифре А и Л не смеју да буду 0.</p>
</div>
<div class="section" id="id14">
<h3>Задатак 6.<a class="headerlink" href="#id14" title="Permalink to this headline">¶</a></h3>
<p>(Задатак реши у IDLE окружењу)
Наћи сва могућа дешифровања следећег сабирања:</p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span>ЛАВ + ... + ЛАВ = МРАВ
</pre></div>
</div>
<p>где са десне стране имамо најмање два сабирка. Наравно, истим словима одговарају исте цифре,
а различитим словима различите. При томе цифре Л и М не смеју да буду 0.</p>
</div>
<div class="section" id="id15">
<h3>Задатак 7.<a class="headerlink" href="#id15" title="Permalink to this headline">¶</a></h3>
<p>Дати су подаци о ученицима једног разреда. Напиши Пајтон програм који одређује
ко је најлакши у разреду, и ко је најтежи дечак у разреду.</p>

<div data-childcomponent="primer4-13" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-13 data-lang="python" 
      data-timelimit=25000    
    data-audio=''   
          data-runortest="ime1 ime2"     data-hidehistory=true
       >
# -*- acsection: general-init -*-
razred = [["Ana",     "ž", 13, 46, 160],
          ["Bojan",   "m", 14, 52, 165],
          ["Vlada",   "m", 13, 47, 157],
          ["Gordana", "ž", 15, 54, 165],
          ["Dejan",   "m", 15, 56, 163],
          ["Đorđe",   "m", 13, 47, 159],
          ["Elena",   "ž", 14, 49, 161],
          ["Žaklina", "ž", 15, 52, 164],
          ["Zoran",   "m", 15, 57, 167],
          ["Ivana",   "ž", 13, 45, 158],
          ["Jasna",   "ž", 14, 51, 162]]
# -*- acsection: main -*-

# овде напиши програм који име најлакшег ученика (независно од пола!) у разреду
# смешта у променљиву ime1, а име најтежег дечака у разреду у променљиву ime2

ime1 = "" # поправи овај ред!
ime2 = "" # поправи овај ред!

# -*- acsection: after-main -*-

print("Najlaksi ucenik u razredu je", ime1)
print("Najtezi decak u razredu je", ime2)
====
from unittest.gui import TestCaseGui
class myTests(TestCaseGui):
    def testOne(self):
        rez1 = "Ivana"
        rez2 = "Zoran"
        run_test = acMainSection(ime1=ime1,ime2=ime2)
        self.assertEqual(run_test["ime1"], rez1, "Вредност променљиве 'ime1' треба да буде %s" % rez1)
        self.assertEqual(run_test["ime2"], rez2, "Вредност променљиве 'ime2' треба да буде %s" % rez2)
myTests().main()
</textarea>
</div>
<p>Задатак можеш решити и у Пајтон окружењу.
Покрени IDLE па из фолдера <code class="docutils literal notranslate"><span class="pre">P04</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Grupa_dece_2.py</span></code>.</p>
</div>
<div class="section" id="id16">
<h3>Задатак 8.<a class="headerlink" href="#id16" title="Permalink to this headline">¶</a></h3>
<p>Ево трошкова живота једне породице током једне године, по месецима (сви износи су представљени у динарима):</p>
<div class="highlight-text notranslate"><div class="highlight"><pre><span></span> Stavka              Jan    Feb    Mar    Apr    Maj    Jun    Jul    Avg    Sep    Okt    Nov    Dec
-----------------   -----  -----  -----  -----  -----  -----  -----  -----  -----  -----  -----  -----
 Stanarina          8.251  8.436  8.524  8.388  8.241  8.196  8.004  7.996  7.991  8.015  8.353  8.456
 Struja             4.321  4.530  4.115  3.990  3.985  3.726  3.351  3.289  3.295  3.485  3.826  3.834
 Telefon (fiksni)   1.425  1.538  1.623  1.489  1.521  1.485  1.491  1.399  1.467  1.531  1.410  1.385
 Telefon (mobilni)  2.181  2.235  2.073  1.951  1.989  1.945  3.017  2.638  2.171  1.831  1.926  1.833
 TV i internet      2.399  2.399  2.399  2.399  2.399  2.399  2.399  2.399  2.399  2.399  2.399  2.399
 Prevoz             1.830  1.830  1.830  1.830  1.950  1.950  1.450  1.450  1.950  1.950  2.050  2.050
 Hrana             23.250 23.780 24.019 24.117 24.389 24.571 24.736 24.951 25.111 25.389 25.531 25.923
 Ostalo             4.500  3.700  5.100  3.500  2.750  4.250  7.320  8.250  3.270  4.290  3.200  8.390
</pre></div>
</div>
<p>Испод су исти подаци представљени табелом. Напиши Пајтон програм који рачуна укупне трошкове ове породице по месецима
(колико је породица укупно потрошила у јануару, колико у фебруару итд).</p>

<div data-childcomponent="primer4-14" class="course-box course-box-problem course-content">

<textarea data-component="activecode" id=primer4-14 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
               data-hidehistory=true
    data-includesrc="troskovi = [
  [&quot;Stanarina&quot;, 8251, 8436, 8524, 8388, 8241, 8196, 8004, 7996, 7991, 8015, 8353, 8456],
  [&quot;Struja&quot;, 4321, 4530, 4115, 3990, 3985, 3726, 3351, 3289, 3295, 3485, 3826, 3834],
  [&quot;Telefon (fiksni)&quot;, 1425, 1538, 1623, 1489, 1521, 1485, 1491, 1399, 1467, 1531, 1410, 1385],
  [&quot;Telefon (mobilni)&quot;, 2181, 2235, 2073, 1951, 1989, 1945, 3017, 2638, 2171, 1831, 1926, 1833],
  [&quot;TV i internet&quot;, 2399, 2399, 2399, 2399, 2399, 2399, 2399, 2399, 2399, 2399, 2399, 2399 ],
  [&quot;Prevoz&quot;, 1830, 1830, 1830, 1830, 1950, 1950, 1450, 1450, 1950, 1950, 2050, 2050],
  [&quot;Hrana&quot;, 23250, 23780, 24019, 24117, 24389, 24571, 24736, 24951, 25111, 25389, 25531, 25923],
  [&quot;Ostalo&quot;, 4500, 3700, 5100, 3500, 2750, 4250, 7320, 8250, 3270, 4290, 3200, 8390]
]

# овде напиши програм
???
"   >


</textarea>
</div>
<p>Задатак можеш решити и у Пајтон окружењу.
Покрени IDLE па из фолдера <code class="docutils literal notranslate"><span class="pre">P04</span></code> учитај програм <code class="docutils literal notranslate"><span class="pre">Troskovi.py</span></code>.</p>
</div>
</div>
<div class="section" id="id17">
<h2>Задаци из Алгоритамске збирке задатака<a class="headerlink" href="#id17" title="Permalink to this headline">¶</a></h2>
<p>Решавањем следећих задатака можете да увежбате додатне вештине у раду са угнежђеним циклусима.</p>
<div class="line-block">
<div class="line"><br /></div>
</div>
<p><strong>Уводни задаци</strong></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/serije_123">https://petlja.org/biblioteka/r/problemi/Zbirka/serije_123</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/serije_neparni_parni">https://petlja.org/biblioteka/r/problemi/Zbirka/serije_neparni_parni</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/tablica_mnozenja">https://petlja.org/biblioteka/r/problemi/Zbirka/tablica_mnozenja</a></p>
<div class="line-block">
<div class="line"><br /></div>
</div>
<p><strong>Комбинаторни објекти</strong></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/ciklicne_permutacije">https://petlja.org/biblioteka/r/problemi/Zbirka/ciklicne_permutacije</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/kockice_za_jamb">https://petlja.org/biblioteka/r/problemi/Zbirka/kockice_za_jamb</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/trouglovi_datog_obima">https://petlja.org/biblioteka/r/problemi/Zbirka/trouglovi_datog_obima</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/kombinacije_poena">https://petlja.org/biblioteka/r/problemi/Zbirka/kombinacije_poena</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/varijacije_trojki">https://petlja.org/biblioteka/r/problemi/Zbirka/varijacije_trojki</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/brojevi_u_datoj_osnovi">https://petlja.org/biblioteka/r/problemi/Zbirka/brojevi_u_datoj_osnovi</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/mali_loto">https://petlja.org/biblioteka/r/problemi/Zbirka/mali_loto</a></p>
<div class="line-block">
<div class="line"><br /></div>
</div>
<p><strong>Сегменти</strong></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/svi_desni_segmenti_1_n">https://petlja.org/biblioteka/r/problemi/Zbirka/svi_desni_segmenti_1_n</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/svi_levi_segmenti_1_n">https://petlja.org/biblioteka/r/problemi/Zbirka/svi_levi_segmenti_1_n</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/sve_podreci">https://petlja.org/biblioteka/r/problemi/Zbirka/sve_podreci</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/sve_podreci_po_opadajucoj_duzini">https://petlja.org/biblioteka/r/problemi/Zbirka/sve_podreci_po_opadajucoj_duzini</a></p>
<div class="line-block">
<div class="line"><br /></div>
</div>
<p><strong>ASCII цртежи</strong></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/kvadrat_od_zvezdica">https://petlja.org/biblioteka/r/problemi/Zbirka/kvadrat_od_zvezdica</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/trougao_od_zvezdica">https://petlja.org/biblioteka/r/problemi/Zbirka/trougao_od_zvezdica</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/romb_od_zvezdica">https://petlja.org/biblioteka/r/problemi/Zbirka/romb_od_zvezdica</a></p>
<p><a class="reference external" href="https://petlja.org/biblioteka/r/problemi/Zbirka/trougao_od_reci">https://petlja.org/biblioteka/r/problemi/Zbirka/trougao_od_reci</a></p>
</div>
</div>


    </div>
    
    <div>
<div class="lecture-prev-next-toggle">
  <a  href="index.html" id="prevLectureLink" class="lecture-toggler-arrows"><i class="fas fa-arrow-left fa"></i> <span> Претходно лекција </span> </a>
  
  <a id="nextLectureLink" class="lecture-toggler-arrows"><span> Следећа лекција</span>  <i class="fas fa-arrow-right fa"></i> </a>
</div>

<script type="text/javascript">
  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
</script>
</div>

    
  </div>
</div>
<div class="petljadoc-resources">
  <span class="pull-left">(Created using  <a href="https://pypi.org/project/Sphinx/">Swinx</a>, <a href="http://runestoneinteractive.org/">RunestoneComponents</a> and <a href="https://github.com/Petlja/PetljaDoc">PetljaDoc</a>) <br> © 2022 Petlja</span>
</div>

<div class="lectureToolbox">

  <label class="font-size-toggle" id="fontSizeToggle">

      <i class="fas fa-font"></i><label>Слова</label>
  </label>
  <div class="font-change d-none" id="fontChangeModal">
      <span class="dec-letter-icon" id="decLetterIcon">A-</span>
      <label id="currentFontSize">16</label>
      <span class="inc-letter-icon" id="incLetterIcon">A+</span>
  </div>
</div>

</div>

  </body>
</html>